<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.17.1: http://docutils.sourceforge.net/" />

    <title>20.12. Write Code Questions &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.cd688d1623c9932e.css?v=630908D8" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.bf73008a1eba9f4e.css?v=7DED055E" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/runtime.439d581f758f2b99.bundle.js?v=2E54442"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.51d9dde4e5162190.bundle.js?v=9CAA2D33"></script>
    <script src="../_static/runestone.42b576791a4ed54b.bundle.js?v=B7293199"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="21. Databases" href="../database/toctree.html" />
    <link rel="prev" title="20.11. Mixed-Up Code Questions" href="mixedupcode.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    
    eBookConfig.useRunestoneServices = {{use_services}};
    eBookConfig.host = '';
    eBookConfig.app = eBookConfig.host + '/runestone';
    eBookConfig.course = '{{ course_name }}';
    eBookConfig.basecourse = '{{ base_course }}';
    eBookConfig.isLoggedIn = {{ is_logged_in}};
    eBookConfig.email = '{{ user_email }}';
    eBookConfig.isInstructor = {{ is_instructor }};
    eBookConfig.username = '{{ user_id}}';
    eBookConfig.readings = {{ readings|safe}};
    eBookConfig.activities = {{ activity_info|safe }}
    eBookConfig.downloadsEnabled = {{downloads_enabled}};
    eBookConfig.allow_pairs = {{allow_pairs}}
    eBookConfig.enableCompareMe = {{enable_compare_me}};
    // _`new_server_prefix`: Defined by the server to be the path to the new server.
    eBookConfig.new_server_prefix = '{{ new_server_prefix }}';
    
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->

    
        {% if serve_ad is defined and serve_ad and settings.adsenseid %}
         <script data-ad-client="{{settings.adsenseid}}" async src="https://pagead2.googlesyndication.com/pagead/js/adsbygoogle.js"></script>
        {% endif %}
    



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
            
            {{ course_name }}
            
          
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="browsing_warning"></span></li>

        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li id="inst_peer_link"><a href='/runestone/peer/instructor.html'>Peer Instruction (Instructor)</a></li>            
            <li><a href='/runestone/peer/student.html'>Peer Instruction (Student)</a></li>            
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/ns/books/published/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=writecode'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
            <script></script>
<li class="dropdown globaltoc-container">
    <a
        href="../index.html"
        class="dropdown-toggle"
        data-toggle="dropdown"
        aria-label="index"
        >This Chapter<b class="caret"></b
    ></a>
    <ul class="dropdown-menu globaltoc">
         
        {% for row in subchapter_list %}
        <li class="toctree-l1">
            <a class="reference internal" href="{{ row['subchap_uri']}}"
                >{{ row['title']}}</a
            >
        </li>
        {% endfor %} 
        
    </ul>
</li>
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">20.11. </span>Mixed-Up Code Questions" data-toggle="tooltip">
    <a href="mixedupcode.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">21. </span>Databases' data-toggle="tooltip" >              
    <a href="../database/toctree.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->

    
        {% if settings.num_banners > 0 and settings.show_rs_banner %}
         <div style="width: 100%;">
            <a href="/runestone/default/donate?banner={{banner_num}}">
            <img class="runestone-banner"
                 src="/runestone/static/images/RunestoneBanner{{banner_num}}.png"
                 alt="Please Support Runestone"
                 style="display: block; width: 728px; margin-left: auto; margin-right: auto; margin-top: 10px;">
            </a>
         </div>
        {% endif %}
    


  
  <section id="write-code-questions">
<h1><span class="section-number">20.12. </span>Write Code Questions<a class="headerlink" href="#write-code-questions" title="Permalink to this headline">¶</a></h1>
<ol class="arabic">
<li>
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=inheritance_writecode1q_v2 data-question_label="20.12.1">
<div id=inheritance_writecode1q_v2_question class="ac_question col-md-12">
<p>Create a subclass of the <code class="docutils literal notranslate"><span class="pre">Parent</span></code> class named <code class="docutils literal notranslate"><span class="pre">Child</span></code> that overrides the inherited <code class="docutils literal notranslate"><span class="pre">eye_color()</span></code> method to return <code class="docutils literal notranslate"><span class="pre">&quot;I</span> <span class="pre">have</span> <span class="pre">green</span> <span class="pre">eyes.&quot;</span></code></p>

</div>
<textarea data-lang="python" id="inheritance_writecode1q_v2_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Parent:

    def eye_color(self):
        return "I have blue eyes."


from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        lil_child = Child()
        parent = Parent()
        self.assertEqual(lil_child.eye_color(), "I have green eyes.", "Testing Child eye_color")
        self.assertEqual(parent.eye_color(), "I have blue eyes.", "Testing Parent eye_color")
        self.assertIsInstance(lil_child, type(parent), "Testing Child inherits from Parent")

myTests().main()
</textarea>
</div>
</div>
</li>
<li>
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=inheritance_greeter_grumpy_ac data-question_label="20.12.2">
<div id=inheritance_greeter_grumpy_ac_question class="ac_question col-md-12">
<p>Create a subclass of the <code class="docutils literal notranslate"><span class="pre">Greeter</span></code> class named <code class="docutils literal notranslate"><span class="pre">GrumpyGreeter</span></code> that overrides the inherited <code class="docutils literal notranslate"><span class="pre">greet</span></code> method to return <code class="docutils literal notranslate"><span class="pre">&quot;Go</span> <span class="pre">Away!&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="inheritance_greeter_grumpy_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Greeter:

    def greet(self):
        return "Welcome!"



from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        gg = GrumpyGreeter()
        g = Greeter()
        self.assertEqual(g.greet(), "Welcome!", "Testing Greeter greet")
        self.assertEqual(gg.greet(), "Go Away!", "Testing GrumpyGreeter greet")
        self.assertIsInstance(gg, type(g), "Testing GrumpyGreeter inherits from Greeter")

myTests().main()
</textarea>
</div>
</div>
</li>
<li>
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=inheritance_animal_cow_ac data-question_label="20.12.3">
<div id=inheritance_animal_cow_ac_question class="ac_question col-md-12">
<p>Create a subclass of the <code class="docutils literal notranslate"><span class="pre">Animal</span></code> class named <code class="docutils literal notranslate"><span class="pre">Cow</span></code> that overrides the inherited <code class="docutils literal notranslate"><span class="pre">noise</span></code> method to return <code class="docutils literal notranslate"><span class="pre">&quot;Moo&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="inheritance_animal_cow_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Animal:

    def noise(self):
        return "Unknown"



from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        cow = Cow()
        p = Animal()
        self.assertEqual(p.noise(), "Unknown", "Testing Animal noise")
        self.assertEqual(c.noise(), "Moo", "Testing Cow noise")
        self.assertIsInstance(cow, type(p), "Testing Cow inherits from Animal")

myTests().main()
</textarea>
</div>
</div>
</li>
<li>
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=inheritance_actor_over_actor_ac data-question_label="20.12.4">
<div id=inheritance_actor_over_actor_ac_question class="ac_question col-md-12">
<p>Create a subclass of the <code class="docutils literal notranslate"><span class="pre">Actor</span></code> class named <code class="docutils literal notranslate"><span class="pre">OverActor</span></code> that overrides the inherited <code class="docutils literal notranslate"><span class="pre">act</span></code> method to return <code class="docutils literal notranslate"><span class="pre">&quot;I</span> <span class="pre">am</span> <span class="pre">super</span> <span class="pre">happy!&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="inheritance_actor_over_actor_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Actor:

    def act(self):
        return "I am happy!"



from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        a = Actor()
        oa = OverActor()
        self.assertEqual(a.act(), "I am happy!", "Testing Actor act")
        self.assertEqual(oa.act(), "I am super happy!", "Testing OverActor act")
        self.assertIsInstance(oa, type(a), "Testing OverActor inherits from Actor")


myTests().main()
</textarea>
</div>
</div>
</li>
<li>
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=inheritance_student_grad_student_ac data-question_label="20.12.5">
<div id=inheritance_student_grad_student_ac_question class="ac_question col-md-12">
<p>Create a subclass of the <code class="docutils literal notranslate"><span class="pre">Student</span></code> class named <code class="docutils literal notranslate"><span class="pre">GradStudent</span></code> that calls the <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method in <code class="docutils literal notranslate"><span class="pre">Student</span></code> to initialize the first and last name and overrides the inherited <code class="docutils literal notranslate"><span class="pre">fav_food</span></code> method to return <code class="docutils literal notranslate"><span class="pre">&quot;Sushi&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="inheritance_student_grad_student_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Student:

    def __init__(self, first, last):
        self.first = first
        self.last = last

    def fav_food(self):
        return "Pizza"


from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        s = Student("Carl", "X")
        gs = GradStudent("Anjali", "Kumar")
        self.assertEqual(s.fav_food(), "Pizza", "Testing Student fav_food")
        self.assertEqual(s.first, "Carl", "Testing Student first name")
        self.assertEqual(s.last, "X", "Testing Student last name")
        self.assertEqual(gs.first, "Anjali", "Testing GradStudent first name")
        self.assertEqual(gs.last, "Kumar", "Testing GradStudent last name")
        self.assertEqual(gs.fav_food(), "Sushi", "Testing GradStudent fav_food")
        self.assertIsInstance(gs, type(s), "Testing GradStudent inherits from Student")

myTests().main()
</textarea>
</div>
</div>
</li>
<li>
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ut_is_ascending_ac data-question_label="20.12.6">
<div id=ut_is_ascending_ac_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">is_ascending(nums)</span></code> that returns <code class="docutils literal notranslate"><span class="pre">True</span></code> if the numbers in the list <code class="docutils literal notranslate"><span class="pre">nums</span></code> are sorted in ascending order and <code class="docutils literal notranslate"><span class="pre">False</span></code> otherwise. If the list <code class="docutils literal notranslate"><span class="pre">nums</span></code> has less than two numbers in it return <code class="docutils literal notranslate"><span class="pre">True</span></code>.  For example, <code class="docutils literal notranslate"><span class="pre">is_ascending([1,2,3])</span></code> should return <code class="docutils literal notranslate"><span class="pre">True</span></code>, <code class="docutils literal notranslate"><span class="pre">is_ascending([1])</span></code> should also return <code class="docutils literal notranslate"><span class="pre">True</span></code>, and <code class="docutils literal notranslate"><span class="pre">is_ascending([3,2,1])</span></code> should return <code class="docutils literal notranslate"><span class="pre">False</span></code>.</p>

</div>
<textarea data-lang="python" id="ut_is_ascending_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def is_ascending(nums):


from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(is_ascending([1,2,3]),True,"is_ascending([1,2,3])")
        self.assertEqual(is_ascending([3,2,1]),False,"is_ascending([3,2,1])")
        self.assertEqual(is_ascending([3,1,2]),False,"is_ascending([3,1,2])")
        self.assertEqual(is_ascending([2,3,1]),False,"is_ascending([2,3,1])")
        self.assertEqual(is_ascending([1,2,4,3]),False,"is_ascending([1,2,4,3])")
        self.assertEqual(is_ascending([1,2,3,0]),False,"is_ascending([1,2,3,0])")
        self.assertEqual(is_ascending([4,1,2,3]),False,"is_ascending(4,1,2,3])")
        self.assertEqual(is_ascending([1,4,2,3]),False,"is_ascending(1,4,2,3])")
        self.assertEqual(is_ascending([1,2,3,4]),True,"is_ascending([1,2,3,4])")
        self.assertEqual(is_ascending([1]),True,"is_ascending([1])")
        self.assertEqual(is_ascending([]),True,"is_ascending([])")


myTests().main()
</textarea>
</div>
</div>
</li>
<li>
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ut_sum_lists_ac data-question_label="20.12.7">
<div id=ut_sum_lists_ac_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">sum_lists(l1,l2)</span></code> that take two lists of numbers, <code class="docutils literal notranslate"><span class="pre">l1</span></code> and <code class="docutils literal notranslate"><span class="pre">l2</span></code> and returns a list with the sum of the corresponding items in <code class="docutils literal notranslate"><span class="pre">l1</span></code> and <code class="docutils literal notranslate"><span class="pre">l2</span></code>.  For example, <code class="docutils literal notranslate"><span class="pre">sum_lists([1,2],[3,4])</span></code> would return [4,6]. If the two lists are of different length then returned list should be the same length as the shortest list. For example, <code class="docutils literal notranslate"><span class="pre">sum_lists([1],[4,3])</span></code> should return <code class="docutils literal notranslate"><span class="pre">[5]</span></code>.</p>

</div>
<textarea data-lang="python" id="ut_sum_lists_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def sum_lists(l1,l2):


from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sum_lists([1,2],[3,4]),[4,6],"sum_lists([1,2],[3,4])")
        self.assertEqual(sum_lists([1],[2]),[3],"sum_lists([1],[2])")
        self.assertEqual(sum_lists([-1,-2],[3,4]),[2,2],"sum_lists([-1,-2],[3,4])")
        self.assertEqual(sum_lists([2,1],[4,3]),[6,4],"sum_lists([2,1],[4,3])")
        self.assertEqual(sum_lists([1],[4,3]),[5],"sum_lists([1],[4,3])")
        self.assertEqual(sum_lists([1,2],[4]),[5],"sum_lists([1,2],[4])")
        self.assertEqual(sum_lists([],[4,3]),[],"sum_lists([],[4,3])")



myTests().main()
</textarea>
</div>
</div>
</li>
<li>
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ut_avg_pos_ac data-question_label="20.12.8">
<div id=ut_avg_pos_ac_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">avg_pos(nums)</span></code> that returns the average of the positive numbers in the list <code class="docutils literal notranslate"><span class="pre">nums</span></code>.  For example, <code class="docutils literal notranslate"><span class="pre">avg_pos([80,</span> <span class="pre">-20,</span> <span class="pre">90])</span></code> should return <code class="docutils literal notranslate"><span class="pre">85.0</span></code>.</p>

</div>
<textarea data-lang="python" id="ut_avg_pos_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def avg_pos(nums):


from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertAlmostEqual(avg_pos([80, -20, 90]),85.0,1,"avg_pos([80, -20, 90])")
        self.assertAlmostEqual(avg_pos([-1, 50, 93]),71.5,1,"avg_pos([-1, 50, 93])")
        self.assertAlmostEqual(avg_pos([-1, -20, 93]),93,1,"avg_pos([-1, -20, 93])")
        self.assertAlmostEqual(avg_pos([27, 34, 93]),51.3,1,"avg_pos([27, 34, 93])")
        self.assertAlmostEqual(avg_pos([83, 95, 93]),90.3,1,"avg_pos([83, 95, 93])")
        self.assertAlmostEqual(avg_pos([72, 81]),76.5,1,"avg_pos([72, 81])")


myTests().main()
</textarea>
</div>
</div>
</li>
<li>
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ut_quartile_ac data-question_label="20.12.9">
<div id=ut_quartile_ac_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">quartile(value)</span></code> that returns <code class="docutils literal notranslate"><span class="pre">0</span></code> if <code class="docutils literal notranslate"><span class="pre">value</span></code> is &lt;= 0.25, <code class="docutils literal notranslate"><span class="pre">1</span></code> if <code class="docutils literal notranslate"><span class="pre">value</span></code> is &gt; 0.25 and &lt;= 0.5, <code class="docutils literal notranslate"><span class="pre">2</span></code> if <code class="docutils literal notranslate"><span class="pre">value</span></code> is &gt; 0.5 and &lt;= 0.75, and <code class="docutils literal notranslate"><span class="pre">3</span></code> if <code class="docutils literal notranslate"><span class="pre">value</span></code> &gt; 0.75.</p>

</div>
<textarea data-lang="python" id="ut_quartile_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def quartille(value):


from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(quartile(-3),0,"quartile(-3)")
        self.assertEqual(quartile(0),0,"quartile(0)")
        self.assertEqual(quartile(0.1),0,"quartile(0.1)")
        self.assertEqual(quartile(0.25),0,"quartile(0.25)")
        self.assertEqual(quartile(0.26),1,"quartile(0.26)")
        self.assertEqual(quartile(0.3),1,"quartile(0.3)")
        self.assertEqual(quartile(0.49),1,"quartile(0.49)")
        self.assertEqual(quartile(0.5),1,"quartile(0.5)")
        self.assertEqual(quartile(0.6),2,"quartile(0.6)")
        self.assertEqual(quartile(0.75),3,"quartile(0.75)")
        self.assertEqual(quartile(0.8),3,"quartile(0.8)")
        self.assertEqual(quartile(2),3,"quartile(2)")


myTests().main()
</textarea>
</div>
</div>
</li>
<li>
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ut_dup_adjacent_fix_ac data-question_label="20.12.10">
<div id=ut_dup_adjacent_fix_ac_question class="ac_question col-md-12">
<p>Fix the function <code class="docutils literal notranslate"><span class="pre">dup_adjacent(nums)</span></code> to return <code class="docutils literal notranslate"><span class="pre">True</span></code> if there are duplicate adjacent values in <code class="docutils literal notranslate"><span class="pre">nums</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">dup_adjacent([1,2,1])</span></code> should return <code class="docutils literal notranslate"><span class="pre">False</span></code> and <code class="docutils literal notranslate"><span class="pre">dup_adjacent([4,</span> <span class="pre">3,</span> <span class="pre">3,</span> <span class="pre">2])</span></code> should return <code class="docutils literal notranslate"><span class="pre">True</span></code>  because of the adjacent 3’s.  Return <code class="docutils literal notranslate"><span class="pre">False</span></code> if the length of the list is less than two.</p>

</div>
<textarea data-lang="python" id="ut_dup_adjacent_fix_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def dup_adjacent(nums):
    if len(nums) == 2:
        return False
    for i in range(len(nums)):
        curr = nums[i]
        prev = nums[i-1]
        if curr == prev:
            return True
        else:
            return False


from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(dup_adjacent([1,2,1]),False,"dup_adjacent([1,2,1])")
        self.assertEqual(dup_adjacent([4, 3, 3, 2]),True,"dup_adjacent([4, 3, 3, 2])")
        self.assertEqual(dup_adjacent([2, 2, 3, 4]),True,"dup_adjacent([2, 2, 3, 4])")
        self.assertEqual(dup_adjacent([2, 1, 4, 4]),True,"dup_adjacent([2, 1, 4, 4])")
        self.assertEqual(dup_adjacent([2, 2]),True,"dup_adjacent([2, 2])")
        self.assertEqual(dup_adjacent([4, 3, 2, 3]),False,"dup_adjacent([4, 3, 2, 3])")
        self.assertEqual(dup_adjacent([]),False,"dup_adjacent([])")
        self.assertEqual(dup_adjacent([3]),False,"dup_adjacent([3])")


myTests().main()
</textarea>
</div>
</div>
</li>
</ol>
</section>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 20.11. Mixed-Up Code Questions' data-toggle="tooltip">
  <a href="mixedupcode.html" aria-label="Previous - 20.11. Mixed-Up Code Questions">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 21. Databases' data-toggle="tooltip" >
    <a href="../database/toctree.html" aria-label="Next - 21. Databases">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>



  
    {% if settings.google_ga %}
      <script>
        var _gaq = _gaq || [];
        _gaq.push(['_setAccount', '{{ settings.google_ga }}']);
        _gaq.push(['_trackPageview']);

        (function() {
          var ga = document.createElement('script'); ga.async = true;
          ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
          var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
        })();
      </script>
    {% endif %}
  

  
    <script type="application/json" class="js-hypothesis-config">
    {
      "showHighlights": false
    }
    </script>

    <script async src="https://hypothes.is/embed.js"></script>
    <!--provide space to prevent hypothes.is sidebar from obscuring content-->
    <style>@media screen and (max-width:767px) { body { padding-right: 21px; }}</style>

  


<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>