<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.17.1: http://docutils.sourceforge.net/" />

    <title>Write Code Questions &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.cd688d1623c9932e.css?v=630908D8" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.bf73008a1eba9f4e.css?v=7DED055E" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/runtime.439d581f758f2b99.bundle.js?v=2E54442"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.51d9dde4e5162190.bundle.js?v=9CAA2D33"></script>
    <script src="../_static/runestone.42b576791a4ed54b.bundle.js?v=B7293199"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    
    eBookConfig.useRunestoneServices = {{use_services}};
    eBookConfig.host = '';
    eBookConfig.app = eBookConfig.host + '/runestone';
    eBookConfig.course = '{{ course_name }}';
    eBookConfig.basecourse = '{{ base_course }}';
    eBookConfig.isLoggedIn = {{ is_logged_in}};
    eBookConfig.email = '{{ user_email }}';
    eBookConfig.isInstructor = {{ is_instructor }};
    eBookConfig.username = '{{ user_id}}';
    eBookConfig.readings = {{ readings|safe}};
    eBookConfig.activities = {{ activity_info|safe }}
    eBookConfig.downloadsEnabled = {{downloads_enabled}};
    eBookConfig.allow_pairs = {{allow_pairs}}
    eBookConfig.enableCompareMe = {{enable_compare_me}};
    // _`new_server_prefix`: Defined by the server to be the path to the new server.
    eBookConfig.new_server_prefix = '{{ new_server_prefix }}';
    
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->

    
        {% if serve_ad is defined and serve_ad and settings.adsenseid %}
         <script data-ad-client="{{settings.adsenseid}}" async src="https://pagead2.googlesyndication.com/pagead/js/adsbygoogle.js"></script>
        {% endif %}
    



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
            
            {{ course_name }}
            
          
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="browsing_warning"></span></li>

        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li id="inst_peer_link"><a href='/runestone/peer/instructor.html'>Peer Instruction (Instructor)</a></li>            
            <li><a href='/runestone/peer/student.html'>Peer Instruction (Student)</a></li>            
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/ns/books/published/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=WriteCode_orig'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
            <script></script>
<li class="dropdown globaltoc-container">
    <a
        href="../index.html"
        class="dropdown-toggle"
        data-toggle="dropdown"
        aria-label="index"
        >This Chapter<b class="caret"></b
    ></a>
    <ul class="dropdown-menu globaltoc">
         
        {% for row in subchapter_list %}
        <li class="toctree-l1">
            <a class="reference internal" href="{{ row['subchap_uri']}}"
                >{{ row['title']}}</a
            >
        </li>
        {% endfor %} 
        
    </ul>
</li>
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->

    
        {% if settings.num_banners > 0 and settings.show_rs_banner %}
         <div style="width: 100%;">
            <a href="/runestone/default/donate?banner={{banner_num}}">
            <img class="runestone-banner"
                 src="/runestone/static/images/RunestoneBanner{{banner_num}}.png"
                 alt="Please Support Runestone"
                 style="display: block; width: 728px; margin-left: auto; margin-right: auto; margin-top: 10px;">
            </a>
         </div>
        {% endif %}
    


  
  <section id="write-code-questions">
<h1>Write Code Questions<a class="headerlink" href="#write-code-questions" title="Permalink to this headline">¶</a></h1>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=str-ex-hiacq data-question_label="1">
<div id=str-ex-hiacq_question class="ac_question col-md-12">
<p>Fix line 2 so that it prints “Hi” instead of “hi”.</p>

</div>
<textarea data-lang="python" id="str-ex-hiacq_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
s1 = "hi"
s1.capitalize()
print(s1)

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(s1,"Hi")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="str-ex-hiansw" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Fix line 2 so that it prints “Hi” instead of “hi”.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="2"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        s1 = "hi"
---
s1 = s1.capitalize()
---
s1.capitalize() #paired
---
print(s1)
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=str-ex-meowacq data-question_label="3">
<div id=str-ex-meowacq_question class="ac_question col-md-12">
<p>Fix the code so that only “meow” is printed.</p>

</div>
<textarea data-lang="python" id="str-ex-meowacq_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
sentence = "The cat goes meow."
s2 = [16:13]sentence
print(s2)

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

        def testOne(self):
            self.assertEqual(s2,"meow")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="str-ex-meowansw" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Fix the code so that only “meow” is printed.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="4"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        sentence = "The cat goes meow."
---
s2 = sentence[13:17]
---
s2 = [16:13]sentence
---
print(s2)
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=str-ex-eggsacq data-question_label="5">
<div id=str-ex-eggsacq_question class="ac_question col-md-12">
<p>Write code to evaluate the length of the string “I like green eggs”
and print it. It should print “The length is 17”.</p>

</div>
<textarea data-lang="python" id="str-ex-eggsacq_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="str-ex-eggsansw" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Write code to evaluate the length of the string “I like green eggs”
and print it. It should print “The length is 17”.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="6"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        sentence = 'I like green eggs'
---
print('The length is ' + str(len(sentence)))
---
print("The length is " + len(sentence)) #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=str-ex-countacq data-question_label="7">
<div id=str-ex-countacq_question class="ac_question col-md-12">
<p>Create a function named <code class="docutils literal notranslate"><span class="pre">count</span></code> that accepts a string and a letter
as arguments, then returns the count of that letter in the string.
For example, if the function call was <code class="docutils literal notranslate"><span class="pre">count(&quot;banana&quot;,</span> <span class="pre">&quot;a&quot;)</span></code> it would
return <code class="docutils literal notranslate"><span class="pre">3</span></code>. <em>Hint: use the</em> <code class="docutils literal notranslate"><span class="pre">count</span></code> <em>method.</em></p>

</div>
<textarea data-lang="python" id="str-ex-countacq_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(count('banana','a'),3,"Tested 'a' in 'banana'")
        self.assertEqual(count('pineapple','s'),0,"Tested 's' in 'pineapple'")
        self.assertEqual(count('pepperoni pizza','p'),4,"Tested 'p' in 'pepperoni pizza'")
        self.assertEqual(count('racecar','r'),2,"Tested 'r' in 'racecar'")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="str-ex-countansw" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function named <code class="docutils literal notranslate"><span class="pre">count</span></code> that accepts a string and a letter
as arguments, then returns the count of that letter in the string.
For example, if the function call was <code class="docutils literal notranslate"><span class="pre">count(&quot;banana&quot;,</span> <span class="pre">&quot;a&quot;)</span></code> it would
return <code class="docutils literal notranslate"><span class="pre">3</span></code>. <em>Hint: use the</em> <code class="docutils literal notranslate"><span class="pre">count</span></code> <em>method.</em></p>

        </div>
        <pre  class="parsonsblocks" data-question_label="8"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def count(str, letter):
---
def count(str, letter) #paired
---
    return str.count(letter)
---
    return count(letter) #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=str-ex-floatacq data-question_label="9">
<div id=str-ex-floatacq_question class="ac_question col-md-12">
<p>Take the following Python code that stores this string: <code class="docutils literal notranslate"><span class="pre">string</span> <span class="pre">=</span> <span class="pre">&quot;X-DSPAM-Confidence:</span> <span class="pre">0.8475&quot;</span></code>.
Use <code class="docutils literal notranslate"><span class="pre">find</span></code> and string slicing to extract the portion of the string after the colon character
and then use the <code class="docutils literal notranslate"><span class="pre">float</span></code> function to convert the extracted string into a floating point number called <code class="docutils literal notranslate"><span class="pre">num</span></code>.</p>

</div>
<textarea data-lang="python" id="str-ex-floatacq_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
string = "X-DSPAM-Confidence: 0.8475"


====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(num,0.8475)

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="str-ex-floatansw" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Take the following Python code that stores this string: <code class="docutils literal notranslate"><span class="pre">string</span> <span class="pre">=</span> <span class="pre">&quot;X-DSPAM-Confidence:</span> <span class="pre">0.8475&quot;</span></code>.
Use <code class="docutils literal notranslate"><span class="pre">find</span></code> and string slicing to extract the portion of the string after the colon character
and then use the <code class="docutils literal notranslate"><span class="pre">float</span></code> function to convert the extracted string into a floating point number called <code class="docutils literal notranslate"><span class="pre">num</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="10"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        string = "X-DSPAM-Confidence: 0.8475"
---
colonpos = string.find(':')
---
digit = string[colonpos + 1:]
---
digit = string[colonpos:] #paired
---
num = float(digit)
---
num = int(digit) #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=str-ex-ducksacq data-question_label="11">
<div id=str-ex-ducksacq_question class="ac_question col-md-12">
<p>In Robert McCloskey’s book <em>Make Way for Ducklings</em>, the names of the ducklings are
Jack, Kack, Lack, Mack, Nack, Ouack, Pack, and Quack.
The following loop tries to output these names in order.
Of course, that’s not quite right because Ouack and Quack are misspelled.
Can you fix it?</p>

</div>
<textarea data-lang="python" id="str-ex-ducksacq_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
prefixes = "JKLMNOPQ"
suffix = "ack"

for letter in prefixes:
    print(letter + suffix)
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="str-ex-ducksansw" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>In Robert McCloskey’s book <em>Make Way for Ducklings</em>, the names of the ducklings are
Jack, Kack, Lack, Mack, Nack, Ouack, Pack, and Quack.
The following loop tries to output these names in order.
Of course, that’s not quite right because Ouack and Quack are misspelled.
Can you fix it?</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="12"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        prefixes = "JKLMNOPQ"
suffix = "ack"
---
for letter in prefixes:
---
    if letter == 'O' or letter == 'Q':
        print(letter + "u" + suffix)
---
    if letter = 'o' or letter = 'Q': #paired
        print(letter + "u" + suffix)
---
    else:
        print(letter + suffix)
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=str-ex-digitsacq data-question_label="13">
<div id=str-ex-digitsacq_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">numDigits</span></code> that will return the number of digits in an integer <code class="docutils literal notranslate"><span class="pre">n</span></code>.</p>

</div>
<textarea data-lang="python" id="str-ex-digitsacq_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def numDigits(n):
    # your code here

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(numDigits(2),1,"Tested numDigits on input of 2")
        self.assertEqual(numDigits(55),2,"Tested numDigits on input of 55")
        self.assertEqual(numDigits(1352),4,"Tested numDigits on input of 1352")
        self.assertEqual(numDigits(444),3,"Tested numDigits on input of 444")
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="str-ex-digitsansw" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Write a function <code class="docutils literal notranslate"><span class="pre">numDigits</span></code> that will return the number of digits in an integer <code class="docutils literal notranslate"><span class="pre">n</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="14"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def numDigits(n):
---
def numDigits(n) #paired
---
    n_str = str(n)
---
    return len(n_str)
---
    return len(n) #paired
---
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=str-ex-nameacq data-question_label="15">
<div id=str-ex-nameacq_question class="ac_question col-md-12">
<p>Write code to print out the statement “Hi my name is Bob and I am 2” using only string methods
or string slicing. You must get every part of the new string from the given strings, not by using
string literals.
Name the final string <code class="docutils literal notranslate"><span class="pre">statement</span></code>.</p>

</div>
<textarea data-lang="python" id="str-ex-nameacq_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
s1 = "hi"
s2 = "My name is Bob, and he and I love to eat ham."

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(statement,"Hi my name is Bob and I am 2")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="str-ex-nameansw" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Write code to print out the statement “Hi my name is Bob and I am 2” using only string methods
or string slicing. You must get every part of the new string from the given strings, not by using
string literals.
Name the final string <code class="docutils literal notranslate"><span class="pre">statement</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="16"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        s1 = "hi"
s2 = "My name is Bob, and he and I love to eat ham."
---
statement = s1.capitalize() + s2[2:3] + s2.lower()[0:3] + s2[3:14] + s2[22:29] + s2[4:6] + s2[2:3] + str(len(s1))
---
statement = s1.capitalize() + s2[7:8] + s2[5:6] + s2[1:2] + s2[3:15] + s2[22:29] + s2[4:6] + s2[2:3] + str(len(s1)) #paired
---
print(statement)
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=str-ex-nameInputacq data-question_label="17">
<div id=str-ex-nameInputacq_question class="ac_question col-md-12">
<p>Write a program that asks a user for their name and from the input
prints the first letter of their name in lowercase.</p>

</div>
<textarea data-lang="python" id="str-ex-nameInputacq_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="str-ex-nameInputansw" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Write a program that asks a user for their name and from the input
prints the first letter of their name in lowercase.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="18"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        prompt = "What is your name?"
---
name = input(prompt)
---
first = name[0]
---
lowerFirst = first.lower()
---
first.lower() #paired
---
print(lowerFirst)
---
print(first) #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=str-ex-loweracq data-question_label="19">
<div id=str-ex-loweracq_question class="ac_question col-md-12">
<p>Write a program that asks for user input and prints their
input in all lowercase, as well as the length of their string.</p>

</div>
<textarea data-lang="python" id="str-ex-loweracq_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="str-ex-loweransw" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Write a program that asks for user input and prints their
input in all lowercase, as well as the length of their string.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="20"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        user_input = input("Enter a string: ")
---
print(user_input.lower() + " " + str(len(user_input)))
---
print(user_input.lower() + " " + len(user_input)) #paired
        </pre>
        </div>
        </div>
    </section>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>



  
    {% if settings.google_ga %}
      <script>
        var _gaq = _gaq || [];
        _gaq.push(['_setAccount', '{{ settings.google_ga }}']);
        _gaq.push(['_trackPageview']);

        (function() {
          var ga = document.createElement('script'); ga.async = true;
          ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
          var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
        })();
      </script>
    {% endif %}
  

  
    <script type="application/json" class="js-hypothesis-config">
    {
      "showHighlights": false
    }
    </script>

    <script async src="https://hypothes.is/embed.js"></script>
    <!--provide space to prevent hypothes.is sidebar from obscuring content-->
    <style>@media screen and (max-width:767px) { body { padding-right: 21px; }}</style>

  


<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>