<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.17.1: http://docutils.sourceforge.net/" />

    <title>Mixed-Up Code Exercises &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.cd688d1623c9932e.css?v=630908D8" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.bf73008a1eba9f4e.css?v=7DED055E" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/runtime.439d581f758f2b99.bundle.js?v=2E54442"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.51d9dde4e5162190.bundle.js?v=9CAA2D33"></script>
    <script src="../_static/runestone.42b576791a4ed54b.bundle.js?v=B7293199"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    
    eBookConfig.useRunestoneServices = {{use_services}};
    eBookConfig.host = '';
    eBookConfig.app = eBookConfig.host + '/runestone';
    eBookConfig.course = '{{ course_name }}';
    eBookConfig.basecourse = '{{ base_course }}';
    eBookConfig.isLoggedIn = {{ is_logged_in}};
    eBookConfig.email = '{{ user_email }}';
    eBookConfig.isInstructor = {{ is_instructor }};
    eBookConfig.username = '{{ user_id}}';
    eBookConfig.readings = {{ readings|safe}};
    eBookConfig.activities = {{ activity_info|safe }}
    eBookConfig.downloadsEnabled = {{downloads_enabled}};
    eBookConfig.allow_pairs = {{allow_pairs}}
    eBookConfig.enableCompareMe = {{enable_compare_me}};
    // _`new_server_prefix`: Defined by the server to be the path to the new server.
    eBookConfig.new_server_prefix = '{{ new_server_prefix }}';
    
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->

    
        {% if serve_ad is defined and serve_ad and settings.adsenseid %}
         <script data-ad-client="{{settings.adsenseid}}" async src="https://pagead2.googlesyndication.com/pagead/js/adsbygoogle.js"></script>
        {% endif %}
    



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
            
            {{ course_name }}
            
          
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="browsing_warning"></span></li>

        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li id="inst_peer_link"><a href='/runestone/peer/instructor.html'>Peer Instruction (Instructor)</a></li>            
            <li><a href='/runestone/peer/student.html'>Peer Instruction (Student)</a></li>            
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/ns/books/published/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=MixedUpCode_orig'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
            <script></script>
<li class="dropdown globaltoc-container">
    <a
        href="../index.html"
        class="dropdown-toggle"
        data-toggle="dropdown"
        aria-label="index"
        >This Chapter<b class="caret"></b
    ></a>
    <ul class="dropdown-menu globaltoc">
         
        {% for row in subchapter_list %}
        <li class="toctree-l1">
            <a class="reference internal" href="{{ row['subchap_uri']}}"
                >{{ row['title']}}</a
            >
        </li>
        {% endfor %} 
        
    </ul>
</li>
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->

    
        {% if settings.num_banners > 0 and settings.show_rs_banner %}
         <div style="width: 100%;">
            <a href="/runestone/default/donate?banner={{banner_num}}">
            <img class="runestone-banner"
                 src="/runestone/static/images/RunestoneBanner{{banner_num}}.png"
                 alt="Please Support Runestone"
                 style="display: block; width: 728px; margin-left: auto; margin-right: auto; margin-top: 10px;">
            </a>
         </div>
        {% endif %}
    


  
  <section id="mixed-up-code-exercises">
<h1>Mixed-Up Code Exercises<a class="headerlink" href="#mixed-up-code-exercises" title="Permalink to this headline">¶</a></h1>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Classes_Basic_Airport_pp" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class <code class="docutils literal notranslate"><span class="pre">Airport</span></code> with an <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method that takes a <code class="docutils literal notranslate"><span class="pre">name</span></code> and <code class="docutils literal notranslate"><span class="pre">code</span></code> as strings and initializes these attributes in the current object.
Then define the <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method to return the <code class="docutils literal notranslate"><span class="pre">name</span> <span class="pre">code</span></code>.  For example, <code class="docutils literal notranslate"><span class="pre">print(a)</span></code> when <code class="docutils literal notranslate"><span class="pre">a</span> <span class="pre">=</span> <span class="pre">Airport(&quot;Detroit&quot;,</span> <span class="pre">&quot;DTW&quot;)</span></code> would print <code class="docutils literal notranslate"><span class="pre">Detroit</span> <span class="pre">DTW</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="1"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        class Airport:
---
def Airport: #paired
---
  def __init__(self, name, code):
---
  def __init__(name, code): #paired
---
    self.name = name
    self.code = code
---
  def __str__(self):
---
  def str(self): #paired
---
    return self.name + " " + self.code
---
    return name + " " + code #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Classes_Basic_Airport_ac data-question_label="2">
<div id=Classes_Basic_Airport_ac_question class="ac_question col-md-12">
<p>Write a class <code class="docutils literal notranslate"><span class="pre">Airport</span></code> with an <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method that takes a <code class="docutils literal notranslate"><span class="pre">name</span></code> and <code class="docutils literal notranslate"><span class="pre">code</span></code> as strings and initializes these attributes in the current object.
Then define the <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method to return the <code class="docutils literal notranslate"><span class="pre">name</span> <span class="pre">code</span></code>.  For example, <code class="docutils literal notranslate"><span class="pre">print(a)</span></code> when <code class="docutils literal notranslate"><span class="pre">a</span> <span class="pre">=</span> <span class="pre">Airport(&quot;Detroit&quot;,</span> <span class="pre">&quot;DTW&quot;)</span></code> would print <code class="docutils literal notranslate"><span class="pre">Detroit</span> <span class="pre">DTW</span></code>.</p>

</div>
<textarea data-lang="python" id="Classes_Basic_Airport_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

a = Airport("Detroit", "DTW")
print(a)

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        a = Airport("Detroit", "DTW")
        self.assertEqual(a.__str__(), "Detroit DTW", "testing __str__ for Detroit")
        a = Airport("Atlanta", "ATL")
        self.assertEqual(a.__str__(), "Atlanta ATL", "testing __str__ for Atlanta")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Classes_Basic_Song_pp" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class <code class="docutils literal notranslate"><span class="pre">Song</span></code> with an <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method that takes a <code class="docutils literal notranslate"><span class="pre">title</span></code> as a string and <code class="docutils literal notranslate"><span class="pre">len</span></code> as a number and initializes these attributes in the current object.
Then define the <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method to return the <code class="docutils literal notranslate"><span class="pre">title,</span> <span class="pre">len</span></code>.  For example, <code class="docutils literal notranslate"><span class="pre">print(s)</span></code> when <code class="docutils literal notranslate"><span class="pre">s</span> <span class="pre">=</span> <span class="pre">Song('Respect',150)</span></code> would print “Respect, 150”.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="3"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        class Song:
---
Class Song: #paired
---
  def __init__(self, title, len):
---
  def __init__(title, len): #paired
---
    self.title = title
    self.len = len
---
  def __str__(self):
---
  def __str__(): #paired
---
    return self.title + ", " + str(self.len)
---
    return title + ", " + len #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Classes_Basic_Song_ac data-question_label="4">
<div id=Classes_Basic_Song_ac_question class="ac_question col-md-12">
<p>Write a class <code class="docutils literal notranslate"><span class="pre">Song</span></code> with an <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method that takes a <code class="docutils literal notranslate"><span class="pre">title</span></code> as a string and <code class="docutils literal notranslate"><span class="pre">len</span></code> as a number and initializes these attributes in the current object.
Then define the <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method to return the <code class="docutils literal notranslate"><span class="pre">title,</span> <span class="pre">len</span></code>.  For example, <code class="docutils literal notranslate"><span class="pre">print(s)</span></code> when <code class="docutils literal notranslate"><span class="pre">s</span> <span class="pre">=</span> <span class="pre">Song('Respect',150)</span></code> would print “Respect, 150”.</p>

</div>
<textarea data-lang="python" id="Classes_Basic_Song_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

s = Song('Respect',150)
print(s)

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        s = Song('Respect',150)
        self.assertEqual(s.__str__(), "Respect, 150", "testing __str__ with Respect")
        s = Song('Truth Hurts',183)
        self.assertEqual(s.__str__(), "Truth Hurts, 183", "testing __str__ with Truth Hurts")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Classes_Basic_Cat_pp" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class Cat with an <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method that takes
<code class="docutils literal notranslate"><span class="pre">name</span></code> as a string and <code class="docutils literal notranslate"><span class="pre">age</span></code> as a number and initializes these attributes in the current object.  Next create the <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns
“name: name age: age”.  For example if <code class="docutils literal notranslate"><span class="pre">c</span> <span class="pre">=</span> <span class="pre">Cat(&quot;Fluffy&quot;,</span> <span class="pre">3)</span></code> then
<code class="docutils literal notranslate"><span class="pre">print(c)</span></code> should print <code class="docutils literal notranslate"><span class="pre">&quot;name:</span> <span class="pre">Fluffy</span> <span class="pre">age:</span> <span class="pre">3&quot;</span></code>.
Then define the <code class="docutils literal notranslate"><span class="pre">make_sound</span></code> method to return <code class="docutils literal notranslate"><span class="pre">&quot;Meow&quot;</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="5"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        class Cat:
---
  def __init__(self, name, age):
---
    self.name = name
    self.age = age
---
  def __str__(self):
---
    return f"name: {self.name}, age: {self.age}"
---
    return f"name: {name}, age: {age}" #paired
---
  def make_sound(self):
---
  def make_sound(): #paired
---
    return "Meow"
---
    return self."Meow" #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Classes_Basic_Cat_ac data-question_label="6">
<div id=Classes_Basic_Cat_ac_question class="ac_question col-md-12">
<p>Write a class Cat with an <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method that takes
<code class="docutils literal notranslate"><span class="pre">name</span></code> as a string and <code class="docutils literal notranslate"><span class="pre">age</span></code> as a number and initializes these attributes in the current object.  Next create the <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns
“name: name, age: age”.  For example if <code class="docutils literal notranslate"><span class="pre">c</span> <span class="pre">=</span> <span class="pre">Cat(&quot;Fluffy&quot;,</span> <span class="pre">3)</span></code> then
<code class="docutils literal notranslate"><span class="pre">print(c)</span></code> should print <code class="docutils literal notranslate"><span class="pre">&quot;name:</span> <span class="pre">Fluffy,</span> <span class="pre">age:</span> <span class="pre">3&quot;</span></code>.
Then define the <code class="docutils literal notranslate"><span class="pre">make_sound</span></code> method to return <code class="docutils literal notranslate"><span class="pre">&quot;Meow&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="Classes_Basic_Cat_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

c = Cat("Fluffy", 3)
print(c)
print(c.make_sound())

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        c = Cat("Fluffy", 3)
        self.assertEqual(c.__str__(), "name: Fluffy, age: 3", "testing __str__ on Fluffy")
        self.assertEqual(c.make_sound(), "Meow", "testing make_sound with Fluffy")
        d = Cat("Spike", 2)
        self.assertEqual(d.__str__(), "name: Spike, age: 2", "testing __str__ on Spike")
        self.assertEqual(d.make_sound(), "Meow", "testing make_sound with Spike")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Classes_Basic_Account_pp" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class <code class="docutils literal notranslate"><span class="pre">Account</span></code> with an <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method that takes <code class="docutils literal notranslate"><span class="pre">id</span></code> and <code class="docutils literal notranslate"><span class="pre">balance</span></code> as numbers. Then create a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns “id, balance”. Next create a <code class="docutils literal notranslate"><span class="pre">deposit</span></code> method takes <code class="docutils literal notranslate"><span class="pre">amount</span></code> as a number and adds that to the <code class="docutils literal notranslate"><span class="pre">balance</span></code>.  For example,
if <code class="docutils literal notranslate"><span class="pre">a</span> <span class="pre">=</span> <span class="pre">Account(32,</span> <span class="pre">100)</span></code> and <code class="docutils literal notranslate"><span class="pre">a.deposit(50)</span></code> is executed, <code class="docutils literal notranslate"><span class="pre">print(a)</span></code> should print “32, 150”.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="7"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        ---
class Account:
---
    def __init__(self, id, balance):
---
    def __init__(self, id, balance) #paired
---
        self.id = id
        self.balance = balance
---
    def __str__(self):
---
        return f"{self.id}, {self.balance}"
---
        return f"{id}, {balance}" #paired
---
    def deposit(self, amount):
---
        self.balance += amount
---
        self.balance += self.amount #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Classes_Basic_Account_ac data-question_label="8">
<div id=Classes_Basic_Account_ac_question class="ac_question col-md-12">
<p>Create a class <code class="docutils literal notranslate"><span class="pre">Account</span></code> with an <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method that takes <code class="docutils literal notranslate"><span class="pre">id</span></code> and <code class="docutils literal notranslate"><span class="pre">balance</span></code> as numbers. Then create a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns “id, balance”. Next create a <code class="docutils literal notranslate"><span class="pre">deposit</span></code> method takes <code class="docutils literal notranslate"><span class="pre">amount</span></code> as a number and adds that to the <code class="docutils literal notranslate"><span class="pre">balance</span></code>.  For example,
if <code class="docutils literal notranslate"><span class="pre">a</span> <span class="pre">=</span> <span class="pre">Account(32,</span> <span class="pre">100)</span></code> and <code class="docutils literal notranslate"><span class="pre">a.deposit(50)</span></code> is executed, <code class="docutils literal notranslate"><span class="pre">print(a)</span></code> should print “32, 150”.</p>

</div>
<textarea data-lang="python" id="Classes_Basic_Account_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

a = Account(32, 100)
a.deposit(50)
print(a)

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        a = Account(32, 100)
        self.assertEqual(a.__str__(), "32, 100", "testing __str__ on a")
        a.deposit(50)
        self.assertEqual(a.__str__(), "32, 150", "testing __str__ on a after deposit")
        b = Account(2, 225.50)
        self.assertEqual(b.__str__(), "2, 225.5", "testing __str__ on b")
        b.deposit(50.25)
        self.assertEqual(b.__str__(), "2, 275.75", "testing __str__ on b after deposit")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Classes_Basic_FortuneTeller_pp" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class <code class="docutils literal notranslate"><span class="pre">FortuneTeller</span></code> with an <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method that takes a list of fortunes as strings and saves that as an attribute.  Then create a <code class="docutils literal notranslate"><span class="pre">tell_fortune</span></code> method that returns one of the fortunes in the list at random.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="9"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        import random
---
class FortuneTeller:
---
    def __init__(self, fortunes):
---
        self.fortunes = fortunes
---
    def tell_fortune(self):
---
    def tell_fortune(): #paired
---
        last = len(self.fortunes) - 1
---
        last = len(self.fortunes) #paired
---
        index = random.randint(0, last)
---
        return self.fortunes[index]
---
        return fortunes[index] #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Classes_Basic_FortuneTeller_ac data-question_label="10">
<div id=Classes_Basic_FortuneTeller_ac_question class="ac_question col-md-12">
<p>Write a class <code class="docutils literal notranslate"><span class="pre">FortuneTeller</span></code> with an <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method that takes a list of fortunes as strings and saves that as an attribute.  Then create a <code class="docutils literal notranslate"><span class="pre">tell_fortune</span></code> method that returns one of the fortunes in the list at random.</p>

</div>
<textarea data-lang="python" id="Classes_Basic_FortuneTeller_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

f = FortuneTeller(["You will get an A", "You will have a great day", "You will fall", "You will laugh"])
for i in range(5):
     print(f.tell_fortune())

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        l = ["You will get an A", "You will have a great day", "You will fall", "You will laugh"]
        f = FortuneTeller(l)
        for i in range(5):
            s = f.tell_fortune()
            self.assertIn(s, l, "checking that the fortune in is the list")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Classes_Basic_Student_pp" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class <code class="docutils literal notranslate"><span class="pre">Student</span></code> with an <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method that takes a student’s name as a string and a list of exam scores as integers. Then create a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns the “name: average exam score” rounded to the nearest integer. For example, <code class="docutils literal notranslate"><span class="pre">print(s)</span></code> when <code class="docutils literal notranslate"><span class="pre">s</span> <span class="pre">=</span> <span class="pre">Student(&quot;Sarah&quot;,</span> <span class="pre">[91,92,97])</span></code> would print <code class="docutils literal notranslate"><span class="pre">Sarah:</span> <span class="pre">93.0</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="11"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        class Student:
---
  def __init__(self, name, scores):
---
  def init(self, name, scores): #paired
---
    self.name = name
    self.scores = scores
---
    name = self.name
    scores = self.scores #paired
---
  def __str__(self):
---
    total_scores  = sum(self.scores)
---
    total_scores = self.scores.sum() #paired
---
    average = round(total_scores / len(self.scores), 0)
---
    return f"{self.name}: {average}"
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Classes_Basic_Student_ac data-question_label="12">
<div id=Classes_Basic_Student_ac_question class="ac_question col-md-12">
<p>Write a class <code class="docutils literal notranslate"><span class="pre">Student</span></code> with an <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method that takes a student’s name as a string and a list of exam scores as integers. Then create a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns the “name: average exam score” rounded to the nearest integer. For example, <code class="docutils literal notranslate"><span class="pre">print(s)</span></code> when <code class="docutils literal notranslate"><span class="pre">s</span> <span class="pre">=</span> <span class="pre">Student(&quot;Sarah&quot;,</span> <span class="pre">[91,92,97])</span></code> would print <code class="docutils literal notranslate"><span class="pre">Sarah:</span> <span class="pre">93.0</span></code>.</p>

</div>
<textarea data-lang="python" id="Classes_Basic_Student_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Student:
     #write code here

s = Student("Sarah", [91,92,97])
print(s)

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        name = "James"
        scores = [92, 80, 76]
        s = Student(name, scores)
        self.assertEqual(s.__str__(), "James: 83.0", "Checking that the Student's average exam score is correct")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Classes_Basic_Encoder_pp" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class <code class="docutils literal notranslate"><span class="pre">Encoder</span></code> with an initializer method that takes in a sentence as a string. Then create a string method that returns “string: length of sentence”. Lastly, create a consonants method that returns the sentence without vowels, all lowercase. For example, <code class="docutils literal notranslate"><span class="pre">print(e)</span></code> when <code class="docutils literal notranslate"><span class="pre">e</span> <span class="pre">=</span> <span class="pre">Encoder(&quot;I</span> <span class="pre">love</span> <span class="pre">Python&quot;)</span></code> would print <code class="docutils literal notranslate"><span class="pre">I</span> <span class="pre">love</span> <span class="pre">Python:</span> <span class="pre">13</span></code> and <code class="docutils literal notranslate"><span class="pre">print(e.consonants())</span></code> would print `` lv pythn``.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="13"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        class Encoder:
---
Class Encoder:  #paired
---
  def __init__(self, sentence):
---
    self.sentence = sentence
---
  def __str__(self):
---
    return f"{self.sentence}: {len(self.sentence)}"
---
    return self.sentence + ": " + len(self.sentence) #paired
---
  def consonants(self):
---
  def consonants(): #paired
---
    no_vowels = ""
---
    for char in self.sentence:
---
    for char in sentence: #paired
---
      char = char.lower()
---
      if char not in ["a", "e", "i", "o", "u"]:
---
        no_vowels += char
---
    return no_vowels
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Classes_Basic_Encoder_ac data-question_label="14">
<div id=Classes_Basic_Encoder_ac_question class="ac_question col-md-12">
<p>Write a class <code class="docutils literal notranslate"><span class="pre">Encoder</span></code> with an initializer method that takes in a sentence as a string. Then create a string method that returns “string: length of sentence”. Lastly, create a consonants method that returns the sentence without vowels, all lowercase. For example, <code class="docutils literal notranslate"><span class="pre">print(e)</span></code> when <code class="docutils literal notranslate"><span class="pre">e</span> <span class="pre">=</span> <span class="pre">Encoder(&quot;I</span> <span class="pre">love</span> <span class="pre">Python&quot;)</span></code> would print <code class="docutils literal notranslate"><span class="pre">I</span> <span class="pre">love</span> <span class="pre">Python:</span> <span class="pre">13</span></code> and <code class="docutils literal notranslate"><span class="pre">print(e.consonants())</span></code> would print `` lv pythn``.</p>

</div>
<textarea data-lang="python" id="Classes_Basic_Encoder_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Encoder:
     #write code here

e = Encoder("I love Python")
print(e)
print(e.consonants())

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        sentence1 = "Hello world!"
        c1 = Encoder(sentence1)
        sentence2 = "Abraham Lincoln was the 16th President of the United States."
        c2 = Encoder(sentence2)
        self.assertEqual(c1.consonants(), "hll wrld!", "Checking that all vowels have been removed from sentence and all chars converted to lowercase")
        self.assertEqual(c2.consonants(), "brhm lncln ws th 16th prsdnt f th ntd stts.", "Checking that all vowels have been removed from sentence and all chars converted to lowercase")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Classes_Basic_CharFrequency_pp" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class <code class="docutils literal notranslate"><span class="pre">CharFrequency</span></code> with an initializer method that takes in a sentence as a string. Then create a string method that returns “The most common character in the sentence is x.” In the case of tiebreakers, return the character that comes last in alphabetical order.
For example, <code class="docutils literal notranslate"><span class="pre">print(c)</span></code> when <code class="docutils literal notranslate"><span class="pre">c</span> <span class="pre">=</span> <span class="pre">CharFreequency(&quot;I</span> <span class="pre">love</span> <span class="pre">Python&quot;)</span></code> would print <code class="docutils literal notranslate"><span class="pre">The</span> <span class="pre">most</span> <span class="pre">common</span> <span class="pre">character</span> <span class="pre">in</span> <span class="pre">the</span> <span class="pre">sentence</span> <span class="pre">is</span> <span class="pre">o.</span></code></p>

        </div>
        <pre  class="parsonsblocks" data-question_label="15"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        class CharFrequency:
---
  def __init__(self, sentence):
---
    self.sentence = sentence
---
  def __str__(self):
---
    most_common = {}
---
    sentence = self.sentence.replace(" ", "")
---
    for char in sentence:
---
      most_common[char] = most_common.get(char, 0) + 1
---
    sorted_list = sorted(most_common.items(), key = lambda k: (k[1], k[0]), reverse = True)
---
    return "The most common character in the sentence is " + str(sorted_list[0][0]) + "."
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Classes_Basic_CharFrequency_ac data-question_label="16">
<div id=Classes_Basic_CharFrequency_ac_question class="ac_question col-md-12">
<p>Write a class <code class="docutils literal notranslate"><span class="pre">CharFrequency</span></code> with an initializer method that takes in a sentence as a string. Then create a string method that returns “The most common character in the sentence is x.” In the case of tiebreakers, return the character that comes last in alphabetical order. For example, <code class="docutils literal notranslate"><span class="pre">print(c)</span></code> when <code class="docutils literal notranslate"><span class="pre">c</span> <span class="pre">=</span> <span class="pre">CharFreequency(&quot;I</span> <span class="pre">love</span> <span class="pre">Python&quot;)</span></code> would print <code class="docutils literal notranslate"><span class="pre">The</span> <span class="pre">most</span> <span class="pre">common</span> <span class="pre">character</span> <span class="pre">in</span> <span class="pre">the</span> <span class="pre">sentence</span> <span class="pre">is</span> <span class="pre">o.</span></code></p>

</div>
<textarea data-lang="python" id="Classes_Basic_CharFrequency_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class CharFrequency:
     #write code here


c = CharFrequency("I love Python")
print(c)

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        sentence1 = "My dog likes to play fetch."
        m1 = CharFrequency(sentence1)
        sentence2 = "5 times 5 equals 25."
        m2 = CharFrequency(sentence2)
        self.assertEqual(m1.__str__(), "The most common character in the sentence is y.", "Checking that the most common char is returned")
        self.assertEqual(m2.__str__(), "The most common character in the sentence is 5.", "Checking that the most common char is returned")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Classes_Basic_Loan_pp" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class <code class="docutils literal notranslate"><span class="pre">Loan</span></code> with an initializer method that takes in an integer loan_amt, a float interest_rate and integer years. Then create a string method that returns “Your loan is ‘’loan_amt’’ for ‘’years’’ years at an interest rate of ‘’interest_rate’’.” Create another method called total_payment that calculates the total amount of loan payment based on the interest rate and number of years. Use the formula (loan_amt * (1 + interest_rate/100) ^ years). Also round your total payment to two decimal places. We’re assuming that you are paying on an annual basis. For example, <code class="docutils literal notranslate"><span class="pre">print(l)</span></code> when <code class="docutils literal notranslate"><span class="pre">l</span> <span class="pre">=</span> <span class="pre">Loan(10000,</span> <span class="pre">5,</span> <span class="pre">10)</span></code> would print <code class="docutils literal notranslate"><span class="pre">Your</span> <span class="pre">loan</span> <span class="pre">is</span> <span class="pre">$10000</span> <span class="pre">for</span> <span class="pre">10</span> <span class="pre">years</span> <span class="pre">at</span> <span class="pre">an</span> <span class="pre">interest</span> <span class="pre">rate</span> <span class="pre">of</span> <span class="pre">5%.</span></code> and <code class="docutils literal notranslate"><span class="pre">print(l.total_payment())</span></code> would print <code class="docutils literal notranslate"><span class="pre">16288.95</span></code> as a float.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="17"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        class Loan:
---
  def __init__(self, loan_amt, interest_rate, years):
---
    self.loan_amt = loan_amt
    self.interest_rate = interest_rate
    self.years = years
---
  def __str__(self):
---
    return f"Your loan is ${self.loan_amt} for {self.years} years at an interest rate of {self.interest_rate}%."
---
    return f"Your loan is $(self.loan_amt) for (self.years) years at an interest rate of (self.interest_rate)%." #paired
---
  def total_payment(self):
---
    total = self.loan_amt * (1 + self.interest_rate/100) ** self.years
---
    total = loan_amt * (1 + interest_rate/100) ** years  #paired
---
    return float('{:.2f}'.format(total))
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Classes_Basic_Loan_ac data-question_label="18">
<div id=Classes_Basic_Loan_ac_question class="ac_question col-md-12">
<p>Write a class <code class="docutils literal notranslate"><span class="pre">Loan</span></code> with an initializer method that takes in an integer loan_amt, a float interest_rate and integer years. Then write a string method that returns “Your loan is ‘’loan_amt’’ for ‘’years’’ years at an interest rate of ‘’interest_rate’’.” Write another method called total_payment that calculates the total amount of loan payment based on the interest rate and number of years. Use the formula (loan_amt * (1 + interest_rate/100) ^ years). Also round your total payment to two decimal places. We’re assuming that you are paying on an annual basis. For example, <code class="docutils literal notranslate"><span class="pre">print(l)</span></code> when <code class="docutils literal notranslate"><span class="pre">l</span> <span class="pre">=</span> <span class="pre">Loan(10000,</span> <span class="pre">5,</span> <span class="pre">10)</span></code> would print <code class="docutils literal notranslate"><span class="pre">Your</span> <span class="pre">loan</span> <span class="pre">is</span> <span class="pre">$10000</span> <span class="pre">for</span> <span class="pre">10</span> <span class="pre">years</span> <span class="pre">at</span> <span class="pre">an</span> <span class="pre">interest</span> <span class="pre">rate</span> <span class="pre">of</span> <span class="pre">5%.</span></code> and <code class="docutils literal notranslate"><span class="pre">print(l.total_payment())</span></code> would print <code class="docutils literal notranslate"><span class="pre">16288.95</span></code> as a float.</p>

</div>
<textarea data-lang="python" id="Classes_Basic_Loan_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Loan:
     #write code here

l = Loan(10000, 5, 10)
print(l)
print(l.total_payment())

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        loan_amt = 10000
        interest_rate = 5
        years = 10
        l1 = Loan(10000, 5, 10)
        l2 = Loan(10000, 2.5, 15)
        self.assertEqual(l1.total_payment(), 16288.95, "Checking the total_payment method")
        self.assertEqual(l2.total_payment(), 14482.98, "Checking the total_payment method")
        self.assertEqual(l1.__str__(), "Your loan is $10000 for 10 years at an interest rate of 5%.", "Checking the string method")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Classes_Basic_Cart1_pp" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class <code class="docutils literal notranslate"><span class="pre">Cart</span></code> that takes in a dictionary price_dict with keys as product names and values as the product price. The initializer method also creates a new empty list called cart_list that will contain tuples in the format (item, quantity). The class also has an add_item and a calculate_total method. The add_item method will take in an item and quantity you wish to purchase and append those to the cart. The calculate_total method should calculate and return the total amount of your purchase for all of the items in your cart. For example, when <code class="docutils literal notranslate"><span class="pre">c</span> <span class="pre">=</span> <span class="pre">Cart({“Notebook”:</span> <span class="pre">1.99,</span> <span class="pre">“Pen”:</span> <span class="pre">2.50,</span> <span class="pre">“Paper”:</span> <span class="pre">1.00})</span></code> and a user wants to purchase 3 pens by calling c.add_item(“Pen”, 3), then print(c.calculate_total) would print <code class="docutils literal notranslate"><span class="pre">7.50</span></code> as a float.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="19"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        class Cart:
---
  def __init__(self, price_dict):
---
    self.price_dict = price_dict
    self.cart_list = []
---
    self.price_dict = price_dict #paired
    self.cart_list = cart_list
---
  def add_item(self, item, quantity):
---
    self.cart_list.append((item, quantity))
---
    self.cart_list.append(item, quantity) #paired
---
  def calculate_total(self):
---
    total = 0
---
    for item, quantity in self.cart_list:
---
      total += quantity * self.price_dict[item]
---
    return float('{:.2f}'.format(total))
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Classes_Basic_Cart1_ac data-question_label="20">
<div id=Classes_Basic_Cart1_ac_question class="ac_question col-md-12">
<p>Create a class <code class="docutils literal notranslate"><span class="pre">Cart</span></code> that takes in a dictionary price_dict with keys as product names and values as the product price. The initializer method also creates a new empty list called cart_list that will contain tuples in the format (item, quantity). The class also has an add_item and a calculate_total method. The add_item method will take in an item and quantity you wish to purchase and append those to the cart. The calculate_total method should calculate and return the total amount of your purchase for all of the items in your cart. For example, when <code class="docutils literal notranslate"><span class="pre">c</span> <span class="pre">=</span> <span class="pre">Cart({“Notebook”:</span> <span class="pre">1.99,</span> <span class="pre">“Pen”:</span> <span class="pre">2.50,</span> <span class="pre">“Paper”:</span> <span class="pre">1.00})</span></code> and a user wants to purchase 3 pens by calling c.add_item(“Pen”, 3), then print(c.calculate_total) would print <code class="docutils literal notranslate"><span class="pre">7.50</span></code> as a float.</p>

</div>
<textarea data-lang="python" id="Classes_Basic_Cart1_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Cart:
     #write code here

c = Cart({"Notebook": 1.99, "Pen": 2.50, "Paper": 1.00})
c.add_item("Pen", 3)
print(c.calculate_total())

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        c = Cart({"Banana": 0.19, "Apple": 1.79, "Blueberries": 3.99, "Strawberries": 4.98})
        c.add_item("Banana", 5)
        c.add_item("Strawberries", 1)
        self.assertEqual(c.cart_list, [('Banana', 5), ('Strawberries', 1)], "Checking that all items were added to the cart")
        self.assertEqual(c.calculate_total(), 5.93, "Checking calculate_total method")

myTests().main()
</textarea>
</div>
</div>
</section>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>



  
    {% if settings.google_ga %}
      <script>
        var _gaq = _gaq || [];
        _gaq.push(['_setAccount', '{{ settings.google_ga }}']);
        _gaq.push(['_trackPageview']);

        (function() {
          var ga = document.createElement('script'); ga.async = true;
          ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
          var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
        })();
      </script>
    {% endif %}
  

  
    <script type="application/json" class="js-hypothesis-config">
    {
      "showHighlights": false
    }
    </script>

    <script async src="https://hypothes.is/embed.js"></script>
    <!--provide space to prevent hypothes.is sidebar from obscuring content-->
    <style>@media screen and (max-width:767px) { body { padding-right: 21px; }}</style>

  


<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>