<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.17.1: http://docutils.sourceforge.net/" />

    <title>Multiple Classes Practice &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.cd688d1623c9932e.css?v=630908D8" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.bf73008a1eba9f4e.css?v=7DED055E" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/runtime.439d581f758f2b99.bundle.js?v=2E54442"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.51d9dde4e5162190.bundle.js?v=9CAA2D33"></script>
    <script src="../_static/runestone.42b576791a4ed54b.bundle.js?v=B7293199"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    
    eBookConfig.useRunestoneServices = {{use_services}};
    eBookConfig.host = '';
    eBookConfig.app = eBookConfig.host + '/runestone';
    eBookConfig.course = '{{ course_name }}';
    eBookConfig.basecourse = '{{ base_course }}';
    eBookConfig.isLoggedIn = {{ is_logged_in}};
    eBookConfig.email = '{{ user_email }}';
    eBookConfig.isInstructor = {{ is_instructor }};
    eBookConfig.username = '{{ user_id}}';
    eBookConfig.readings = {{ readings|safe}};
    eBookConfig.activities = {{ activity_info|safe }}
    eBookConfig.downloadsEnabled = {{downloads_enabled}};
    eBookConfig.allow_pairs = {{allow_pairs}}
    eBookConfig.enableCompareMe = {{enable_compare_me}};
    // _`new_server_prefix`: Defined by the server to be the path to the new server.
    eBookConfig.new_server_prefix = '{{ new_server_prefix }}';
    
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->

    
        {% if serve_ad is defined and serve_ad and settings.adsenseid %}
         <script data-ad-client="{{settings.adsenseid}}" async src="https://pagead2.googlesyndication.com/pagead/js/adsbygoogle.js"></script>
        {% endif %}
    



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
            
            {{ course_name }}
            
          
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="browsing_warning"></span></li>

        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li id="inst_peer_link"><a href='/runestone/peer/instructor.html'>Peer Instruction (Instructor)</a></li>            
            <li><a href='/runestone/peer/student.html'>Peer Instruction (Student)</a></li>            
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/ns/books/published/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=MultipleClassesPractice_orig'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
            <script></script>
<li class="dropdown globaltoc-container">
    <a
        href="../index.html"
        class="dropdown-toggle"
        data-toggle="dropdown"
        aria-label="index"
        >This Chapter<b class="caret"></b
    ></a>
    <ul class="dropdown-menu globaltoc">
         
        {% for row in subchapter_list %}
        <li class="toctree-l1">
            <a class="reference internal" href="{{ row['subchap_uri']}}"
                >{{ row['title']}}</a
            >
        </li>
        {% endfor %} 
        
    </ul>
</li>
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->

    
        {% if settings.num_banners > 0 and settings.show_rs_banner %}
         <div style="width: 100%;">
            <a href="/runestone/default/donate?banner={{banner_num}}">
            <img class="runestone-banner"
                 src="/runestone/static/images/RunestoneBanner{{banner_num}}.png"
                 alt="Please Support Runestone"
                 style="display: block; width: 728px; margin-left: auto; margin-right: auto; margin-top: 10px;">
            </a>
         </div>
        {% endif %}
    


  
  <section id="multiple-classes-practice">
<span id="index-0"></span><h1>Multiple Classes Practice<a class="headerlink" href="#multiple-classes-practice" title="Permalink to this headline">¶</a></h1>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=mult_class_track_and_album_len_ac data-question_label="1">
<div id=mult_class_track_and_album_len_ac_question class="ac_question col-md-12">
<p>Given a <code class="docutils literal notranslate"><span class="pre">Song</span></code> class that has object attributes of <code class="docutils literal notranslate"><span class="pre">name</span></code> and <code class="docutils literal notranslate"><span class="pre">length</span></code> as shown below, finish a <code class="docutils literal notranslate"><span class="pre">Album</span></code> class with an <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method that takes a name and list of <code class="docutils literal notranslate"><span class="pre">Song</span></code> objects.  Then write a <code class="docutils literal notranslate"><span class="pre">length</span></code> method in the <code class="docutils literal notranslate"><span class="pre">Album</span></code> class that returns the total length for the album based on the length of the songs in the album.  Given the objects created below it should print 242.</p>

</div>
<textarea data-lang="python" id="mult_class_track_and_album_len_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Song:

    def __init__(self, name, length):
        self.name = name
        self.length = length

    def __str__(self):
        return f"{self.name}: {self.length}"

class Album:

s1 = Song("Something", 120)
s2 = Song("Here Comes the Sun", 122)
a1 = Album("Abbey Road", [s1, s2])
print(a1.length())

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        s1 = Song("Something", 120)
        s2 = Song("Here Comes the Sun", 122)
        a1 = Album("Abbey Road", [s1, s2])
        self.assertEqual(a1.length(), 242)

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="mult_class_track_and_album_len_pp" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Given a <code class="docutils literal notranslate"><span class="pre">Song</span></code> class that has object attributes of <code class="docutils literal notranslate"><span class="pre">name</span></code> and <code class="docutils literal notranslate"><span class="pre">length</span></code>, create an <code class="docutils literal notranslate"><span class="pre">Album</span></code> class with an <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method that takes a name and list of <code class="docutils literal notranslate"><span class="pre">Song</span></code> objects.  Then create a <code class="docutils literal notranslate"><span class="pre">length</span></code> method in the <code class="docutils literal notranslate"><span class="pre">Album</span></code> class that returns the total length for the album based on the length of the songs in the album.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="2"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        class Album
---
     def __init__(self, name, songs):
---
         self.name = name
         self.songs = songs
---
     def length(self):
---
     def length(): #paired
---
         total = 0
---
         for song in self.songs:
---
         for song in songs: #paired
---
             total += song.length
---
             total += song #paired
---
         return total
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=mult_class_point_and_line_practice_ac data-question_label="3">
<div id=mult_class_point_and_line_practice_ac_question class="ac_question col-md-12">
<p>Given a <code class="docutils literal notranslate"><span class="pre">Point</span></code> class below with a <code class="docutils literal notranslate"><span class="pre">distance</span></code> object method that takes another point and returns the distance between the two points, create a <code class="docutils literal notranslate"><span class="pre">Line</span></code> class with an <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method that takes two points.  Also write the <code class="docutils literal notranslate"><span class="pre">length</span></code> method in the <code class="docutils literal notranslate"><span class="pre">Line</span></code> class to return the length of the line (hint: use the <code class="docutils literal notranslate"><span class="pre">distance</span></code> method in the <code class="docutils literal notranslate"><span class="pre">Point</span></code> class).</p>

</div>
<textarea data-lang="python" id="mult_class_point_and_line_practice_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
import math
class Point:

    def __init__(self, initX, initY):
        """ Create a new point at the given coordinates. """
        self.x = initX
        self.y = initY

    def __str__(self):
        return f"x = {self.x}, y = {self.y}"

    def distance(self, other):
        mx = (self.x - other.x) ** 2
        my = (self.y - other.y) ** 2
        return math.sqrt(mx + my)

class Line:

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        point1 = Point(0,0)
        point2 = Point(10,5)
        point3 = Point(5,10)
        point4 = Point(20,10)
        l1 = Line(point1, point2)
        l2 = Line(point3, point4)

        self.assertAlmostEqual(l1.length(), 11.18, 2, "l1.length")
        self.assertAlmostEqual(l2.length(), 15.0, 1, "l2.length()")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="mult_class_point_and_line_practice_pp" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Given a <code class="docutils literal notranslate"><span class="pre">Point</span></code> class that has a <code class="docutils literal notranslate"><span class="pre">distance</span></code> object method that takes another point and returns the distance between the two points, create a <code class="docutils literal notranslate"><span class="pre">Line</span></code> class with an <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method that takes two points.  Also write the <code class="docutils literal notranslate"><span class="pre">length</span></code> method in the <code class="docutils literal notranslate"><span class="pre">Line</span></code> class to return the length of the line (hint: use the <code class="docutils literal notranslate"><span class="pre">distance</span></code> method in the <code class="docutils literal notranslate"><span class="pre">Point</span></code> class).</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="4"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        class Line:
---
    def __init__(self, p1, p2):
---
    def __init__(p1, p2): #paired
---
        self.p1 = p1
        self.p2 = p2
---
    def length(self):
---
    def length(): #paired
---
        return self.p1.distance(self.p2)
---
        return p1.distance(p2) #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=mult_class_point_and_triangle_practice_ac data-question_label="5">
<div id=mult_class_point_and_triangle_practice_ac_question class="ac_question col-md-12">
<p>Given a <code class="docutils literal notranslate"><span class="pre">Point</span></code> class below with a <code class="docutils literal notranslate"><span class="pre">distance</span></code> object method that takes another point and returns the distance between the two points, create a <code class="docutils literal notranslate"><span class="pre">Triangle</span></code> class with an <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method that takes three points.  Also write the <code class="docutils literal notranslate"><span class="pre">perimeter</span></code> method in the <code class="docutils literal notranslate"><span class="pre">Triangle</span></code> class to return the sum of the lengths of the sides of the triangle. Use the <code class="docutils literal notranslate"><span class="pre">distance</span></code> method in the <code class="docutils literal notranslate"><span class="pre">Point</span></code> class to calculate the distance between two points.</p>

</div>
<textarea data-lang="python" id="mult_class_point_and_triangle_practice_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
import math
class Point:

    def __init__(self, initX, initY):
        """ Create a new point at the given coordinates. """
        self.x = initX
        self.y = initY

    def distanceFromOrigin(self):
        return ((self.x ** 2) + (self.y ** 2)) ** 0.5

    def __str__(self):
        return f"x = {self.x}, y = {self.y}"

    def halfway(self, target):
        mx = (self.x + target.x) / 2
        my = (self.y + target.y) / 2
        return Point(mx, my)

    def distance(self, other):
        mx = (self.x - other.x) ** 2
        my = (self.y - other.y) ** 2
        return math.sqrt(mx + my)

class Triangle:

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        point1 = Point(0,0)
        point2 = Point(10,5)
        point3 = Point(5,10)
        point4 = Point(20,10)
        t1 = Triangle(point1, point2, point3)
        t2 = Triangle(point1, point4, point3)

        self.assertAlmostEqual(t1.perimeter(), 29.43, 2, "t1.perimeter()")
        self.assertAlmostEqual(t2.perimeter(), 48.54, 2, "t2.perimeter()")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="mult_class_point_and_triangle_practice_pp" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Given a <code class="docutils literal notranslate"><span class="pre">Point</span></code> class that has a <code class="docutils literal notranslate"><span class="pre">distance</span></code> object method that takes another point and returns the distance between the two points, create a <code class="docutils literal notranslate"><span class="pre">Triangle</span></code> class with an <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method that takes three points.  Also write the <code class="docutils literal notranslate"><span class="pre">perimeter</span></code> method in the <code class="docutils literal notranslate"><span class="pre">Triangle</span></code> class to return the sum of the lengths of the sides of the triangle. Use the <code class="docutils literal notranslate"><span class="pre">distance</span></code> method in the <code class="docutils literal notranslate"><span class="pre">Point</span></code> class to calculate the distance between two points.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="6"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        class Triangle:
---
     def __init__(self, p1, p2, p3):
---
     def __init__(p1, p2, p3): #paired
---
         self.p1 = p1
         self.p2 = p2
         self.p3 = p3
---
     def perimeter(self):
---
     def perimeter(): #paired
---
         total = self.p1.distance(self.p2)
---
         total = distance(self.p1, self.p2) #paired
---
         total += self.p2.distance(self.p3)
         total += self.p3.distance(self.p1)
---
         total += distance(self.p2, self.p3) #paired
         total += distance(self.p3, self.p1)
---
         return total
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=mult_class_item_order_line_order_practice_ac data-question_label="7">
<div id=mult_class_item_order_line_order_practice_ac_question class="ac_question col-md-12">
<p>Given the <code class="docutils literal notranslate"><span class="pre">Item</span></code> and <code class="docutils literal notranslate"><span class="pre">OrderLine</span></code> classes shown below, write an <code class="docutils literal notranslate"><span class="pre">Order</span></code> class with an <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method that takes a list of <code class="docutils literal notranslate"><span class="pre">OrderLine</span></code> objects and a <code class="docutils literal notranslate"><span class="pre">total</span></code> method that returns the total for the order.</p>

</div>
<textarea data-lang="python" id="mult_class_item_order_line_order_practice_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Item:

    def __init__(self, name, price):
        self.name = name
        self.price = price

    def __str__(self):
        return f"{self.name}: {self.price}"

class OrderLine:

    def __init__(self, item, quantity):
        self.item = item
        self.quantity = quantity

    def total(self):
        return self.item.price * self.quantity

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        i1 = Item("Burger", 6.99)
        i2 = Item("Fries", 2.99)
        i3 = Item("Coke", 1.99)
        ol1 = OrderLine(i1,2)
        ol2 = OrderLine(i2,3)
        ol3 = OrderLine(i3,1)
        o1 = Order([ol1, ol2, ol3])
        self.assertAlmostEqual(o1.total(), 24.94, 2, "o1.total()")


myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="mult_class_item_order_line_order_practice_pp" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Given an <code class="docutils literal notranslate"><span class="pre">Item</span></code> class with object attributes of <code class="docutils literal notranslate"><span class="pre">name</span></code> and <code class="docutils literal notranslate"><span class="pre">price</span></code> and an <code class="docutils literal notranslate"><span class="pre">OrderLine</span></code> class with an object method <code class="docutils literal notranslate"><span class="pre">total</span></code> that returns the total for the <code class="docutils literal notranslate"><span class="pre">OrderLine</span></code> object (the item’s price times the quantity), create an <code class="docutils literal notranslate"><span class="pre">Order</span></code> class with an <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method that takes a list of <code class="docutils literal notranslate"><span class="pre">OrderLine</span></code> objects and a <code class="docutils literal notranslate"><span class="pre">total</span></code> method that returns the total for the order.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="8"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        class Order:
---
     def __init__(self, orderLines):
---
         self.orderLines = orderLines
---
         orderLines = self.orderLines #paired
---
     def total(self):
---
     def total(): #paired
---
         sum = 0
---
         for orderLine in self.orderLines:
---
         for self.orderLine in self.orderLines: #paired
---
             sum += orderLine.total()
---
             sum += orderLine.total #paired
---
         return sum
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=mult_class_card_and_deck_practice_ac data-question_label="9">
<div id=mult_class_card_and_deck_practice_ac_question class="ac_question col-md-12">
<p>Write a <code class="docutils literal notranslate"><span class="pre">shuffle</span></code> method in the <code class="docutils literal notranslate"><span class="pre">Deck</span></code> class that splits the list of cards in half and creates a new list of cards with a card from the first list followed by a card from the second list.  For example, if the list of cards was {a, b, c, d, e, f} it would reset the list to {a, d, b, e, c, f}.  If the list of cards was {a, b, c, d, e} it would reset the list to {a, c, b, d, e}.</p>

</div>
<textarea data-lang="python" id="mult_class_card_and_deck_practice_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Card:

    suit_names = ["Clubs", "Diamonds", "Hearts", "Spades"]
    rank_names = ["Ace", "2", "3", "4", "5", "6", "7", "8", "9",  "10", "Jack", "Queen", "King"]

    def __init__(self, suit=0, rank = 0):
        self.suit = suit
        self.rank = rank

    def __str__(self):
        return Card.rank_names[self.rank] + " of " + Card.suit_names[self.suit]

class Deck:

    def __init__(self):
        self.card_list = []
        for i in range(len(Card.suit_names)):
            for j in range(len(Card.rank_names)):
                self.card_list.append(Card(i, j))

    def __str__(self):
        output = ""
        for card in self.card_list:
            output += card.__str__() + ", "
        return output

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        d = Deck()
        d.shuffle()
        self.assertEqual(d.card_list[0].__str__(), "Ace of Clubs","testing d.card_list[0]")
        self.assertEqual(d.card_list[1].__str__(), "Ace of Hearts","testing d.card_list[1]")
        self.assertEqual(d.card_list[2].__str__(), "2 of Clubs","testing d.card_list[2]")
        self.assertEqual(d.card_list[3].__str__(), "2 of Hearts","testing d.card_list[3]")
        self.assertEqual(len(d.card_list), 52, "testing number of cards")
        d = Deck()
        d.card_list = d.card_list[0:5]
        d.shuffle()
        self.assertEqual(d.card_list[0].__str__(), "Ace of Clubs","testing d.card_list[0]")
        self.assertEqual(d.card_list[1].__str__(), "3 of Clubs","testing d.card_list[1]")
        self.assertEqual(d.card_list[2].__str__(), "2 of Clubs","testing d.card_list[2]")
        self.assertEqual(d.card_list[3].__str__(), "4 of Clubs","testing d.card_list[3]")
        self.assertEqual(d.card_list[4].__str__(), "5 of Clubs","testing d.card_list[4]")
        self.assertEqual(len(d.card_list), 5, "testing number of cards")


myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="mult_class_card_and_deck_practice_pp" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a <code class="docutils literal notranslate"><span class="pre">shuffle</span></code> method for the <code class="docutils literal notranslate"><span class="pre">Deck</span></code> class that splits the list of cards in half and creates a new list of cards with a card from the first list followed by a card from the second list.  For example, if the list of cards was {a, b, c, d, e, f} it would reset the list to {a, d, b, e, c, f}.  If the list of cards was {a, b, c, d, e} it would reset the list to {a, c, b, d, e}.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="10"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def shuffle(self):
---
def shuffle(): #paired
---
      mid = len(self.card_list) // 2
---
      mid = len(self.card_list) / 2 #paired
---
      l1 = self.card_list[:mid]
      l2 = self.card_list[mid:]
      l3 = []
---
      for i in range(len(l1)):
---
      for i in range(l1): #paired
---
          l3.append(l1[i])
---
          l3.append(l2[i])
---
      if len(l2) > len(l1):
---
      if len(l1) > len(l2): #paired
---
          l3.append(l2[-1])
---
      self.card_list = l3
---
      self.card_list = self.l3 #paired
        </pre>
        </div>
        </div>
    </section>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>



  
    {% if settings.google_ga %}
      <script>
        var _gaq = _gaq || [];
        _gaq.push(['_setAccount', '{{ settings.google_ga }}']);
        _gaq.push(['_trackPageview']);

        (function() {
          var ga = document.createElement('script'); ga.async = true;
          ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
          var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
        })();
      </script>
    {% endif %}
  

  
    <script type="application/json" class="js-hypothesis-config">
    {
      "showHighlights": false
    }
    </script>

    <script async src="https://hypothes.is/embed.js"></script>
    <!--provide space to prevent hypothes.is sidebar from obscuring content-->
    <style>@media screen and (max-width:767px) { body { padding-right: 21px; }}</style>

  


<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>