<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Mixed-Up Code Questions Original &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/637.css?v=bddd2011b20db76a20ea" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.css?v=7c5388ad4389eae1a318" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/runtime.bundle.js?v=aa7c717c38c5405af9e6"></script>
    <script src="../_static/637.bundle.js?v=1123ad75f9b821cdd636"></script>
    <script src="../_static/runestone.bundle.js?v=9c8e466e8cbed4a3af5a"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'py4e-int';
    eBookConfig.basecourse = 'py4e-int';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          py4e-int
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=mixedupcode_orig'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
  
<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="mixed-up-code-questions-original">
<h1>Mixed-Up Code Questions Original<a class="headerlink" href="#mixed-up-code-questions-original" title="Permalink to this headline">¶</a></h1>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="mixedupcode_network_1" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">decode</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">link</span></code> and returns a string that contains
the contents of the <code class="docutils literal notranslate"><span class="pre">link</span></code> using urllib. Between each word, there should be a space. Also, a space at the end is okay.
For example, <code class="docutils literal notranslate"><span class="pre">decode('http://data.pr4e.org/romeo.txt')</span></code> should return <code class="docutils literal notranslate"><span class="pre">'But</span> <span class="pre">soft</span> <span class="pre">what</span> <span class="pre">light</span> <span class="pre">through</span> <span class="pre">yonder</span> <span class="pre">window</span> <span class="pre">breaks</span> <span class="pre">It</span> <span class="pre">is</span> <span class="pre">the</span> <span class="pre">east</span> <span class="pre">and</span> <span class="pre">Juliet</span> <span class="pre">is</span> <span class="pre">the</span> <span class="pre">sun</span> <span class="pre">Arise</span> <span class="pre">fair</span> <span class="pre">sun</span> <span class="pre">and</span> <span class="pre">kill</span> <span class="pre">the</span> <span class="pre">envious</span> <span class="pre">moon</span> <span class="pre">Who</span> <span class="pre">is</span> <span class="pre">already</span> <span class="pre">sick</span> <span class="pre">and</span> <span class="pre">pale</span> <span class="pre">with</span> <span class="pre">grief</span> <span class="pre">'</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="1"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        import urllib.request
---
def decode(link):
---
    fhand = urllib.request.urlopen(link)
---
    string = ''
---
    for line in fhand:
---
        string += line.decode().replace('\n', ' ')
---
        string += line.decode().strip() #paired
---
    return string
---
    return fhand.decode().strip() #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=mixedupcode_network_1_ac data-question_label="2">
<div id=mixedupcode_network_1_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">decode</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">link</span></code> and returns a string that contains
the contents of the <code class="docutils literal notranslate"><span class="pre">link</span></code> using urllib. Between each word, there should be a space. Also, a space at the end is okay.
For example, <code class="docutils literal notranslate"><span class="pre">decode('http://data.pr4e.org/romeo.txt')</span></code> should return <code class="docutils literal notranslate"><span class="pre">'But</span> <span class="pre">soft</span> <span class="pre">what</span> <span class="pre">light</span> <span class="pre">through</span> <span class="pre">yonder</span> <span class="pre">window</span> <span class="pre">breaks</span> <span class="pre">It</span> <span class="pre">is</span> <span class="pre">the</span> <span class="pre">east</span> <span class="pre">and</span> <span class="pre">Juliet</span> <span class="pre">is</span> <span class="pre">the</span> <span class="pre">sun</span> <span class="pre">Arise</span> <span class="pre">fair</span> <span class="pre">sun</span> <span class="pre">and</span> <span class="pre">kill</span> <span class="pre">the</span> <span class="pre">envious</span> <span class="pre">moon</span> <span class="pre">Who</span> <span class="pre">is</span> <span class="pre">already</span> <span class="pre">sick</span> <span class="pre">and</span> <span class="pre">pale</span> <span class="pre">with</span> <span class="pre">grief</span> <span class="pre">'</span></code>.</p>

</div>
<textarea data-lang="python" id="mixedupcode_network_1_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(decode('http://data.pr4e.org/clown.txt'),'the clown ran after the car and the car ran into the tent and the tent fell down on the clown and the car ', "decode('http://data.pr4e.org/clown.txt')")
        self.assertEqual(decode('http://data.pr4e.org/romeo.txt'),'But soft what light through yonder window breaks It is the east and Juliet is the sun Arise fair sun and kill the envious moon Who is already sick and pale with grief ', "decode('http://data.pr4e.org/romeo.txt')")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="mixedupcode_network_2" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">write_txt</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">link</span></code>, retrieves a file from the <code class="docutils literal notranslate"><span class="pre">link</span></code> using urllib,
and writes the contents of the <code class="docutils literal notranslate"><span class="pre">link</span></code> to a file called <code class="docutils literal notranslate"><span class="pre">'clown.txt'</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">write_txt('http://data.pr4e.org/clown.txt')</span></code> should create
a file called <code class="docutils literal notranslate"><span class="pre">'clown.txt'</span></code> that has the string <code class="docutils literal notranslate"><span class="pre">'the</span> <span class="pre">clown</span> <span class="pre">ran</span> <span class="pre">after</span> <span class="pre">the</span> <span class="pre">car</span> <span class="pre">and</span> <span class="pre">the</span> <span class="pre">car</span> <span class="pre">ran</span> <span class="pre">into</span> <span class="pre">the</span> <span class="pre">tent</span> <span class="pre">and</span> <span class="pre">the</span> <span class="pre">tent</span> <span class="pre">fell</span> <span class="pre">down</span> <span class="pre">on</span> <span class="pre">the</span> <span class="pre">clown</span> <span class="pre">and</span> <span class="pre">the</span> <span class="pre">car'</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="3"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        import urllib.request
import os
---
def write_txt(link):
---
    text = urllib.request.urlopen(link).read()
---
    fhand = open(os.path.join(os.path.dirname(__file__), 'clown.txt'), 'wb')
---
    fhand = open(os.path.join(os.path.dirname(__file__), 'clown.txt'), 'w') #paired
---
    fhand.write(text)
---
    fhand.close()
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=mixedupcode_network_2_ac data-question_label="4">
<div id=mixedupcode_network_2_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">write_txt</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">link</span></code>, retrieves a file from the <code class="docutils literal notranslate"><span class="pre">link</span></code> using urllib,
and writes the contents of the <code class="docutils literal notranslate"><span class="pre">link</span></code> to a file called <code class="docutils literal notranslate"><span class="pre">'clown.txt'</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">write_txt('http://data.pr4e.org/clown.txt')</span></code> should create
a file called <code class="docutils literal notranslate"><span class="pre">'clown.txt'</span></code> that has the string <code class="docutils literal notranslate"><span class="pre">'the</span> <span class="pre">clown</span> <span class="pre">ran</span> <span class="pre">after</span> <span class="pre">the</span> <span class="pre">car</span> <span class="pre">and</span> <span class="pre">the</span> <span class="pre">car</span> <span class="pre">ran</span> <span class="pre">into</span> <span class="pre">the</span> <span class="pre">tent</span> <span class="pre">and</span> <span class="pre">the</span> <span class="pre">tent</span> <span class="pre">fell</span> <span class="pre">down</span> <span class="pre">on</span> <span class="pre">the</span> <span class="pre">clown</span> <span class="pre">and</span> <span class="pre">the</span> <span class="pre">car'</span></code>.</p>

</div>
<textarea data-lang="python" id="mixedupcode_network_2_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        fhand = open(os.path.join(os.path.dirname(__file__), 'clown.txt'), 'r')
        text = fhand.read().strip()
        self.assertEqual(text, 'the clown ran after the car and the car ran into the tent and the tent fell down on the clown and the car', "write_txt('http://data.pr4e.org/clown.txt')")
        fhand.close()

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="mixedupcode_network_3" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">count_words</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">link</span></code>, retrieves a file from the <code class="docutils literal notranslate"><span class="pre">link</span></code> using urllib, and
returns a dictionary with words as keys and the number of times they appear in the <code class="docutils literal notranslate"><span class="pre">link</span></code> as values. For example, <code class="docutils literal notranslate"><span class="pre">count_words('http://data.pr4e.org/romeo.txt')</span></code>
should return <code class="docutils literal notranslate"><span class="pre">{'But':</span> <span class="pre">1,</span> <span class="pre">'soft':</span> <span class="pre">1,</span> <span class="pre">'what':</span> <span class="pre">1,</span> <span class="pre">'light':</span> <span class="pre">1,</span> <span class="pre">'through':</span> <span class="pre">1,</span> <span class="pre">'yonder':</span> <span class="pre">1,</span> <span class="pre">'window':</span> <span class="pre">1,</span> <span class="pre">'breaks':</span> <span class="pre">1,</span> <span class="pre">'It':</span> <span class="pre">1,</span> <span class="pre">'is':</span> <span class="pre">3,</span> <span class="pre">'the':</span> <span class="pre">3,</span> <span class="pre">'east':</span> <span class="pre">1,</span> <span class="pre">'and':</span> <span class="pre">3,</span> <span class="pre">'Juliet':</span> <span class="pre">1,</span> <span class="pre">'sun':</span> <span class="pre">2,</span> <span class="pre">'Arise':</span> <span class="pre">1,</span> <span class="pre">'fair':</span> <span class="pre">1,</span> <span class="pre">'kill':</span> <span class="pre">1,</span> <span class="pre">'envious':</span> <span class="pre">1,</span> <span class="pre">'moon':</span> <span class="pre">1,</span> <span class="pre">'Who':</span> <span class="pre">1,</span> <span class="pre">'already':</span> <span class="pre">1,</span> <span class="pre">'sick':</span> <span class="pre">1,</span> <span class="pre">'pale':</span> <span class="pre">1,</span> <span class="pre">'with':</span> <span class="pre">1,</span> <span class="pre">'grief':</span> <span class="pre">1}</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="5"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        import urllib.request
---
def count_words(link):
---
    fhand = urllib.request.urlopen(link)
---
    fhand = urllib.request.urlopen(link).read() #paired
---
    counts = {}
---
    for line in fhand:
---
        words = line.decode().split()
---
        words = line.decode() #paired
---
        words = line.split() #paired
---
        for word in words:
---
            counts[word] = counts.get(word, 0) + 1
---
    return counts
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=mixedupcode_network_3_ac data-question_label="6">
<div id=mixedupcode_network_3_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">count_words</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">link</span></code>, retrieves a file from the <code class="docutils literal notranslate"><span class="pre">link</span></code> using urllib, and
returns a dictionary with words as keys and the number of times they appear in the <code class="docutils literal notranslate"><span class="pre">link</span></code> as values. For example, <code class="docutils literal notranslate"><span class="pre">count_words('http://data.pr4e.org/romeo.txt')</span></code>
should return <code class="docutils literal notranslate"><span class="pre">{'But':</span> <span class="pre">1,</span> <span class="pre">'soft':</span> <span class="pre">1,</span> <span class="pre">'what':</span> <span class="pre">1,</span> <span class="pre">'light':</span> <span class="pre">1,</span> <span class="pre">'through':</span> <span class="pre">1,</span> <span class="pre">'yonder':</span> <span class="pre">1,</span> <span class="pre">'window':</span> <span class="pre">1,</span> <span class="pre">'breaks':</span> <span class="pre">1,</span> <span class="pre">'It':</span> <span class="pre">1,</span> <span class="pre">'is':</span> <span class="pre">3,</span> <span class="pre">'the':</span> <span class="pre">3,</span> <span class="pre">'east':</span> <span class="pre">1,</span> <span class="pre">'and':</span> <span class="pre">3,</span> <span class="pre">'Juliet':</span> <span class="pre">1,</span> <span class="pre">'sun':</span> <span class="pre">2,</span> <span class="pre">'Arise':</span> <span class="pre">1,</span> <span class="pre">'fair':</span> <span class="pre">1,</span> <span class="pre">'kill':</span> <span class="pre">1,</span> <span class="pre">'envious':</span> <span class="pre">1,</span> <span class="pre">'moon':</span> <span class="pre">1,</span> <span class="pre">'Who':</span> <span class="pre">1,</span> <span class="pre">'already':</span> <span class="pre">1,</span> <span class="pre">'sick':</span> <span class="pre">1,</span> <span class="pre">'pale':</span> <span class="pre">1,</span> <span class="pre">'with':</span> <span class="pre">1,</span> <span class="pre">'grief':</span> <span class="pre">1}</span></code>.</p>

</div>
<textarea data-lang="python" id="mixedupcode_network_3_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(count_words('http://data.pr4e.org/romeo.txt'),{'But': 1, 'soft': 1, 'what': 1, 'light': 1, 'through': 1, 'yonder': 1, 'window': 1, 'breaks': 1, 'It': 1, 'is': 3, 'the': 3, 'east': 1, 'and': 3, 'Juliet': 1, 'sun': 2, 'Arise': 1, 'fair': 1, 'kill': 1, 'envious': 1, 'moon': 1, 'Who': 1, 'already': 1, 'sick': 1, 'pale': 1, 'with': 1, 'grief': 1}, "count_words('http://data.pr4e.org/romeo.txt')")
        self.assertEqual(count_words('http://data.pr4e.org/clown.txt'),{'the': 7, 'clown': 2, 'ran': 2, 'after': 1, 'car': 3, 'and': 3, 'into': 1, 'tent': 2, 'fell': 1, 'down': 1, 'on': 1},"count_words('http://data.pr4e.org/clown.txt')")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="mixedupcode_network_4" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">write_jpg</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">img_link</span></code>, retrieves a file from the <code class="docutils literal notranslate"><span class="pre">img_link</span></code> using urllib,
and writes the contents of the <code class="docutils literal notranslate"><span class="pre">img_link</span></code> to a file called <code class="docutils literal notranslate"><span class="pre">'cover.jpg'</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">write_jpg('http://data.pr4e.org/cover3.jpg')</span></code> should create
a file called <code class="docutils literal notranslate"><span class="pre">'cover.jpg'</span></code> that has an image of the cover for <code class="docutils literal notranslate"><span class="pre">&quot;PYTHON</span> <span class="pre">FOR</span> <span class="pre">EVERYBODY&quot;</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="7"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        import urllib.request
import os
---
def write_jpg(img_link):
---
    img = urllib.request.urlopen(img_link).read()
---
    fhand = open(os.path.join(os.path.dirname(__file__), 'cover.jpg'), 'wb')
---
    fhand = open(os.path.join(os.path.dirname(__file__), 'cover.jpg'), 'w') #paired
---
    fhand = open(os.path.join(os.path.dirname(__file__), 'cover.jpg'), 'r') #paired
---
    fhand.write(img)
---
    fhand.close()
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=mixedupcode_network_4_ac data-question_label="8">
<div id=mixedupcode_network_4_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">write_jpg</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">img_link</span></code>, retrieves a file from the <code class="docutils literal notranslate"><span class="pre">img_link</span></code> using urllib,
and writes the contents of the <code class="docutils literal notranslate"><span class="pre">img_link</span></code> to a file called <code class="docutils literal notranslate"><span class="pre">'cover.jpg'</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">write_jpg('http://data.pr4e.org/cover3.jpg')</span></code> should create
a file called <code class="docutils literal notranslate"><span class="pre">'cover.jpg'</span></code> that has an image of the cover for <code class="docutils literal notranslate"><span class="pre">&quot;PYTHON</span> <span class="pre">FOR</span> <span class="pre">EVERYBODY&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="mixedupcode_network_4_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="mixedupcode_network_5" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">num_chars</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">link</span></code>, retrieves a file from the <code class="docutils literal notranslate"><span class="pre">link</span></code> using urllib, and
returns the number of characters in <code class="docutils literal notranslate"><span class="pre">link</span></code> in the format <code class="docutils literal notranslate"><span class="pre">&quot;(num)</span> <span class="pre">characters&quot;</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">num_chars('http://data.pr4e.org/romeo-full.txt')</span></code>
should return <code class="docutils literal notranslate"><span class="pre">&quot;8864</span> <span class="pre">characters&quot;</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="9"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        import urllib.request
---
def num_chars(link):
---
    txt = urllib.request.urlopen(link)
---
    num = len(txt.read())
---
    num = len(txt.decode()) #paired
---
    num = len(txt) #paired
---
    return str(num) + ' characters'
---
    return num + ' characters' #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=mixedupcode_network_5_ac data-question_label="10">
<div id=mixedupcode_network_5_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">num_chars</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">link</span></code>, retrieves a file from the <code class="docutils literal notranslate"><span class="pre">link</span></code> using urllib, and
returns the number of characters in <code class="docutils literal notranslate"><span class="pre">link</span></code> in the format <code class="docutils literal notranslate"><span class="pre">&quot;(num)</span> <span class="pre">characters&quot;</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">num_chars('http://data.pr4e.org/romeo-full.txt')</span></code>
should return <code class="docutils literal notranslate"><span class="pre">&quot;8864</span> <span class="pre">characters&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="mixedupcode_network_5_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(num_chars('http://data.pr4e.org/romeo-full.txt'),"8864 characters", "num_chars('http://data.pr4e.org/romeo-full.txt')")
        self.assertEqual(num_chars('http://data.pr4e.org/clown.txt'),"106 characters","num_chars('http://data.pr4e.org/clown.txt')")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="mixedupcode_network_6" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">contents</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">link</span></code>, retrieves a file from the <code class="docutils literal notranslate"><span class="pre">link</span></code> using sockets, and returns as a string
the contents of the <code class="docutils literal notranslate"><span class="pre">link</span></code> (specifically 10000 characters). For example, <code class="docutils literal notranslate"><span class="pre">contents('http://data.pr4e.org/clown.txt')</span></code> should return
<code class="docutils literal notranslate"><span class="pre">&quot;HTTP/1.1</span> <span class="pre">200</span> <span class="pre">OK\nDate:</span> <span class="pre">Thu,</span> <span class="pre">12</span> <span class="pre">Aug</span> <span class="pre">2021</span> <span class="pre">01:24:15</span> <span class="pre">GMT\nServer:</span> <span class="pre">Apache/2.4.18</span> <span class="pre">(Ubuntu)\nLast-Modified:</span> <span class="pre">Sat,</span> <span class="pre">13</span> <span class="pre">May</span> <span class="pre">2017</span> <span class="pre">11:22:22</span> <span class="pre">GMT\nETag:</span> <span class="pre">'6a-54f6609240717'\nAccept-Ranges:</span> <span class="pre">bytes\nContent-Length:</span> <span class="pre">106\nCache-Control:</span> <span class="pre">max-age=0,</span> <span class="pre">no-cache,</span> <span class="pre">no-store,</span> <span class="pre">must-revalidate\nPragma:</span> <span class="pre">no-cache\nExpires:</span> <span class="pre">Wed,</span> <span class="pre">11</span> <span class="pre">Jan</span> <span class="pre">1984</span> <span class="pre">05:00:00</span> <span class="pre">GMT\nConnection:</span> <span class="pre">close\nContent-Type:</span> <span class="pre">text/plain\n\nthe</span> <span class="pre">clown</span> <span class="pre">ran</span> <span class="pre">after</span> <span class="pre">the</span> <span class="pre">car</span> <span class="pre">and</span> <span class="pre">the</span> <span class="pre">car</span> <span class="pre">ran</span> <span class="pre">into</span> <span class="pre">the</span> <span class="pre">tent</span> <span class="pre">and</span> <span class="pre">the</span> <span class="pre">tent</span> <span class="pre">fell</span> <span class="pre">down</span> <span class="pre">on</span> <span class="pre">the</span> <span class="pre">clown</span> <span class="pre">and</span> <span class="pre">the</span> <span class="pre">car\n&quot;</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="11"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        import socket
---
def contents(link):
---
    mysock = socket.socket(socket.AF_INET, socket.SOCK_STREAM)
---
    mysock.connect(('data.pr4e.org', 80))
---
    cmd = 'GET ' + link + ' HTTP/1.0\r\n\r\n'
---
    cmd = 'GET ' + link + ' HTTP/1.0\r\n\r\n'.encode() #paired
---
    cmd = cmd.encode()
---
    mysock.send(cmd)
---
    data = mysock.recv(10000)
---
    data = mysock.recv() #paired
---
    return data.decode()
---
    return data #paired
---
    mysock.close()
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=mixedupcode_network_6_ac data-question_label="12">
<div id=mixedupcode_network_6_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">contents</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">link</span></code>, retrieves a file from the <code class="docutils literal notranslate"><span class="pre">link</span></code> using sockets, and returns as a string
the contents of the <code class="docutils literal notranslate"><span class="pre">link</span></code> (specifically 10000 characters). For example, <code class="docutils literal notranslate"><span class="pre">contents('http://data.pr4e.org/clown.txt')</span></code> should return
<code class="docutils literal notranslate"><span class="pre">&quot;HTTP/1.1</span> <span class="pre">200</span> <span class="pre">OK\nDate:</span> <span class="pre">Thu,</span> <span class="pre">12</span> <span class="pre">Aug</span> <span class="pre">2021</span> <span class="pre">01:24:15</span> <span class="pre">GMT\nServer:</span> <span class="pre">Apache/2.4.18</span> <span class="pre">(Ubuntu)\nLast-Modified:</span> <span class="pre">Sat,</span> <span class="pre">13</span> <span class="pre">May</span> <span class="pre">2017</span> <span class="pre">11:22:22</span> <span class="pre">GMT\nETag:</span> <span class="pre">'6a-54f6609240717'\nAccept-Ranges:</span> <span class="pre">bytes\nContent-Length:</span> <span class="pre">106\nCache-Control:</span> <span class="pre">max-age=0,</span> <span class="pre">no-cache,</span> <span class="pre">no-store,</span> <span class="pre">must-revalidate\nPragma:</span> <span class="pre">no-cache\nExpires:</span> <span class="pre">Wed,</span> <span class="pre">11</span> <span class="pre">Jan</span> <span class="pre">1984</span> <span class="pre">05:00:00</span> <span class="pre">GMT\nConnection:</span> <span class="pre">close\nContent-Type:</span> <span class="pre">text/plain\n\nthe</span> <span class="pre">clown</span> <span class="pre">ran</span> <span class="pre">after</span> <span class="pre">the</span> <span class="pre">car</span> <span class="pre">and</span> <span class="pre">the</span> <span class="pre">car</span> <span class="pre">ran</span> <span class="pre">into</span> <span class="pre">the</span> <span class="pre">tent</span> <span class="pre">and</span> <span class="pre">the</span> <span class="pre">tent</span> <span class="pre">fell</span> <span class="pre">down</span> <span class="pre">on</span> <span class="pre">the</span> <span class="pre">clown</span> <span class="pre">and</span> <span class="pre">the</span> <span class="pre">car\n&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="mixedupcode_network_6_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(contents('http://data.pr4e.org/romeo.txt'),"HTTP/1.1 200 OK\nDate: Thu, 12 Aug 2021 01:24:14 GMT\nServer: Apache/2.4.18 (Ubuntu)\nLast-Modified: Sat, 13 May 2017 11:22:22 GMT\nETag: 'a7-54f6609245537'\nAccept-Ranges: bytes\nContent-Length: 167\nCache-Control: max-age=0, no-cache, no-store, must-revalidate\nPragma: no-cache\nExpires: Wed, 11 Jan 1984 05:00:00 GMT\nConnection: close\nContent-Type: text/plain\nBut soft what light through yonder window breaks\nIt is the east and Juliet is the sun\nArise fair sun and kill the envious moon\nWho is already sick and pale with grief\n", "contents('http://data.pr4e.org/romeo.txt')")
        self.assertEqual(contents('http://data.pr4e.org/clown.txt'),"HTTP/1.1 200 OK\nDate: Thu, 12 Aug 2021 01:24:15 GMT\nServer: Apache/2.4.18 (Ubuntu)\nLast-Modified: Sat, 13 May 2017 11:22:22 GMT\nETag: '6a-54f6609240717'\nAccept-Ranges: bytes\nContent-Length: 106\nCache-Control: max-age=0, no-cache, no-store, must-revalidate\nPragma: no-cache\nExpires: Wed, 11 Jan 1984 05:00:00 GMT\nConnection: close\nContent-Type: text/plain\n\nthe clown ran after the car and the car ran into the tent and the tent fell down on the clown and the car\n","contents('http://data.pr4e.org/clown.txt')")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="mixedupcode_network_7" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">reg_num_links</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">url</span></code> and returns the number of ‘href’
attributes that start with ‘http’ using regular expressions. Since websites are frequently updated,
the returned number may change as links get added and deleted.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="13"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        import urllib.request
import re
---
def reg_num_links(url):
---
    html = urllib.request.urlopen(url).read()
---
    html = urllib.request.urlopen(url) #paired
---
    links = re.findall(b'href="(http[s]?://.*?)"', html)
---
    links = re.findall(b'href="(http.*)"', html) #paired
---
    return len(links)
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=mixedupcode_network_7_ac data-question_label="14">
<div id=mixedupcode_network_7_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">reg_num_links</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">url</span></code> and returns the number of ‘href’
attributes that start with ‘http’ using regular expressions. Since websites are frequently updated,
the returned number may change as links get added and deleted.</p>

</div>
<textarea data-lang="python" id="mixedupcode_network_7_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        list1 = ['https://www.nytimes.com/', 'https://www.washingtonpost.com/']
        for link in list1:
            html = urllib.request.urlopen(link).read()
            links = re.findall(b'href="(http[s]?://.*?)"', html)
            length = len(links)
            self.assertEqual(reg_num_links(link),length,"reg_num_links(link)")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="mixedupcode_network_8" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">bsoup_num_links</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">url</span></code> and returns the number of ‘href’
attributes that start with ‘http’ using BeautifulSoup. Since websites are frequently updated,
the returned number may change as links get added and deleted.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="15"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        import requests
from bs4 import BeautifulSoup
---
def bsoup_num_links(url):
---
    list1 = []
---
    resp = requests.get(url)
---
    soup = BeautifulSoup(resp.content, 'html.parser')
---
    tags = soup.find_all(href=True)
---
    tags = soup.find_all(href) #paired
---
    tags = soup.find_all('href'=True) #paired
---
    tags = soup.find_all('href') #paired
---
    for tag in tags:
---
        if tag.get('href', None).startswith('http'):
---
        if tag.startswith('http'): #paired
---
            list1.append((tag.get('href', None)))
---
    return len(list1)
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=mixedupcode_network_8_ac data-question_label="16">
<div id=mixedupcode_network_8_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">bsoup_num_links</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">url</span></code> and returns the number of ‘href’
attributes that start with ‘http’ using BeautifulSoup. Since websites are frequently updated,
the returned number may change as links get added and deleted.</p>

</div>
<textarea data-lang="python" id="mixedupcode_network_8_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        list1 = ['https://www.nytimes.com/', 'https://www.washingtonpost.com/']
        for link in list1:
            list2 = []
            resp = requests.get(link)
            soup = BeautifulSoup(resp.content, 'html.parser')
            tags = soup.find_all(href=True)
            for tag in tags:
                if tag.get('href', None).startswith('http'):
                    list2.append((tag.get('href', None)))
            length = len(list2)
            self.assertEqual(bsoup_num_links(link),length,"bsoup_num_links(link)")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="mixedupcode_network_9" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">img_links</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">url</span></code> and returns a list that contains all image links
using BeautifulSoup. Since websites are frequently updated, the returned list of image links may change as image links get added and deleted.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="17"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        import requests
from bs4 import BeautifulSoup
---
def img_links(url):
---
    list1 = []
---
    resp = requests.get(url)
---
    soup = BeautifulSoup(resp.content, 'html.parser')
---
    tags = soup.find_all('img')
---
    tags = soup.findall('img') #paired
---
    for tag in tags:
---
        list1.append(tag.get('src', None))
---
        list1.append(tag) #paired
---
    return list1
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=mixedupcode_network_9_ac data-question_label="18">
<div id=mixedupcode_network_9_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">img_links</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">url</span></code> and returns a list that contains all image links
using BeautifulSoup. Since websites are frequently updated, the returned list of image links may change as image links get added and deleted.</p>

</div>
<textarea data-lang="python" id="mixedupcode_network_9_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        lst = ['https://www.nytimes.com/', 'https://canvas.it.umich.edu/']
        for link in lst:
            list1 = []
            resp = requests.get(link)
            soup = BeautifulSoup(resp.content, 'html.parser')
            tags = soup.find_all('img')
            for tag in tags:
                list1.append(tag.get('src', None))
            self.assertEqual(img_links(link),list1,"img_links(link)")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="mixedupcode_network_10" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">span_attrs</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">url</span></code> and returns a list of dictionaries using BeautifulSoup.
Each dictionary is equivalent to each span tag. The keys of the dictionary are the attributes of the span tag,
and the values of the dictionary are the values of the attributes. Since websites are frequently updated,
the returned list of dictionaries may change as span tags, attributes, and values get added, deleted, or modified.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="19"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        import requests
from bs4 import BeautifulSoup
---
def span_attrs(url):
---
    list1 = []
---
    resp = requests.get(url)
---
    soup = BeautifulSoup(resp.content, "html.parser")
---
    tags = soup.find_all('span')
---
    tags = soup.find_all(span) #paired
---
    for tag in tags:
---
        list1.append(tag.attrs)
---
        list1.append(tag.attributes) #paired
---
        list1.append({tag.attrs: tag.values}) #paired
---
        list1.append({tag.attrs(): tag.values()}) #paired
---
        list1.append({tag.attributes: tag.values}) #paired
---
        list1.append({tag.attributes(): tag.values()}) #paired
---
    return list1
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=mixedupcode_network_10_ac data-question_label="20">
<div id=mixedupcode_network_10_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">span_attrs</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">url</span></code> and returns a list of dictionaries using BeautifulSoup.
Each dictionary is equivalent to each span tag. The keys of the dictionary are the attributes of the span tag,
and the values of the dictionary are the values of the attributes. Since websites are frequently updated,
the returned list of dictionaries may change as span tags, attributes, and values get added, deleted, or modified.</p>

</div>
<textarea data-lang="python" id="mixedupcode_network_10_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        lst = ['https://www.nytimes.com/', 'https://www.w3schools.com/TAGS/default.ASP']
        for link in lst:
            list1 = []
            resp = requests.get(link)
            soup = BeautifulSoup(resp.content, "html.parser")
            tags = soup.find_all('span')
            for tag in tags:
                list1.append(tag.attrs)
            self.assertEqual(span_attrs(link),list1,"span_attrs(link)")

myTests().main()
</textarea>
</div>
</div>
</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>