<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>5.21. Group Work: Functions and Strings &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/637.css?v=bddd2011b20db76a20ea" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.css?v=7c5388ad4389eae1a318" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/runtime.bundle.js?v=aa7c717c38c5405af9e6"></script>
    <script src="../_static/637.bundle.js?v=1123ad75f9b821cdd636"></script>
    <script src="../_static/runestone.bundle.js?v=9c8e466e8cbed4a3af5a"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="5.22. Functions and Strings Multiple Choice Questions" href="function_string_exercises.html" />
    <link rel="prev" title="5.20. Functions Write Code Questions" href="functions_writecode.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'py4e-int';
    eBookConfig.basecourse = 'py4e-int';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          py4e-int
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=func-string'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">5.20. </span>Functions Write Code Questions" data-toggle="tooltip">
    <a href="functions_writecode.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">5.22. </span>Functions and Strings Multiple Choice Questions' data-toggle="tooltip" >              
    <a href="function_string_exercises.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
  
<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="group-work-functions-and-strings">
<h1><span class="section-number">5.21. </span>Group Work: Functions and Strings<a class="headerlink" href="#group-work-functions-and-strings" title="Permalink to this headline">¶</a></h1>
<p>It is best to use a POGIL approach with the following. In POGIL students work
in groups on activities and each member has an assigned role.  For more information see <a class="reference external" href="https://cspogil.org/Home">https://cspogil.org/Home</a>.</p>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>If you work in a group, have only one member of the group fill in the answers on this page.  You will be able to share your answers with the group at the bottom of the page.</p>
</div>
<p><strong>Learning Objectives</strong></p>
<p>Students will know and be able to do the following.</p>
<p><em>Content Objectives:</em></p>
<ul class="simple">
<li><p>Recognize strings.</p></li>
<li><p>Use string indices.</p></li>
<li><p>Understand how slice works for both positive and negative indices.</p></li>
<li><p>Use input and convert between strings and numbers.</p></li>
<li><p>Recognize common string methods.</p></li>
</ul>
<p><em>Process Objectives:</em></p>
<ul class="simple">
<li><p>Predict the output from code.</p></li>
<li><p>Modify code to produce the correct output.</p></li>
</ul>
<p>A <strong>string</strong> is a sequence of characters enclosed in quotes.  In Python you can use pairs of single
or double quotes to enclose a string.  You can even use tripe quotes when
a string covers more than one line.</p>
<div class="section" id="string-indices">
<h2><span class="section-number">5.21.1. </span>String Indices<a class="headerlink" href="#string-indices" title="Permalink to this headline">¶</a></h2>

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.21.1.1" id="funct_fitb_pogil_initials_v2"  style="visibility: hidden;">
            <p>Q-1: What is the last thing that will be printed when the code below runs?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^\\s*Jk\\s*$", "regexFlags": "", "feedback": "<p>It prints the first letter of the first name and last letter of the last name.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>String indices start with 0 and -1 is the index of the last letter in a string.</p>\n"}]]
        </script>

        </div>
        </div>
            
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=funct_ac_pogil_initials_v2 data-question_label="5.21.1.2">
<div id=funct_ac_pogil_initials_v2_question class="ac_question col-md-12">
<p>Run the code below to see what it prints.  Then fix it to pass the given test.  It should return a string with the
first character of the first name and first character of the last name.</p>

</div>
<textarea data-lang="python" id="funct_ac_pogil_initials_v2_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
         data-caption='get_initials'  data-wasm=/_static
     style="visibility: hidden;">
# function definition
def get_initials(first, last):
    return first[0] + last[-1]

# function definition
def main():
    print(type("Hello"))
    print(type('Class'))
    print(type(42))

    print(get_initials("J'Quan",'Alik'))


# function call
main()

from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(get_initials("J'Quan",'Alik'),"JA",'''get_initials("J'Quan",'Alik')''')

myTests().main()
</textarea>
</div>
</div>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>Use [index] to get a character from a string.  The first character in a string is at index 0 and the last is at the length of the string minus 1 (also know as index -1 in Python).  Also note that you can use pairs of single quotes or double quotes to create a string.  This is useful when the string includes a single quote.</p>
</div>
<p>Fix the function <code class="docutils literal notranslate"><span class="pre">get_initials</span></code> above to return a string with the first letter of the first name followed by the first letter of the last name.</p>

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.21.1.3" id="funct_fitb_string_type"  style="visibility: hidden;">
            <p>Q-3: What type of thing is “Hello”, i.e. what class does it belong to?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^\\s*'?str'?\\s*$", "regexFlags": "", "feedback": "<p>A string in Python is an object of the str class.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>What is the class name printed by print(type(\u201cHello\u201d))?</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.21.1.4" id="funct_fitb_int_type"  style="visibility: hidden;">
            <p>Q-4: What type of thing is 42, i.e. what class does it belong to?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^\\s*'?int'?\\s*$", "regexFlags": "", "feedback": "<p>An integer in Python is an object of the int class.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>What is the class name printed by print(type(42))?</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.21.1.5" id="funct_fitb_add_strings"  style="visibility: hidden;">
            <p>Q-5: What is the symbol that is used to append (concatenate) strings together?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^\\s*\\+\\s*$", "regexFlags": "", "feedback": "<p>Use + to append one string after another in Python.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Check the code above and try again.</p>\n"}]]
        </script>

        </div>
        </div>
            </div>
<div class="section" id="string-slices">
<h2><span class="section-number">5.21.2. </span>String Slices<a class="headerlink" href="#string-slices" title="Permalink to this headline">¶</a></h2>

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.21.2.1" id="funct_fitb_pogil_short_name"  style="visibility: hidden;">
            <p>Q-6: What is the last thing that will be printed when the code below runs?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^\\s*Sibs\\s*$", "regexFlags": "", "feedback": "<p>The last thing it prints are the first two letters of the first name and the last two letters of the last name.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>A slice starts with the first number and ends before the second.  If the second is left off it goes to the end of the string. The last character in a string is at index -1.</p>\n"}]]
        </script>

        </div>
        </div>
            
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=funct_ac_pogil_short_name data-question_label="5.21.2.2">
<div id=funct_ac_pogil_short_name_question class="ac_question col-md-12">
<p>Run the code below to see what it prints.</p>

</div>
<textarea data-lang="python" id="funct_ac_pogil_short_name_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
         data-caption='get_short_name'  data-wasm=/_static
     style="visibility: hidden;">
# function definition
def get_short_name(first, last):
    print(len(first))
    print(len(last))
    return first[:2] + last[-2:]

# function definition
def main():
    print(get_short_name('Simona',"Jacobs"))

# function call
main()
</textarea>
</div>
</div>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>Use the slice [start:end] operator to get a slice (substring) from a string. It will return a new string starting at the start and including all the characters up to just before the end (end - 1).</p>
</div>

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.21.2.3" id="funct_fitb_three_char_slice"  style="visibility: hidden;">
            <p>Q-8: Use the slice operator to return the first three characters from the variable <code class="docutils literal notranslate"><span class="pre">dna</span></code>?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^\\s*dna\\[0\\:3\\]|\\[:3\\]\\s*$", "regexFlags": "", "feedback": "<p>This will return a new string with the characters from index 0 to 2.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Look at the note above and try again.</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.21.2.4" id="funct_fitb_last_three_char_slice"  style="visibility: hidden;">
            <p>Q-9: Use the slice operator to return a new string with just the last three characters from the variable named <code class="docutils literal notranslate"><span class="pre">dna</span></code> using a negative index?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^\\s*dna\\[-3\\:\\]\\s*$", "regexFlags": "", "feedback": "<p>This will return a new string with the last three characters in it.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Read about negative indicies in Python and try again.</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.21.2.5" id="funct_fitb_start_slice_def"  style="visibility: hidden;">
            <p>Q-10: What index will the slice (substring) start with if the start index isn’t specified?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [0, 0], "feedback": "<p>A slice starts at index 0 if the start isn\u2019t specified.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Look at the code above and try again.</p>\n"}]]
        </script>

        </div>
        </div>
            
            <div class="runestone alert alert-warning">
            <ul data-component="multiplechoice" data-question_label="5.21.2.6" data-multipleanswers="false"  id="funt_mc_end_slice"  style="visibility: hidden;">
            <p>Q-11: What index is the default end index if it isn’t specified in a slice?</p>
<li data-component="answer"  id="funt_mc_end_slice_opt_a"><p>0</p>
</li><li data-component="feedback" id="funt_mc_end_slice_opt_a">
<p>This is the default start index</p></li><li data-component="answer"  id="funt_mc_end_slice_opt_b"><p>-1 (the index of the last character)</p>
</li><li data-component="feedback" id="funt_mc_end_slice_opt_b">
<p>Since a slice does not include the character at the specified end index this would stop at the second to last character in the string</p></li><li data-component="answer" data-correct id="funt_mc_end_slice_opt_c"><p>the length of the string</p>
</li><li data-component="feedback" id="funt_mc_end_slice_opt_c">
<p>The slice returns a string with all the characters from the start to the end minus one.  That is the same as the length of the string minus one.</p></li>

            </ul>
            </div>
            
        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.21.2.7" id="funct_fitb_len_string"  style="visibility: hidden;">
            <p>Q-12: What built-in function tells you the number of characters in a string?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^\\s*len\\s*$", "regexFlags": "", "feedback": "<p>The len function takes a string and returns the number of characters in it.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Look at the example code above.</p>\n"}]]
        </script>

        </div>
        </div>
            
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=fuct_ac_list_join_strings_withtout_first data-question_label="5.21.2.8">
<div id=fuct_ac_list_join_strings_withtout_first_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">join_no_first</span></code> that takes two strings <code class="docutils literal notranslate"><span class="pre">a</span></code> and <code class="docutils literal notranslate"><span class="pre">b</span></code> and returns a new string with all the characters in string <code class="docutils literal notranslate"><span class="pre">a</span></code> except the first one followed by all the characters in <code class="docutils literal notranslate"><span class="pre">b</span></code> except the first one.  For example, join_no_first(‘hi’, ‘bye’) would return ‘iye’.</p>

</div>
<textarea data-lang="python" id="fuct_ac_list_join_strings_withtout_first_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def join_no_first(a, b):

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(join_no_first('hi', 'bye'), 'iye', "join_no_first('hi', 'bye')")
        self.assertEqual(join_no_first('abc', 'cba'), 'bcba', "join_no_first('abc', 'cba')")
        self.assertEqual(join_no_first('Watch out', 'bye'), 'atch outye', "join_no_first('Watch out', 'bye')")
        self.assertEqual(join_no_first('123', '456'), '2356', "join_no_first('123', '456')")
        self.assertEqual(join_no_first('!Hi', 'Bye!'), 'Hiye!', "join_no_first('!Hi', 'Bye!')")
        self.assertEqual(join_no_first('a', 'x'), '', "join_no_first('a', 'x')")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="input-and-converting-between-strings-and-numbers">
<h2><span class="section-number">5.21.3. </span>Input and Converting Between Strings and Numbers<a class="headerlink" href="#input-and-converting-between-strings-and-numbers" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=funct_ac_pogil_get_name data-question_label="5.21.3.1">
<div id=funct_ac_pogil_get_name_question class="ac_question col-md-12">
<p>Run the code below to see what it prints.</p>

</div>
<textarea data-lang="python" id="funct_ac_pogil_get_name_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
         data-caption='get_initials'  data-wasm=/_static
     style="visibility: hidden;">
from datetime import datetime

# function definition
def get_name():
    first = input("What is your first name?")
    print("Hello " + first)
    age_str = input("What is your age?")
    today = datetime.today()
    age = int(age_str)
    birth_year = today.year - age
    print("You were born in " + str(birth_year) + " or " + str(birth_year - 1))

# function call
get_name()
</textarea>
</div>
</div>

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.21.3.2" id="funct_fitb_convert_string_int"  style="visibility: hidden;">
            <p>Q-15: What is the name of the built-in function that will convert a string to an integer?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^\\s*int\\s*$", "regexFlags": "", "feedback": "<p>This will convert a string to an integer</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Look at the code above and try again.</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.21.3.3" id="funct_fitb_convert_int_string"  style="visibility: hidden;">
            <p>Q-16: What is the name of the built-in function that will convert an integer to a string?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^\\s*str\\s*$", "regexFlags": "", "feedback": "<p>This will convert a number to a string.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Look at the code above and try again.</p>\n"}]]
        </script>

        </div>
        </div>
            </div>
<div class="section" id="string-methods">
<h2><span class="section-number">5.21.4. </span>String Methods<a class="headerlink" href="#string-methods" title="Permalink to this headline">¶</a></h2>
<p>Strings have methods (functions) that operate on a string object using <strong>dot-notation</strong> as shown in the example code below.</p>

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.21.4.1" id="funct_fitb_string_method_pred"  style="visibility: hidden;">
            <p>Q-17: What will be returned from the <code class="docutils literal notranslate"><span class="pre">get_user_name</span></code> function below?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^\\s*malana_coffy\\s*$", "regexFlags": "", "feedback": "<p>This will return the first name in lowercase followed by a _ and then the last name in lowercase.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Remember that string methods don\u2019t change the original string.</p>\n"}]]
        </script>

        </div>
        </div>
            
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=funct_ac_pogil_user_name data-question_label="5.21.4.2">
<div id=funct_ac_pogil_user_name_question class="ac_question col-md-12">
<p>Run the code below to see what it prints.</p>

</div>
<textarea data-lang="python" id="funct_ac_pogil_user_name_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
         data-caption='get_user_name'  data-wasm=/_static
     style="visibility: hidden;">
# function definition
def get_user_name(first, last):
    print(first.lower())
    print(last.upper())
    print(first.find('a'))
    print(first.find('z'))
    print(last.replace("f", "1"))
    print(first.split("a"))
    user = first.lower() + "_" + last.lower()
    return user

# function definition
def main():
    print("run away!".capitalize())
    print(get_user_name("Malana", "Coffy"))

# function call
main()
</textarea>
</div>
</div>

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.21.4.3" id="funct_fitb_pogil_find_return"  style="visibility: hidden;">
            <p>Q-19: What does the function <code class="docutils literal notranslate"><span class="pre">find</span></code> return if the character is not found in the string?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [-1, -1], "feedback": "<p>If the character isn\u2019t in the string find returns -1.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Check the output of find(\u2018z\u2019) above</p>\n"}]]
        </script>

        </div>
        </div>
            <div class="admonition note">
<p class="admonition-title">Note</p>
<p>Strings are immutable, which means that they do not change.  String methods that appear to change a string return a new string.</p>
</div>

<div class="runestone">
<ul data-component="dragndrop" data-question_label="5.21.4.4" id="funct_string_methods_dnd_v2"  style="visibility: hidden;">
    <span data-subcomponent="question">Q-20: Drag each string method to the description of what it does.</span>
	<span data-subcomponent=feedback>Read the chapter on strings and try again</span>

    <li data-subcomponent="draggable" id="funct_string_methods_dnd_v2_drag1">replace</li>
    <li data-subcomponent="dropzone" for="funct_string_methods_dnd_v2_drag1">Returns a new string with all instances of the first specified character replaced with the second specified character.</li>

    <li data-subcomponent="draggable" id="funct_string_methods_dnd_v2_drag2">lower</li>
    <li data-subcomponent="dropzone" for="funct_string_methods_dnd_v2_drag2">Returns a new string with all of the characters in lowercase.</li>

    <li data-subcomponent="draggable" id="funct_string_methods_dnd_v2_drag3">upper</li>
    <li data-subcomponent="dropzone" for="funct_string_methods_dnd_v2_drag3">Returns a new string with all of the characters in uppercase.</li>

    <li data-subcomponent="draggable" id="funct_string_methods_dnd_v2_drag4">split</li>
    <li data-subcomponent="dropzone" for="funct_string_methods_dnd_v2_drag4">Returns a list of strings created by splitting the original string at the specified character.</li>

    <li data-subcomponent="draggable" id="funct_string_methods_dnd_v2_drag5">capitalize</li>
    <li data-subcomponent="dropzone" for="funct_string_methods_dnd_v2_drag5">Returns a new string with the first letter capitalized (uppercase).</li>
</ul></div>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=funct_ac_pogil_strip_methods_v2 data-question_label="5.21.4.5">
<div id=funct_ac_pogil_strip_methods_v2_question class="ac_question col-md-12">
<p>Run the code below to see what it prints.</p>

</div>
<textarea data-lang="python" id="funct_ac_pogil_strip_methods_v2_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
         data-caption='Example with strip, rstrip, lstrip, and find'  data-wasm=/_static
     style="visibility: hidden;">
# function definition
def test(the_string):
    print(len(the_string))
    s1 = the_string.strip()
    print(s1)
    print(len(s1))
    s2 = the_string.rstrip()
    print(s2)
    print(len(s2))
    s3 = the_string.lstrip()
    print(s3)
    print(s3)
    print(the_string.find('H'))

# function definition
def main():
    test(" Help! ")

# function call
main()
</textarea>
</div>
</div>

<div class="runestone">
<ul data-component="dragndrop" data-question_label="5.21.4.6" id="funct_string_methods2_dnd"  style="visibility: hidden;">
    <span data-subcomponent="question">Q-22: Drag each string method to the description of what it does.</span>
	<span data-subcomponent=feedback>Read the chapter on strings and try again</span>

    <li data-subcomponent="draggable" id="funct_string_methods2_dnd_drag1">strip</li>
    <li data-subcomponent="dropzone" for="funct_string_methods2_dnd_drag1">Returns a new string with leading and trailing spaces removed.</li>

    <li data-subcomponent="draggable" id="funct_string_methods2_dnd_drag2">rstrip</li>
    <li data-subcomponent="dropzone" for="funct_string_methods2_dnd_drag2">Returns a new string with trailing spaces removed.</li>

    <li data-subcomponent="draggable" id="funct_string_methods2_dnd_drag3">lstrip</li>
    <li data-subcomponent="dropzone" for="funct_string_methods2_dnd_drag3">Returns a new string with leading spaces removed.</li>

    <li data-subcomponent="draggable" id="funct_string_methods2_dnd_drag4">find</li>
    <li data-subcomponent="dropzone" for="funct_string_methods2_dnd_drag4">Returns the starting index of a target string if found or -1 if not found.</li>
</ul></div><p>For more information on string methods see <a class="reference external" href="https://www.w3schools.com/python/python_ref_string.asp">https://www.w3schools.com/python/python_ref_string.asp</a>.</p>
<p>If you worked in a group, you can copy the answers from this page to the other group members.  Select the group members below and click the button to share the answers.</p>

<div class="runestone alert alert-warning sqcontainer">
    <div data-component="groupsub" id=func_string_groupsub data-size_limit=4>
        <div class="col-sm-6">
            <select id="assignment_group" multiple class="assignment_partner_select" style="width: 100%">
            </select>
        </div>
        <div id="groupsub_button" class="col-sm-6">
        </div>
        <p>The Submit Group button will submit the answer for each each question
        on this page for each member of your group. It also logs you as the official
        group submitter.</p>
    </div>
</div>
</div>
</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 5.20. Functions Write Code Questions' data-toggle="tooltip">
  <a href="functions_writecode.html" aria-label="Previous - 5.20. Functions Write Code Questions">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 5.22. Functions and Strings Multiple Choice Questions' data-toggle="tooltip" >
    <a href="function_string_exercises.html" aria-label="Next - 5.22. Functions and Strings Multiple Choice Questions">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>