<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>5.20. Functions Write Code Questions &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/637.css?v=bddd2011b20db76a20ea" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.css?v=7c5388ad4389eae1a318" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/runtime.bundle.js?v=aa7c717c38c5405af9e6"></script>
    <script src="../_static/637.bundle.js?v=1123ad75f9b821cdd636"></script>
    <script src="../_static/runestone.bundle.js?v=9c8e466e8cbed4a3af5a"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="5.21. Group Work: Functions and Strings" href="func-string.html" />
    <link rel="prev" title="5.19. Functions Mixed-Up Code Questions" href="functions_mixedupcode.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'py4e-int';
    eBookConfig.basecourse = 'py4e-int';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          py4e-int
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=functions_writecode'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">5.19. </span>Functions Mixed-Up Code Questions" data-toggle="tooltip">
    <a href="functions_mixedupcode.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">5.21. </span>Group Work: Functions and Strings' data-toggle="tooltip" >              
    <a href="func-string.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
  
<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="functions-write-code-questions">
<h1><span class="section-number">5.20. </span>Functions Write Code Questions<a class="headerlink" href="#functions-write-code-questions" title="Permalink to this headline">¶</a></h1>
<ol class="arabic">
<li><div id='funct_writecode1' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=funct_writecode1q data-question_label="5.20.1">
<div id=funct_writecode1q_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">remainder</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">x</span></code> and <code class="docutils literal notranslate"><span class="pre">y</span></code> (where <code class="docutils literal notranslate"><span class="pre">y</span></code> defaults to 4) and
returns the remainder of <code class="docutils literal notranslate"><span class="pre">x</span></code> divided by <code class="docutils literal notranslate"><span class="pre">y</span></code>. Ignore cases for when <code class="docutils literal notranslate"><span class="pre">x</span></code> is negative or <code class="docutils literal notranslate"><span class="pre">y</span></code> is less than or equal to 0.
For example, <code class="docutils literal notranslate"><span class="pre">remainder(5)</span></code> should return <code class="docutils literal notranslate"><span class="pre">1</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_writecode1q_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(remainder(5, 1),0,"remainder(5, 1)")
        self.assertEqual(remainder(0, 4),0,"remainder(0, 4)")
        self.assertEqual(remainder(2, 3),2,"remainder(2, 3)")
        self.assertEqual(remainder(5),1,"remainder(5)")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=funct_writecode1a data-question_label="5.20.2">
<div id=funct_writecode1a_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">remainder</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">x</span></code> and <code class="docutils literal notranslate"><span class="pre">y</span></code> (where <code class="docutils literal notranslate"><span class="pre">y</span></code> defaults to 4) and
returns the remainder of <code class="docutils literal notranslate"><span class="pre">x</span></code> divided by <code class="docutils literal notranslate"><span class="pre">y</span></code>. Ignore cases for when <code class="docutils literal notranslate"><span class="pre">x</span></code> is negative or <code class="docutils literal notranslate"><span class="pre">y</span></code> is less than or equal to 0.
For example, <code class="docutils literal notranslate"><span class="pre">remainder(5)</span></code> should return <code class="docutils literal notranslate"><span class="pre">1</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_writecode1a_editor" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
# Function headers end with a colon (:) and a default value for y
def remainder(x, y = 4):

    # Indent four spaces after the header
    # Use modulus division (%) instead of double slashes (//) to get the remainder
    z = x % y

    # Keyword return is lowercase and returns z (the remainder of x divided by y)
    return z
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=funct_writecode2q data-question_label="5.20.3">
<div id=funct_writecode2q_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">area_of_rect</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">width</span></code> and <code class="docutils literal notranslate"><span class="pre">length</span></code> and
returns the area (of the rectangle). For example, <code class="docutils literal notranslate"><span class="pre">area_of_rect(5,4)</span></code> should return <code class="docutils literal notranslate"><span class="pre">20</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_writecode2q_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(area_of_rect(1,7),7,"area_of_rect(1,7)")
        self.assertEqual(area_of_rect(5,4),20,"area_of_rect(5,4)")
        self.assertAlmostEqual(area_of_rect(5.4,4.3),23.22,2,"area_of_rect(5.4,4.3)")

myTests().main()
</textarea>
</div>
</div>
</li>
<li><div id='funct_writecode3' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=funct_writecode3q data-question_label="5.20.4">
<div id=funct_writecode3q_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">welcome_message</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">name</span></code> and
returns <code class="docutils literal notranslate"><span class="pre">&quot;Hello</span> <span class="pre">(name)!</span> <span class="pre">We</span> <span class="pre">hope</span> <span class="pre">that</span> <span class="pre">you</span> <span class="pre">will</span> <span class="pre">enjoy</span> <span class="pre">this</span> <span class="pre">course.</span> <span class="pre">:)&quot;</span></code>. For example,
<code class="docutils literal notranslate"><span class="pre">welcome_message('Aurora')</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;Hello</span> <span class="pre">Aurora!</span> <span class="pre">We</span> <span class="pre">hope</span> <span class="pre">that</span> <span class="pre">you</span> <span class="pre">will</span> <span class="pre">enjoy</span> <span class="pre">this</span> <span class="pre">course.</span> <span class="pre">:)&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_writecode3q_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(welcome_message('Aurora'),"Hello Aurora! We hope that you will enjoy this course. :)","welcome_message('Aurora')")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=funct_writecode3a data-question_label="5.20.5">
<div id=funct_writecode3a_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">welcome_message</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">name</span></code> and
returns <code class="docutils literal notranslate"><span class="pre">&quot;Hello</span> <span class="pre">(name)!</span> <span class="pre">We</span> <span class="pre">hope</span> <span class="pre">that</span> <span class="pre">you</span> <span class="pre">will</span> <span class="pre">enjoy</span> <span class="pre">this</span> <span class="pre">course.</span> <span class="pre">:)&quot;</span></code>. For example,
<code class="docutils literal notranslate"><span class="pre">welcome_message('Aurora')</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;Hello</span> <span class="pre">Aurora!</span> <span class="pre">We</span> <span class="pre">hope</span> <span class="pre">that</span> <span class="pre">you</span> <span class="pre">will</span> <span class="pre">enjoy</span> <span class="pre">this</span> <span class="pre">course.</span> <span class="pre">:)&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_writecode3a_editor" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def welcome_message(name):
    return "Hello " + name + "! We hope that you will enjoy this course. :)"
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=funct_writecode4q data-question_label="5.20.6">
<div id=funct_writecode4q_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">birthday</span></code> that takes in three parameters (<code class="docutils literal notranslate"><span class="pre">month</span></code>, <code class="docutils literal notranslate"><span class="pre">day</span></code>, and <code class="docutils literal notranslate"><span class="pre">year</span></code>) and returns it in the format <code class="docutils literal notranslate"><span class="pre">month</span></code>/<code class="docutils literal notranslate"><span class="pre">day</span></code>/<code class="docutils literal notranslate"><span class="pre">year</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">birthday(11,</span> <span class="pre">17,</span> <span class="pre">1990)</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;11/17/1990&quot;</span></code> and <code class="docutils literal notranslate"><span class="pre">birthday(7,</span> <span class="pre">5,</span> <span class="pre">2004)</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;7/5/2004&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_writecode4q_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(birthday(11, 17, 1990),"11/17/1990","birthday(11, 17, 1990)")
        self.assertEqual(birthday(1, 20, 2020),"1/20/2020","birthday(1, 20, 2020)")
        self.assertEqual(birthday(12, 1, 2000),"12/1/2000","birthday(12, 1, 2000)")
        self.assertEqual(birthday(7, 5, 2004),"7/5/2004","birthday(7, 5, 2004)")

myTests().main()
</textarea>
</div>
</div>
</li>
<li><div id='funct_writecode5' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=funct_writecode5q data-question_label="5.20.7">
<div id=funct_writecode5q_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">address</span></code> that combines 3 different string address parameters (<code class="docutils literal notranslate"><span class="pre">city</span></code>, <code class="docutils literal notranslate"><span class="pre">state</span></code>, and <code class="docutils literal notranslate"><span class="pre">zip</span></code>) and returns a user’s address.
After the <code class="docutils literal notranslate"><span class="pre">city</span></code> and <code class="docutils literal notranslate"><span class="pre">state</span></code> inputs, add a comma and a space. For example, <code class="docutils literal notranslate"><span class="pre">address('Seattle',</span> <span class="pre">'WA',</span> <span class="pre">'98105')</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;Seattle,</span> <span class="pre">WA,</span> <span class="pre">98105&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_writecode5q_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(len(address('Seattle', 'WA', '98105').split(',')),3,"len(address('Seattle', 'WA', '98105').split(','))")
        self.assertEqual(address('Seattle', 'WA', '98105'),"Seattle, WA, 98105","address('Seattle', 'WA', '98105')")
        self.assertEqual(address('Ann Arbor', 'MI', '48105'),"Ann Arbor, MI, 48105","address('Ann Arbor', 'MI', '48105')")
        self.assertEqual(address('Los Angeles', 'CA', '90005'),"Los Angeles, CA, 90005","address('Los Angeles', 'CA', '90005')")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=funct_writecode5a data-question_label="5.20.8">
<div id=funct_writecode5a_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">address</span></code> that combines 3 different string address parameters (<code class="docutils literal notranslate"><span class="pre">city</span></code>, <code class="docutils literal notranslate"><span class="pre">state</span></code>, and <code class="docutils literal notranslate"><span class="pre">zip</span></code>) and returns a user’s address.
After the <code class="docutils literal notranslate"><span class="pre">city</span></code> and <code class="docutils literal notranslate"><span class="pre">state</span></code> inputs, add a comma and a space. For example, <code class="docutils literal notranslate"><span class="pre">address('Seattle',</span> <span class="pre">'WA',</span> <span class="pre">'98105')</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;Seattle,</span> <span class="pre">WA,</span> <span class="pre">98105&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_writecode5a_editor" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def address(city, state, zip):
    user_address = city + ", " + state + ", " + zip
    return user_address
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=funct_writecode6q data-question_label="5.20.9">
<div id=funct_writecode6q_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">squareArea</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">length</span></code> and calculates the area of the square.
It returns a string with this format: “The total area of the square with length (length) is (area).”. For example,
<code class="docutils literal notranslate"><span class="pre">squareArea(10)</span></code> would return <code class="docutils literal notranslate"><span class="pre">&quot;The</span> <span class="pre">total</span> <span class="pre">area</span> <span class="pre">of</span> <span class="pre">the</span> <span class="pre">square</span> <span class="pre">with</span> <span class="pre">length</span> <span class="pre">10</span> <span class="pre">is</span> <span class="pre">100.&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_writecode6q_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(squareArea(10),"The total area of the square with length 10 is 100.","squareArea(10)")
        self.assertEqual(squareArea(5),"The total area of the square with length 5 is 25.","squareArea(5)")
        self.assertEqual(squareArea(4.5),"The total area of the square with length 4.5 is 20.25.","squareArea(4.5)")


myTests().main()
</textarea>
</div>
</div>
</li>
<li><div id='funct_writecode7' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=funct_writecode7q data-question_label="5.20.10">
<div id=funct_writecode7q_question class="ac_question col-md-12">
<p>Write three functions called <code class="docutils literal notranslate"><span class="pre">addNumbers</span></code>, <code class="docutils literal notranslate"><span class="pre">subtractNumbers</span></code>, and <code class="docutils literal notranslate"><span class="pre">calculate</span></code>.
The function <code class="docutils literal notranslate"><span class="pre">addNumbers</span></code> should take two numbers (<code class="docutils literal notranslate"><span class="pre">x</span></code> and <code class="docutils literal notranslate"><span class="pre">y</span></code>) as parameters and return the value of adding them together,
while <code class="docutils literal notranslate"><span class="pre">subtractNumbers</span></code> should also take two numbers (<code class="docutils literal notranslate"><span class="pre">x</span></code> and <code class="docutils literal notranslate"><span class="pre">y</span></code>) and return the value of <code class="docutils literal notranslate"><span class="pre">x</span></code> minus <code class="docutils literal notranslate"><span class="pre">y</span></code>.
Lastly, define a function called <code class="docutils literal notranslate"><span class="pre">calculate</span></code> that takes three numbers (<code class="docutils literal notranslate"><span class="pre">a</span></code>, <code class="docutils literal notranslate"><span class="pre">b</span></code>, and <code class="docutils literal notranslate"><span class="pre">c</span></code>) and uses
<code class="docutils literal notranslate"><span class="pre">addNumbers</span></code> and <code class="docutils literal notranslate"><span class="pre">subtractNumbers</span></code> to add <code class="docutils literal notranslate"><span class="pre">a</span></code> and <code class="docutils literal notranslate"><span class="pre">b</span></code> and subtract <code class="docutils literal notranslate"><span class="pre">c</span></code>. The value should be returned. For example,
<code class="docutils literal notranslate"><span class="pre">calculate(2,3,4)</span></code> should return <code class="docutils literal notranslate"><span class="pre">1</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_writecode7q_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def addNumbers(x, y):
    # your code here

def subtractNumbers(x, y):
    # your code here

def calculate(a, b, c):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(addNumbers(2,3),5,"addNumbers(2,3)")
        self.assertEqual(subtractNumbers(6,5),1,"subtractNumbers(6,5)")
        self.assertEqual(subtractNumbers(3,5),-2,"subtractNumbers(3,5)")
        self.assertEqual(calculate(2,3,4),1,"calculate(2,3,4)")
        self.assertAlmostEqual(calculate(1.0, 0.0, 3.5),-2.5,1,"calculate(1.0, 0.0, 3.5)")
        self.assertEqual(calculate(-10, -9, -3),-16,"calculate(-10, -9, -3)")
        self.assertAlmostEqual(calculate(-10, -7.3, 3),-20.3,1,"calculate(-10, -7.3, 3)")
        self.assertAlmostEqual(calculate(10, -7.3, -3),5.7,1,"calculate(10, -7.3, -3)")
        self.assertAlmostEqual(calculate(-10.1, 5.2, -5),0.1,1,"calculate(-10.1, 5.2, -5)")
        self.assertAlmostEqual(calculate(10.6, -7.3, 3.8),-0.5,1,"calculate(10.6, -7.3, 3.8)")
        self.assertAlmostEqual(calculate(-9.8, 7.3, 3),-5.5,1,"calculate(-9.8, 7.3, 3)")
        self.assertAlmostEqual(calculate(8, 7.3, -0.2),15.5,1,"calculate(8, 7.3, -0.2)")
        self.assertAlmostEqual(calculate(10.5, 7.3, 0.2),17.6,1,"calculate(10.5, 7.3, 0.2)")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=funct_writecode7a data-question_label="5.20.11">
<div id=funct_writecode7a_question class="ac_question col-md-12">
<p>Write three functions called <code class="docutils literal notranslate"><span class="pre">addNumbers</span></code>, <code class="docutils literal notranslate"><span class="pre">subtractNumbers</span></code>, and <code class="docutils literal notranslate"><span class="pre">calculate</span></code>.
The function <code class="docutils literal notranslate"><span class="pre">addNumbers</span></code> should take two numbers (<code class="docutils literal notranslate"><span class="pre">x</span></code> and <code class="docutils literal notranslate"><span class="pre">y</span></code>) as parameters and return the value of adding them together,
while <code class="docutils literal notranslate"><span class="pre">subtractNumbers</span></code> should also take two numbers (<code class="docutils literal notranslate"><span class="pre">x</span></code> and <code class="docutils literal notranslate"><span class="pre">y</span></code>) and return the value of <code class="docutils literal notranslate"><span class="pre">x</span></code> minus <code class="docutils literal notranslate"><span class="pre">y</span></code>.
Lastly, define a function called <code class="docutils literal notranslate"><span class="pre">calculate</span></code> that takes three numbers (<code class="docutils literal notranslate"><span class="pre">a</span></code>, <code class="docutils literal notranslate"><span class="pre">b</span></code>, and <code class="docutils literal notranslate"><span class="pre">c</span></code>) and uses
<code class="docutils literal notranslate"><span class="pre">addNumbers</span></code> and <code class="docutils literal notranslate"><span class="pre">subtractNumbers</span></code> to add <code class="docutils literal notranslate"><span class="pre">a</span></code> and <code class="docutils literal notranslate"><span class="pre">b</span></code> and subtract <code class="docutils literal notranslate"><span class="pre">c</span></code>. The value should be returned. For example,
<code class="docutils literal notranslate"><span class="pre">calculate(2,3,4)</span></code> should return <code class="docutils literal notranslate"><span class="pre">1</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_writecode7a_editor" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def addNumbers(x, y):

    # Return the value of adding x and y together
    return x + y

def subtractNumbers(x, y):
    # Return the value of subtracting y from x
    return x - y

def calculate(a, b, c):

    # Call addNumbers, passing in a and b as parameters, and set the answer equal to the variable result
    result = addNumbers(a, b)
    # Call subtractNumbers on the result of the addition and c, the un
    finalResult = subtractNumbers(result, c)
    # Returns the finalResult
    return finalResult
</textarea>
</div>
</div>
</div>
    </div>

</li>
</ol>
</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 5.19. Functions Mixed-Up Code Questions' data-toggle="tooltip">
  <a href="functions_mixedupcode.html" aria-label="Previous - 5.19. Functions Mixed-Up Code Questions">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 5.21. Group Work: Functions and Strings' data-toggle="tooltip" >
    <a href="func-string.html" aria-label="Next - 5.21. Group Work: Functions and Strings">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>