<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>5.29. Group Work: Functions and Lists &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />
    <link rel="stylesheet" type="text/css" href="../_static/637.css?v=bddd2011b20db76a20ea" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.css?v=5b1dc1bb539ccf7c69f7" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/runtime.bundle.js?v=f3c8dd9e2667ee8dc613"></script>
    <script src="../_static/637.bundle.js?v=1123ad75f9b821cdd636"></script>
    <script src="../_static/runestone.bundle.js?v=ef2d58494168454a1d6a"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="5.30. Functions with Lists Multiple Choice Questions" href="function_list_exercises.html" />
    <link rel="prev" title="5.28. Functions and Conditionals Write Code Questions" href="function_conditional_writecode.html" />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'py4e-int';
    eBookConfig.basecourse = 'py4e-int';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar" aria-label="navbar toggle">toggle navbar</span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
          <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          py4e-int
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=funcWithLists'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">5.28. </span>Functions and Conditionals Write Code Questions" data-toggle="tooltip">
    <a href="function_conditional_writecode.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">5.30. </span>Functions with Lists Multiple Choice Questions' data-toggle="tooltip" >              
    <a href="function_list_exercises.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  document.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="group-work-functions-and-lists">
<h1><span class="section-number">5.29. </span>Group Work: Functions and Lists<a class="headerlink" href="#group-work-functions-and-lists" title="Permalink to this headline">¶</a></h1>
<p>It is best to use a POGIL approach with the following. In POGIL students work
in groups on activities and each member has an assigned role.  For more information see <a class="reference external" href="https://cspogil.org/Home">https://cspogil.org/Home</a>.</p>
<p><strong>Learning Objectives</strong></p>
<p>Students will know and be able to do the following.</p>
<p><em>Content Objectives:</em></p>
<ul class="simple">
<li><p>Learn how to access items in a list with positive and negative indices.</p></li>
<li><p>Recognize built-in functions that work on lists (len, max, min, sum).</p></li>
<li><p>Recognize list methods (append, pop, extend, sort, reverse, index).</p></li>
<li><p>Learn how to use the slice operator to copy parts of a list.</p></li>
</ul>
<p><em>Process Objectives:</em></p>
<ul class="simple">
<li><p>Predict the outcome of code with lists</p></li>
</ul>
<div class="section" id="lists">
<h2><span class="section-number">5.29.1. </span>Lists<a class="headerlink" href="#lists" title="Permalink to this headline">¶</a></h2>
<p>A list holds items in order and you can get the value at an index, just like you can with strings.</p>

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.29.1.1" id="funct_ll_list_check_first_line"  style="visibility: hidden;">
            <p>p9-1: What is the first thing that will be printed when the code below runs?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [1, 1], "feedback": "<p>It will print the item at index 0 in alist.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>What is the item at index 0 in the passed list?</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.29.1.2" id="funct_ll_list_check_last_line"  style="visibility: hidden;">
            <p>p9-2: What is the last thing that will be printed when the code below runs?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^\\s*'?bye'?\\s*$", "regexFlags": "", "feedback": "<p>It will print the item at the end of the list in the second call to the function.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>What is the last item in the list in the second call?</p>\n"}]]
        </script>

        </div>
        </div>
            
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=func_ll_ac_list_check data-question_label="5.29.1.3">
<div id=func_ll_ac_list_check_question class="ac_question col-md-12">
<p>Run this code to see what it prints.</p>

</div>
<textarea data-lang="python" id="func_ll_ac_list_check_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
         data-caption='Accessing list values at indices'  data-wasm=/_static
     style="visibility: hidden;">
# function definition
def list_check(alist):
    print(alist[0])
    print(alist[1])
    print(len(alist))
    print(alist[-1])

# function definition
def main():
    list_check([1, 2, 5])
    list_check(['hi', 'bye'])

# function call
main()
</textarea>
</div>
</div>

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.29.1.4" id="funct_ll_list_char"  style="visibility: hidden;">
            <p>p9-4: What characters are used to indicate the start and end of a list?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^\\s*\\[\\]|\\]\\[\\s*$", "regexFlags": "", "feedback": "<p>Square brackets are used at the start and end of a list.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Look at the parameter to the function call in the code above.</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.29.1.5" id="funct_ll_first_index"  style="visibility: hidden;">
            <p>p9-5: What index is used to get the first item in a list?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [0, 0], "feedback": "<p>The first item in a list is at index 0.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Try again!</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.29.1.6" id="funct_ll_last_index"  style="visibility: hidden;">
            <p>p9-6: What negative index is used to get the last item in a list?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [-1, -1], "feedback": "<p>The last item in a list is at index -1, which is the length of the list - 1.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Try again!</p>\n"}]]
        </script>

        </div>
        </div>
            </div>
<div class="section" id="built-in-functions-that-work-on-lists">
<h2><span class="section-number">5.29.2. </span>Built-in Functions That Work on Lists<a class="headerlink" href="#built-in-functions-that-work-on-lists" title="Permalink to this headline">¶</a></h2>
<p>There are several built-in functions in Python that work on lists.</p>

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.29.2.1" id="funct_ll_list_func_first"  style="visibility: hidden;">
            <p>p9-7: What is the first value that will be printed when the code below runs?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [3, 3], "feedback": "<p>It will print the maximum value in the first list, which is 3.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>What is the maximum value in the first list that is passed to the function?</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.29.2.2" id="funct_ll_list_func_last"  style="visibility: hidden;">
            <p>p9-8: What is the last value that will be printed when the code below runs?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^\\s*95\\.0\\s*$", "regexFlags": "", "feedback": "<p>It will print the average of the values in the last list that is passed to the function.</p>\n"}, {"number": [95, 95], "feedback": "<p>This will print a floating point (will have a decimal and at least one digit after the decimal) result not an integer result.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>What is the average of the items in the last list that is passed to the function?</p>\n"}]]
        </script>

        </div>
        </div>
            
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=func_ll_ac_list_func data-question_label="5.29.2.3">
<div id=func_ll_ac_list_func_question class="ac_question col-md-12">
<p>Run this code to see what it prints.</p>

</div>
<textarea data-lang="python" id="func_ll_ac_list_func_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
         data-caption='List methods'  data-wasm=/_static
     style="visibility: hidden;">
# function definition
def list_func(my_list):
    print(max(my_list))
    print(min(my_list))
    print(len(my_list))
    print(sum(my_list))
    print(sum(my_list) / len(my_list))

# function definition
def main():
    list_func([1, 2, 3])
    list_func([90, 100])

# function call
main()
</textarea>
</div>
</div>

<div class="runestone">
<ul data-component="dragndrop" data-question_label="5.29.2.4" id="func_ll_dnd_funct"  style="visibility: hidden;">
    <span data-subcomponent="question">p9-10: Drag each built-in function name to what it does.</span>
	<span data-subcomponent=feedback>Read the chapter on functions and try again.</span>

    <li data-subcomponent="draggable" id="func_ll_dnd_funct_drag1">len</li>
    <li data-subcomponent="dropzone" for="func_ll_dnd_funct_drag1">Returns the length of an object (the number of items in the list).</li>

    <li data-subcomponent="draggable" id="func_ll_dnd_funct_drag2">sum</li>
    <li data-subcomponent="dropzone" for="func_ll_dnd_funct_drag2">Returns the total of the values in the list.</li>

    <li data-subcomponent="draggable" id="func_ll_dnd_funct_drag3">min</li>
    <li data-subcomponent="dropzone" for="func_ll_dnd_funct_drag3">Returns the minimum value in the list.</li>

    <li data-subcomponent="draggable" id="func_ll_dnd_funct_drag4">max</li>
    <li data-subcomponent="dropzone" for="func_ll_dnd_funct_drag4">Returns the maximum value in the list.</li>
</ul></div></div>
<div class="section" id="list-methods">
<h2><span class="section-number">5.29.3. </span>List Methods<a class="headerlink" href="#list-methods" title="Permalink to this headline">¶</a></h2>
<p>Lists are objects of the <code class="docutils literal notranslate"><span class="pre">list</span></code> class and have methods that operate on list objects.</p>

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.29.3.1" id="funct_ll_list_methods_first"  style="visibility: hidden;">
            <p>p9-11: What is the second thing that will be printed when the code below runs?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^\\s*\\[1\\]\\s*$", "regexFlags": "", "feedback": "<p>It will print the contents of the passed list</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>What does alist equal at the start of the function the first time the function is called?</p>\n"}]]
        </script>

        </div>
        </div>
            
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=func_ll_ac_list_methods data-question_label="5.29.3.2">
<div id=func_ll_ac_list_methods_question class="ac_question col-md-12">
<p>Run this code to see what it prints.</p>

</div>
<textarea data-lang="python" id="func_ll_ac_list_methods_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
         data-caption='List methods'  data-wasm=/_static
     style="visibility: hidden;">
# function definition
def list_methods(alist):
    print(type(alist))
    print(alist)
    alist.append(3)
    print(alist)
    alist.append([2])
    print(alist)
    alist.pop(1)
    print(alist)
    alist.extend([8, 11])
    print(alist)

# function definition
def main():
    l1 = [1]
    list_methods(l1)
    print(l1)
    l1 = ['hi', 'bye']
    list_methods(l1)
    print(l1)

# function call
main()
</textarea>
</div>
</div>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>Lists are mutable (changeable).  List methods like append and pop change the current list.</p>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=func_ll_ac_list_methods2 data-question_label="5.29.3.3">
<div id=func_ll_ac_list_methods2_question class="ac_question col-md-12">
<p>Run this code to see what it prints.</p>

</div>
<textarea data-lang="python" id="func_ll_ac_list_methods2_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
         data-caption='More list methods'  data-wasm=/_static
     style="visibility: hidden;">
# function definition
def list_methods2(alist):
    print(alist)
    alist.reverse()
    print(alist)
    alist.sort()
    print(alist)

# function definition
def main():
    l1 = [-2, 50, -20, 30]
    list_methods2(l1)
    print(l1)
    l1 = ['hi', 'bye', 'apple']
    list_methods2(l1)
    print(l1)

# function call
main()
</textarea>
</div>
</div>

<div class="runestone">
<ul data-component="dragndrop" data-question_label="5.29.3.4" id="func_ll_dnd_func_methods"  style="visibility: hidden;">
    <span data-subcomponent="question">p9-14: Drag each built-in function name to what it does.</span>
	<span data-subcomponent=feedback>Read the chapter on functions and try again.</span>

    <li data-subcomponent="draggable" id="func_ll_dnd_func_methods_drag1">pop</li>
    <li data-subcomponent="dropzone" for="func_ll_dnd_func_methods_drag1">Removes the value at the specified index.</li>

    <li data-subcomponent="draggable" id="func_ll_dnd_func_methods_drag2">append</li>
    <li data-subcomponent="dropzone" for="func_ll_dnd_func_methods_drag2">Adds the value or list at the end of the current list.</li>

    <li data-subcomponent="draggable" id="func_ll_dnd_func_methods_drag3">extend</li>
    <li data-subcomponent="dropzone" for="func_ll_dnd_func_methods_drag3">Adds the contents of the passed list to the end of the current list.</li>

    <li data-subcomponent="draggable" id="func_ll_dnd_func_methods_drag4">sort</li>
    <li data-subcomponent="dropzone" for="func_ll_dnd_func_methods_drag4">Sort the contents of the list in ascending order.</li>

    <li data-subcomponent="draggable" id="func_ll_dnd_func_methods_drag5">reverse</li>
    <li data-subcomponent="dropzone" for="func_ll_dnd_func_methods_drag5">Reverse the contents of the list.</li>
</ul></div></div>
<div class="section" id="using-the-slice-operator">
<h2><span class="section-number">5.29.4. </span>Using the Slice Operator<a class="headerlink" href="#using-the-slice-operator" title="Permalink to this headline">¶</a></h2>
<p>You can use the slice operator[n:m] with lists to make a copy of the list with all items
from n to m - 1, just like you can with strings.</p>

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.29.4.1" id="funct_ll_fitb_slice_first"  style="visibility: hidden;">
            <p>p9-15: How many items will be in the first list that is printed below?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [4, 4], "feedback": "<p>It will print the result of copying the first list which has 4 items.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Run the code to check.</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.29.4.2" id="funct_ll_fitb_slice_last"  style="visibility: hidden;">
            <p>p9-16: How many items will be in the last list that is printed below?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [4, 4], "feedback": "<p>The slice operator does not change the original list.  It makes a new list.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Run the code to check.</p>\n"}]]
        </script>

        </div>
        </div>
            
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=func_ll_ac_list_slice data-question_label="5.29.4.3">
<div id=func_ll_ac_list_slice_question class="ac_question col-md-12">
<p>Run this code to see what it prints.</p>

</div>
<textarea data-lang="python" id="func_ll_ac_list_slice_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
         data-caption='Add tests'  data-wasm=/_static
     style="visibility: hidden;">
# function definition
def list_slice(num_list):
    a = num_list[:]
    print(a)
    b = num_list[:2]
    print(b)
    c = num_list[2:]
    print(c)

# function definition
def main():
    l1 = [-2, 50, -20, 30]
    list_slice(l1)
    print(l1)
    l1 = ['hi', 'bye', 'apple', 'dog']
    list_slice(l1)
    print(l1)

# function call
main()
</textarea>
</div>
</div>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>The slice operator always returns a new object.  It doesn’t change the current object (list or string).</p>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=fuct_ac_list_first_half data-question_label="5.29.4.4">
<div id=fuct_ac_list_first_half_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">first_half</span></code> that takes a list and returns a new list (use the slice operator) with just the items from the first half of the original list.  For example, first_half([1,2,3,4]) would return [1, 2] and first_half([7,8,9]) should return [7].</p>

</div>
<textarea data-lang="python" id="fuct_ac_list_first_half_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def first_half(alist):

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(first_half([1,2,3,4]), [1,2], 'first_half([1,2,3,4])')
        self.assertEqual(first_half([7,8,9]), [7], 'first_half([7,8,9])')
        self.assertEqual(first_half([]), [], 'first_half([])')
        self.assertEqual(first_half([6]), [], 'first_half([6])')
        self.assertEqual(first_half([1,2,3,4,5]), [1,2], 'first_half([1,2,3,4,5])')
        self.assertEqual(first_half([1,2,3,4,5,6]), [1,2,3], 'first_half([1,2,3,4,5,6])')

myTests().main()
</textarea>
</div>
</div>
</div>
</div>


  <div id=scprogresscontainer>
    You have attempted <span id=scprogresstotal></span> of <span id=scprogressposs></span> activities on this page <div id="subchapterprogress"></div>
  </div>
  
      
<ul>
  <li id="relations-prev" style="opacity:0.4" onmouseover="mouseOverPrev()" onmouseout="mouseOutPrev()"class="navLink" title='Previous Section - 5.28. Functions and Conditionals Write Code Questions' data-toggle="tooltip">
    <a href="function_conditional_writecode.html" aria-label="5.28. Functions and Conditionals Write Code Questions">
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">
            <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
        </div>
    </a>
  </li>
</ul>

  
<ul>
  <li id="relations-next" style="opacity:0.4" onmouseover="mouseOverNext()" onmouseout="mouseOutNext()" class="navLink" title='Next Section - 5.30. Functions with Lists Multiple Choice Questions' data-toggle="tooltip" >
    <a href="function_list_exercises.html" aria-label="5.30. Functions with Lists Multiple Choice Questions">
      <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">
          <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
      </div>
    </a>
  </li>
</ul>

<script>
document.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>




<script> 
  document.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>