<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>5.33. Group Work: Functions with Loops &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/637.css?v=bddd2011b20db76a20ea" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.css?v=7c5388ad4389eae1a318" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/runtime.bundle.js?v=aa7c717c38c5405af9e6"></script>
    <script src="../_static/637.bundle.js?v=1123ad75f9b821cdd636"></script>
    <script src="../_static/runestone.bundle.js?v=9c8e466e8cbed4a3af5a"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="5.34. Functions with Loops Multiple Choice Questions" href="functions_loops_exercises.html" />
    <link rel="prev" title="5.32. Functions and Lists Write Code Questions" href="function_list_writecode.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'py4e-int';
    eBookConfig.basecourse = 'py4e-int';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          py4e-int
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=funcWithLoops'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">5.32. </span>Functions and Lists Write Code Questions" data-toggle="tooltip">
    <a href="function_list_writecode.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">5.34. </span>Functions with Loops Multiple Choice Questions' data-toggle="tooltip" >              
    <a href="functions_loops_exercises.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
  
<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="group-work-functions-with-loops">
<h1><span class="section-number">5.33. </span>Group Work: Functions with Loops<a class="headerlink" href="#group-work-functions-with-loops" title="Permalink to this headline">¶</a></h1>
<p>It is best to use a POGIL approach with the following. In POGIL students work
in groups on activities and each member has an assigned role.  For more information see <a class="reference external" href="https://cspogil.org/Home">https://cspogil.org/Home</a>.</p>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>If you work in a group, have only one member of the group fill in the answers on this page.  You will be able to share your answers with the group at the bottom of the page.</p>
</div>
<p><strong>Learning Objectives</strong></p>
<p>Students will know and be able to do the following.</p>
<p><em>Content Objectives:</em></p>
<ul class="simple">
<li><p>Recognize a for-each loop.</p></li>
<li><p>Recognize a range.</p></li>
<li><p>Recognize a for loop.</p></li>
<li><p>Recognize a while loop.</p></li>
</ul>
<p><em>Process Objectives:</em></p>
<ul class="simple">
<li><p>Predict the outcome of code with lists and loops.</p></li>
<li><p>Modify code with loops.</p></li>
<li><p>Fix code that processes a list with loops.</p></li>
</ul>
<div class="section" id="the-for-each-loop">
<h2><span class="section-number">5.33.1. </span>The For-Each Loop<a class="headerlink" href="#the-for-each-loop" title="Permalink to this headline">¶</a></h2>
<p>A for-each loop in Python will loop though the items in a list starting with the item at index 0, then index 1, and so on till the last item in the list.</p>

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.33.1.1" id="funct_ll_fitb_count_odd_first"  style="visibility: hidden;">
            <p>Q-1: What is the first thing that will be printed by the code below?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [1, 1], "feedback": "<p>It will print the number of values that are odd in the first list, which is one.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Run the code to check.</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.33.1.2" id="funct_ll_fitb_count_odd_last"  style="visibility: hidden;">
            <p>Q-2: What is the last thing that will be printed by the code below?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [0, 0], "feedback": "<p>It will print the number of values that are odd in the last list which is zero.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Run the code to check.</p>\n"}]]
        </script>

        </div>
        </div>
            
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=func_ll_ac_count_odd data-question_label="5.33.1.3">
<div id=func_ll_ac_count_odd_question class="ac_question col-md-12">
<p>Run this code to see what it prints.</p>

</div>
<textarea data-lang="python" id="func_ll_ac_count_odd_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
         data-caption='Count odd numbers'  data-wasm=/_static
     style="visibility: hidden;">
# function definition
def count_odd(num_list):
    count = 0
    for num in num_list:
        if num % 2 == 1:
            count += 1
    return count

# function definition
def main():
    list1 = [2, 8, 9]
    print(count_odd(list1))
    list1 = [1, 3, 5]
    print(count_odd(list1))
    list1 = [0]
    print(count_odd(list1))

# function call
main()
</textarea>
</div>
</div>

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.33.1.4" id="funct_ll_fitb_count_odd_var"  style="visibility: hidden;">
            <p>Q-4: What is the name of the variable that can be used to refer to the current list item in the for-each loop?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^\\s*num\\s*$", "regexFlags": "", "feedback": "<p>The for-each loop uses num to refer to the current list item.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>What is the first name in the for-each loop?</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="func_ll_pp_total_even" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Drag the blocks from the left and put them in the correct order on the right to define a function <code class="docutils literal notranslate"><span class="pre">total_even</span></code> that returns the total of the even numbers in the passed list.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="5.33.1.5"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def total_even(alist):
---
Def total_even(alist): #paired
---
    total = 0
---
    for num in alist:
---
    for num in alist #paired
---
        if num % 2 == 0:
---
        if num % 2 == 1: #paired
---
            total += num
---
    return total
---
    Return total #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=fuct_ac_ll_sum67_fix data-question_label="5.33.1.6">
<div id=fuct_ac_ll_sum67_fix_question class="ac_question col-md-12">
<p>Fix the <code class="docutils literal notranslate"><span class="pre">sum67</span></code> function below that takes a list and should return the total of the items in the list except for all the numbers between a 6 and 7 (inclusive).   For example, sum67([1,2]) should return 3, but sum67([2, 6, 8, 7, 2]) should return 4.</p>

</div>
<textarea data-lang="python" id="fuct_ac_ll_sum67_fix_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def sum67(nums):
    total = 0               # initialize the total
    found_6 = True          # initialize a Boolean flag
    for num in nums:        # loop through the items in a list
        if found_6 && num == 7:
            found_6 = False # set the Boolean flag to false
        elif num = 6:
            found_6 = True  # set the Boolean flag to True
        elif found_6:
            continue        # go back to the top of the loop
        else:
            total += num    # add num to total
        return total        # return the total

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sum67([1,2]), 3, 'sum67([1,2])')
        self.assertEqual(sum67([2, 6, 8, 7, 2]), 4, 'sum67([2, 6, 8, 7, 2])')
        self.assertEqual(sum67([6, 2, 7]), 0, 'sum67([6, 2, 7])')
        self.assertEqual(sum67([6, 2, 7, 5]), 5, 'sum67([6, 2, 7, 5])')
        self.assertEqual(sum67([7, 2, 6]), 9, 'sum67([7, 2, 6])')
        self.assertEqual(sum67([4, 6, 7]), 4, 'sum67([4, 6, 7])')
        self.assertEqual(sum67([]), 0, 'sum67([])')

myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="range-and-for">
<h2><span class="section-number">5.33.2. </span>Range and For<a class="headerlink" href="#range-and-for" title="Permalink to this headline">¶</a></h2>
<p>How do you loop just a set number of times?  You can use
the built-in <code class="docutils literal notranslate"><span class="pre">range</span></code> function to do this.</p>

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.33.2.1" id="funct_ll_fitb_print_to_first"  style="visibility: hidden;">
            <p>Q-7: What is first value that that following code prints?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [0, 0], "feedback": "<p>The code prints from 0 to the passed end.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Run the code and see.</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.33.2.2" id="funct_ll_fitb_print_to_last"  style="visibility: hidden;">
            <p>Q-8: What is last value that that following code prints?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [2, 2], "feedback": "<p>The code prints from 0 to the passed end (exclusive).  The end is 3 which means the last value is 2.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>What is the parameter in the second call to count_to?</p>\n"}]]
        </script>

        </div>
        </div>
            
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=func_ll_ac_count_to data-question_label="5.33.2.3">
<div id=func_ll_ac_count_to_question class="ac_question col-md-12">
<p>Run this code to see what it prints.</p>

</div>
<textarea data-lang="python" id="func_ll_ac_count_to_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
         data-caption='Range example with end'  data-wasm=/_static
     style="visibility: hidden;">
# function definition
def print_to(end):
    for x in range(end):
        print(x)

# function definition
def main():
    print_to(5)
    print()
    print_to(3)

# function call
main()
</textarea>
</div>
</div>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>The <code class="docutils literal notranslate"><span class="pre">range(end)</span></code> function will produce values from 0 to end - 1.</p>
</div>

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.33.2.4" id="funct_ll_fitb_print_from_to_first"  style="visibility: hidden;">
            <p>Q-10: What is first value that that following code prints?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [1, 1], "feedback": "<p>The code prints from start (inclusive) to the passed end (exclusive) and the start is one.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Run the code and see.</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.33.2.5" id="funct_ll_fitb_print_from_to_last"  style="visibility: hidden;">
            <p>Q-11: What is last value that that following code prints?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [9, 9], "feedback": "<p>The code prints from start (inclusive) to the passed end (exclusive) so it ends with 9.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Run the code and see.</p>\n"}]]
        </script>

        </div>
        </div>
            
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=func_ll_ac_print_from_to data-question_label="5.33.2.6">
<div id=func_ll_ac_print_from_to_question class="ac_question col-md-12">
<p>Run this code to see what it prints.</p>

</div>
<textarea data-lang="python" id="func_ll_ac_print_from_to_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
         data-caption='Range example with start and end'  data-wasm=/_static
     style="visibility: hidden;">
# function definition
def print_from_to(start, end):
    for x in range(start, end):
        print(x)

# function definition
def main():
    print_from_to(1,5)
    print()
    print_from_to(5,10)

# function call
main()
</textarea>
</div>
</div>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>The function range(start, end) will return a range object (an iterator) that allows you to loop from start (inclusive) to end (exclusive).</p>
</div>

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.33.2.7" id="funct_ll_fitb_print_from_to_by_first"  style="visibility: hidden;">
            <p>Q-13: What is first value that that following code prints?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [10, 10], "feedback": "<p>The code prints from start (inclusive) to the passed end (exclusive) and the start is ten.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Run the code and see.</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.33.2.8" id="funct_ll_fitb_print_from_to_by_last"  style="visibility: hidden;">
            <p>Q-14: What is last value that that following code prints?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [9, 9], "feedback": "<p>The code prints from start (inclusive) which is 1 to the passed end (exclusive) which is 10 and changes by 2 each time so the end is 9.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Run the code and see.</p>\n"}]]
        </script>

        </div>
        </div>
            
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=func_ll_ac_print_from_to_by data-question_label="5.33.2.9">
<div id=func_ll_ac_print_from_to_by_question class="ac_question col-md-12">
<p>Run this code to see what it prints.</p>

</div>
<textarea data-lang="python" id="func_ll_ac_print_from_to_by_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
         data-caption='Range example with start, end, and by'  data-wasm=/_static
     style="visibility: hidden;">
# function definition
def print_from_to_by(start, end, by):
    for x in range(start, end, by):
        print(x)

# function definition
def main():
    print_from_to_by(10,0,-1)
    print()
    print_from_to_by(1,10,2)

# function call
main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="func_ll_pp_total_at_odd_indices" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Drag the blocks from the left and put them in the correct order on the right to define a function <code class="docutils literal notranslate"><span class="pre">total_at_odd_indices</span></code> that returns the total of the numbers at odd indices in the passed list.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="5.33.2.10"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def total_at_odd_indices(alist):
---
def total_at_odd_indices(alist) #paired
---
     total = 0
---
     for i in range(1,len(alist),2):
---
     for i in range(1,len(alist)): #paired
---
         total += alist[i]
---
         total += i #paired
---
     return total
        </pre>
        </div>
        </div>
    </div>
<div class="section" id="while-loops">
<h2><span class="section-number">5.33.3. </span>While Loops<a class="headerlink" href="#while-loops" title="Permalink to this headline">¶</a></h2>
<p>A while loop repeats while a Boolean expression is True.</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=func_ll_ac_inifite data-question_label="5.33.3.1">
<div id=func_ll_ac_inifite_question class="ac_question col-md-12">
<p>Try running the code below.</p>

</div>
<textarea data-lang="python" id="func_ll_ac_inifite_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
         data-caption='Example infinite loop'  data-wasm=/_static
     style="visibility: hidden;">
def example():
    count = 0
    while (True):
        print("This is the song that never ends", count)
        count += 1
        if (count > 100):
            break

def main():
    example()

main()
</textarea>
</div>
</div>

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5.33.3.2" id="funct_ll_fitb_while_break"  style="visibility: hidden;">
            <p>Q-18: What keyword is used to stop the loop in the above code?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^\\s*break\\s*$", "regexFlags": "", "feedback": "<p>The break keyword will stop the enclosing loop.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>How does the code stop above?</p>\n"}]]
        </script>

        </div>
        </div>
            <p>What do you think would happen if you deleted lines 6 and 7 in the above code?</p>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>A loop that never ends is called an infinite loop.  A while loop should have some way to end.  If you have an infinite loop you may need to refresh the page to stop it.</p>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=func_ll_ac_while data-question_label="5.33.3.3">
<div id=func_ll_ac_while_question class="ac_question col-md-12">
<p>Run this code to see what it prints.</p>

</div>
<textarea data-lang="python" id="func_ll_ac_while_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
         data-caption='While Example'  data-wasm=/_static
     style="visibility: hidden;">
import random
def check_value(target, actual):
    if target == actual:
        return("Correct")
    elif (actual < target):
        return("Too Low")
    else:
        return("Too High")

# function definition
def main():
    # pick a random number
    target = random.randint(1,10)

    # get the user's guess
    num = int(input("Enter a number from 1 to 10 (inclusive)"))

    # get the result
    result = check_value(target, num)

    # loop while the guess is not correct
    while result != "Correct":

        # Tell the result and get a new number
        num = int(input(result + ".  Enter a number from 1 to 10"))

        # get the result
        result = check_value(target, num)

    # Tell the user the number
    print("You guessed it!  It was", target)

# function call
main()
</textarea>
</div>
</div>
<p>Modify the code above to keep a count of the number of guesses and print the number of guesses it took to guess the correct value.</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=func_ll_ac_blastoff data-question_label="5.33.3.4">
<div id=func_ll_ac_blastoff_question class="ac_question col-md-12">
<p>Fix the code below to print from start (inclusive) to 0 and then “Blastoff”</p>

</div>
<textarea data-lang="python" id="func_ll_ac_blastoff_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
         data-caption='Fixing a while loop'  data-wasm=/_static
     style="visibility: hidden;">
def countdown(start):
    value = start
    while value > 0:
        print(value)
    print(Blastoff)

def main():
    countdown(5)
</textarea>
</div>
</div>
<p>If you worked in a group, you can copy the answers from this page to the other group members.  Select the group members below and click the button to share the answers.</p>

<div class="runestone alert alert-warning sqcontainer">
    <div data-component="groupsub" id=func_loop_groupsub data-size_limit=4>
        <div class="col-sm-6">
            <select id="assignment_group" multiple class="assignment_partner_select" style="width: 100%">
            </select>
        </div>
        <div id="groupsub_button" class="col-sm-6">
        </div>
        <p>The Submit Group button will submit the answer for each each question
        on this page for each member of your group. It also logs you as the official
        group submitter.</p>
    </div>
</div>
</div>
</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 5.32. Functions and Lists Write Code Questions' data-toggle="tooltip">
  <a href="function_list_writecode.html" aria-label="Previous - 5.32. Functions and Lists Write Code Questions">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 5.34. Functions with Loops Multiple Choice Questions' data-toggle="tooltip" >
    <a href="functions_loops_exercises.html" aria-label="Next - 5.34. Functions with Loops Multiple Choice Questions">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>