<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Functions with Tuples and Dictionaries Mixed-Up Code Questions &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/637.css?v=bddd2011b20db76a20ea" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.css?v=7c5388ad4389eae1a318" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/runtime.bundle.js?v=aa7c717c38c5405af9e6"></script>
    <script src="../_static/637.bundle.js?v=1123ad75f9b821cdd636"></script>
    <script src="../_static/runestone.bundle.js?v=9c8e466e8cbed4a3af5a"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'py4e-int';
    eBookConfig.basecourse = 'py4e-int';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          py4e-int
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=functions_tups_dicts_mixedupcode_orig'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
  
<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="functions-with-tuples-and-dictionaries-mixed-up-code-questions">
<h1>Functions with Tuples and Dictionaries Mixed-Up Code Questions<a class="headerlink" href="#functions-with-tuples-and-dictionaries-mixed-up-code-questions" title="Permalink to this headline">¶</a></h1>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="Functions_Tups_Dictionaries_MixedUpCode_1_v2" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">addition_tup</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">lst</span></code> and <code class="docutils literal notranslate"><span class="pre">num</span></code> and returns a tuple that contains tuples with the first
value being a number from the <code class="docutils literal notranslate"><span class="pre">lst</span></code> and the second value being the number from the <code class="docutils literal notranslate"><span class="pre">lst</span></code> plus <code class="docutils literal notranslate"><span class="pre">num</span></code>. Round each value within the tuple to two decimal places.
For example, <code class="docutils literal notranslate"><span class="pre">addition_tup([1,2,3],</span> <span class="pre">5)</span></code> should return <code class="docutils literal notranslate"><span class="pre">((1,</span> <span class="pre">6),</span> <span class="pre">(2,</span> <span class="pre">7),</span> <span class="pre">(3,</span> <span class="pre">8))</span></code> and <code class="docutils literal notranslate"><span class="pre">addition_tup([-1,-2,</span> <span class="pre">-3,</span> <span class="pre">0,</span> <span class="pre">5],</span> <span class="pre">-5.2)</span></code> should return
<code class="docutils literal notranslate"><span class="pre">((-1,</span> <span class="pre">-6.2),</span> <span class="pre">(-2,</span> <span class="pre">-7.2),</span> <span class="pre">(-3,</span> <span class="pre">-8.2),</span> <span class="pre">(0,</span> <span class="pre">-5.2),</span> <span class="pre">(5,</span> <span class="pre">-0.2))</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="1"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def addition_tup(lst, num):
---
    lst_of_tups = []
---
    tup_of_tups = () #paired
---
    for i in lst:
---
        lst_of_tups.append((round(i,2), round(i + num,2)))
---
        lst_of_tups.append(round(i,2), round(i + num,2))) #paired
---
    return tuple(lst_of_tups)
---
    return lst_of_tups #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=Functions_Tups_Dictionaries_MixedUpCode_1_v2_ac data-question_label="2">
<div id=Functions_Tups_Dictionaries_MixedUpCode_1_v2_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">addition_tup</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">lst</span></code> and <code class="docutils literal notranslate"><span class="pre">num</span></code> and returns a list that contains tuples with the first
value being a number from the <code class="docutils literal notranslate"><span class="pre">lst</span></code> and the second value being the number from the <code class="docutils literal notranslate"><span class="pre">lst</span></code> plus <code class="docutils literal notranslate"><span class="pre">num</span></code>. Round each value within the tuple to two decimal places.
For example, <code class="docutils literal notranslate"><span class="pre">addition_tup([1,2,3],</span> <span class="pre">5)</span></code> should return <code class="docutils literal notranslate"><span class="pre">((1,</span> <span class="pre">6),</span> <span class="pre">(2,</span> <span class="pre">7),</span> <span class="pre">(3,</span> <span class="pre">8))</span></code> and <code class="docutils literal notranslate"><span class="pre">addition_tup([-1,-2,</span> <span class="pre">-3,</span> <span class="pre">0,</span> <span class="pre">5],</span> <span class="pre">-5.2)</span></code> should return
<code class="docutils literal notranslate"><span class="pre">((-1,</span> <span class="pre">-6.2),</span> <span class="pre">(-2,</span> <span class="pre">-7.2),</span> <span class="pre">(-3,</span> <span class="pre">-8.2),</span> <span class="pre">(0,</span> <span class="pre">-5.2),</span> <span class="pre">(5,</span> <span class="pre">-0.2))</span></code>.</p>

</div>
<textarea data-lang="python" id="Functions_Tups_Dictionaries_MixedUpCode_1_v2_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(addition_tup([1,2,3], 5),((1, 6), (2, 7), (3, 8)),"addition_tup([1,2,3], 5)")
        self.assertEqual(addition_tup([-1,-2, -3, 0, 5], -5.2),((-1, -6.2), (-2, -7.2), (-3, -8.2), (0, -5.2), (5, -0.2)),"addition_tup([-1,-2, -3, 0, 5], -5.2)")
        self.assertEqual(addition_tup([-1,-2.5, -3, 0.5, 5], 0),((-1, -1), (-2.5, -2.5), (-3, -3), (0.5, 0.5), (5, 5)),"addition_tup([-1,-2.5, -3, 0.5, 5], 0)")


myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="Functions_Tups_Dictionaries_MixedUpCode_2" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">greater_than_five</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">dictionary</span></code> as a parameter and returns a list of keys that
correspond to values in the <code class="docutils literal notranslate"><span class="pre">dictionary</span></code> that are greater than 5.  For example, <code class="docutils literal notranslate"><span class="pre">greater_than_five({'a':1,</span> <span class="pre">'b':</span> <span class="pre">2,</span> <span class="pre">'c':7})</span></code>
should return <code class="docutils literal notranslate"><span class="pre">['c']</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="3"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def greater_than_five(dictionary):
---
    lst = []
---
    lst = () #paired
---
    for i in dictionary.items():
---
    for i in dictionary.keys(): #paired
---
        if i[1] > 5:
---
        if i[0] > 5: #paired
---
            lst.append(i[0])
---
            lst.append(i[1]) #paired
---
    return lst
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=Functions_Tups_Dictionaries_MixedUpCode_2_ac data-question_label="4">
<div id=Functions_Tups_Dictionaries_MixedUpCode_2_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">greater_than_five</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">dictionary</span></code> as a parameter and returns a list of keys that
correspond to values in the <code class="docutils literal notranslate"><span class="pre">dictionary</span></code> that are greater than 5. For example, <code class="docutils literal notranslate"><span class="pre">greater_than_five({'a':1,</span> <span class="pre">'b':</span> <span class="pre">2,</span> <span class="pre">'c':7})</span></code>
should return <code class="docutils literal notranslate"><span class="pre">['c']</span></code>.</p>

</div>
<textarea data-lang="python" id="Functions_Tups_Dictionaries_MixedUpCode_2_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(greater_than_five({'a':1, 'b': 2, 'c':7}),['c'],"greater_than_five({'a':1, 'b': 2, 'c':7})")
        self.assertEqual(greater_than_five({'a':1, 'b': 2, 'c':0}),[],"greater_than_five({'a':1, 'b': 2, 'c':0})")
        self.assertEqual(greater_than_five({7:6, "hi":6, 10:6}),[7, 'hi', 10],"greater_than_five({7:6, 'hi': 6, 10:6})")
        self.assertEqual(greater_than_five({-7:-5, "hi":5.2, 10:5}),['hi'],"greater_than_five({-7:-5, 'hi': 5.2, 10:5})")


myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="Functions_Tups_Dictionaries_MixedUpCode_3" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">dict_to_tup_of_tuples</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">dictionary</span></code> as a parameter and returns a tuple that contains
tuples with the first value being the keys of the <code class="docutils literal notranslate"><span class="pre">dictionary</span></code> and the second being the values of the <code class="docutils literal notranslate"><span class="pre">dictionary</span></code>.  For example,
<code class="docutils literal notranslate"><span class="pre">dict_to_tup_of_tuples({'a':1,</span> <span class="pre">'b':</span> <span class="pre">2,</span> <span class="pre">'c':7})</span></code> would return <code class="docutils literal notranslate"><span class="pre">(('a',</span> <span class="pre">1),</span> <span class="pre">('b',</span> <span class="pre">2),</span> <span class="pre">('c',</span> <span class="pre">7))</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="5"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def dict_to_tup_of_tuples(dictionary):
---
    lst1 = []
---
    lst2 = []
---
    for item in dictionary.items():
---
    for item in dictionary: #paired
---
        lst1.append(item[0])
---
        lst2.append(item[1])
---
    zipped_lst = zip(lst1,lst2)
---
    zipped_lst = zip(lst2,lst1) #paired
---
    return tuple(zipped_lst)
---
    return zipped_lst #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=Functions_Tups_Dictionaries_MixedUpCode_3_ac data-question_label="6">
<div id=Functions_Tups_Dictionaries_MixedUpCode_3_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">dict_to_tup_of_tuples</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">dictionary</span></code> as a parameter and returns a tuple that contains
tuples with the first value being the keys of the <code class="docutils literal notranslate"><span class="pre">dictionary</span></code> and the second being the values of the <code class="docutils literal notranslate"><span class="pre">dictionary</span></code>. For example,
<code class="docutils literal notranslate"><span class="pre">dict_to_tup_of_tuples({'a':1,</span> <span class="pre">'b':</span> <span class="pre">2,</span> <span class="pre">'c':7})</span></code> would return <code class="docutils literal notranslate"><span class="pre">(('a',</span> <span class="pre">1),</span> <span class="pre">('b',</span> <span class="pre">2),</span> <span class="pre">('c',</span> <span class="pre">7))</span></code>.</p>

</div>
<textarea data-lang="python" id="Functions_Tups_Dictionaries_MixedUpCode_3_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(dict_to_tup_of_tuples({'a':1, 'b': 2, 'c':7}),(('a', 1), ('b', 2), ('c', 7)),"dict_to_tup_of_tuples({'a':1, 'b': 2, 'c':7})")
        self.assertEqual(dict_to_tup_of_tuples({'hi':0, 'b': 2.2, 'c':-7}),(('hi', 0), ('b', 2.2), ('c', -7)),"dict_to_tup_of_tuples({'hi':0, 'b': 2.2, 'c':-7})")
        self.assertEqual(dict_to_tup_of_tuples({2:1, -4: 2, 3.2:7}),((2, 1), (-4, 2), (3.2, 7)),"dict_to_tup_of_tuples({2:1, -4: 2, 3.2:7})")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="Functions_Tups_Dictionaries_MixedUpCode_4" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">swap_keys_values</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">dictionary</span></code> as a parameter and returns <code class="docutils literal notranslate"><span class="pre">swapped_dictionary</span></code> that
swaps the keys and values of the passed in <code class="docutils literal notranslate"><span class="pre">dictionary</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">swap_keys_values({'a':1,</span> <span class="pre">'b':</span> <span class="pre">2,</span> <span class="pre">'c':7})</span></code> would return
<code class="docutils literal notranslate"><span class="pre">{1:</span> <span class="pre">'a',</span> <span class="pre">2:</span> <span class="pre">'b',</span> <span class="pre">7:</span> <span class="pre">'c'}</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="7"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def swap_keys_values(dictionary):
---
    swapped_dictionary = {}
---
    for item in dictionary.items():
---
    for item in dictionary(): #paired
---
        swapped_dictionary[item[1]] = item[0]
---
        swapped_dictionary[item[0]] = item[1] #paired
---
        swapped_dictionary[item[1]].append(item[0]) #paired
---
    return swapped_dictionary
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=Functions_Tups_Dictionaries_MixedUpCode_4_ac data-question_label="8">
<div id=Functions_Tups_Dictionaries_MixedUpCode_4_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">swap_keys_values</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">dictionary</span></code> as a parameter and returns <code class="docutils literal notranslate"><span class="pre">swapped_dictionary</span></code> that
swaps the keys and values of the passed in <code class="docutils literal notranslate"><span class="pre">dictionary</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">swap_keys_values({'a':1,</span> <span class="pre">'b':</span> <span class="pre">2,</span> <span class="pre">'c':7})</span></code> would return
<code class="docutils literal notranslate"><span class="pre">{1:</span> <span class="pre">'a',</span> <span class="pre">2:</span> <span class="pre">'b',</span> <span class="pre">7:</span> <span class="pre">'c'}</span></code>.</p>

</div>
<textarea data-lang="python" id="Functions_Tups_Dictionaries_MixedUpCode_4_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(swap_keys_values({'a':1, 'b': 2, 'c':7}),{1: 'a', 2: 'b', 7: 'c'},"swap_keys_values({'a':1, 'b': 2, 'c':7})")
        self.assertEqual(swap_keys_values({'hi':0, 'b': 2.2, 'c':-7}),{0: 'hi', 2.2: 'b', -7: 'c'},"swap_keys_values({'hi':0, 'b': 2.2, 'c':-7})")
        self.assertEqual(swap_keys_values({2:1, -4: 2, 3.2:7}),{1: 2, 2: -4, 7: 3.2},"swap_keys_values({2:1, -4: 2, 3.2:7})")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="Functions_Tups_Dictionaries_MixedUpCode_5" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">common_value_in_tups</span></code> that takes in two tuples, <code class="docutils literal notranslate"><span class="pre">tup1</span></code> and <code class="docutils literal notranslate"><span class="pre">tup2</span></code>, as parameters and
returns <code class="docutils literal notranslate"><span class="pre">True</span></code> if any of the values in both tuples match and <code class="docutils literal notranslate"><span class="pre">False</span></code> otherwise. (Note: Disregard position of values.)  For example,
<code class="docutils literal notranslate"><span class="pre">common_value_in_tups((1,2,3),(6,4,5))</span></code> should return <code class="docutils literal notranslate"><span class="pre">False</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="9"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def common_value_in_tups(tup1, tup2):
---
    for value in tup1:
---
    if tup1 == tup2: #paired
---
        if value in tup2:
---
        for value2 in tup2: #paired
---
            return True
---
    return False
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=Functions_Tups_Dictionaries_MixedUpCode_5_ac data-question_label="10">
<div id=Functions_Tups_Dictionaries_MixedUpCode_5_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">common_value_in_tups</span></code> that takes in two tuples, <code class="docutils literal notranslate"><span class="pre">tup1</span></code> and <code class="docutils literal notranslate"><span class="pre">tup2</span></code>, as parameters and
returns <code class="docutils literal notranslate"><span class="pre">True</span></code> if any of the values in both tuples match and <code class="docutils literal notranslate"><span class="pre">False</span></code> otherwise. (Note: Disregard position of values.) For example,
<code class="docutils literal notranslate"><span class="pre">common_value_in_tups((1,2,3),(6,4,5))</span></code> should return <code class="docutils literal notranslate"><span class="pre">False</span></code>.</p>

</div>
<textarea data-lang="python" id="Functions_Tups_Dictionaries_MixedUpCode_5_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(common_value_in_tups((1,2,3),(6,4,5)),False,"common_value_in_tups((1,2,3),(6,4,5)")
        self.assertEqual(common_value_in_tups((0,0,0),(0,0,0)),True,"common_value_in_tups((0,0,0),(0,0,0))")
        self.assertEqual(common_value_in_tups((-1,0,2),(5,10.2,0)),True,"common_value_in_tups((-1,0,2),(5,10.2,0))")
        self.assertEqual(common_value_in_tups((1,2,3.2),(3.2,4,5)),True,"common_value_in_tups((1,2,3.2),(3.2,4,5))")

myTests().main()
</textarea>
</div>
</div>
</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>