<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Mixed-up Code Questions &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/637.css?v=bddd2011b20db76a20ea" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.css?v=7c5388ad4389eae1a318" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/runtime.bundle.js?v=aa7c717c38c5405af9e6"></script>
    <script src="../_static/637.bundle.js?v=1123ad75f9b821cdd636"></script>
    <script src="../_static/runestone.bundle.js?v=9c8e466e8cbed4a3af5a"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'py4e-int';
    eBookConfig.basecourse = 'py4e-int';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          py4e-int
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=mixedupcode_orig'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
  
<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="mixed-up-code-questions">
<h1>Mixed-up Code Questions<a class="headerlink" href="#mixed-up-code-questions" title="Permalink to this headline">¶</a></h1>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="functEx1muc_v2" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">len_str</span></code> that uses a built-in function to return the number of characters in the passed string <code class="docutils literal notranslate"><span class="pre">str</span></code>. Then print the result of a call to <code class="docutils literal notranslate"><span class="pre">len_str</span></code> passing in the string <code class="docutils literal notranslate"><span class="pre">I</span> <span class="pre">love</span> <span class="pre">Python!</span></code>.  It should return 14.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="1"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def len_str(str):
---
def len_str(str) #paired
---
    return len(str)
---
    Return len(str) #paired
---
print(len_str("I love Python!"))
---
print(len_str(I love Python!)) #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=functEx1muc_v2_ac data-question_label="2">
<div id=functEx1muc_v2_ac_question class="ac_question col-md-12">
<p>Write  a function <code class="docutils literal notranslate"><span class="pre">len_str</span></code> that uses a built-in function to return the number of characters in the passed string <code class="docutils literal notranslate"><span class="pre">str</span></code>. Then print the result of a call to <code class="docutils literal notranslate"><span class="pre">len_str</span></code> passing in the string <code class="docutils literal notranslate"><span class="pre">I</span> <span class="pre">love</span> <span class="pre">Python!</span></code>.  It should return 14.</p>

</div>
<textarea data-lang="python" id="functEx1muc_v2_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(len_str(""),0,"len_str('')")
        self.assertEqual(len_str("a"),1,"len_str('a')")
        self.assertEqual(len_str("Hi"),2,"len_str('Hi')")
        self.assertEqual(len_str("I love Python!"),14,"len_str('I love Python!')")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="functEx2muc" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">bmi</span></code> that takes <code class="docutils literal notranslate"><span class="pre">height</span></code> (in inches) and <code class="docutils literal notranslate"><span class="pre">weight</span></code> (in pounds) as parameters. It should calculate the ‘’bmi’’ by squaring the height then dividing the weight by the height squared and then converting to metric by multiplying by 703.
Watch your indentation!</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="3"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def bmi(height, weight):
---
    heightSquared = height * height
---
    BMI = weight / heightSquared
---
    BMImetric = BMI * 703
---
    return BMImetric
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=functEx2muc_ac data-question_label="4">
<div id=functEx2muc_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">bmi</span></code> that takes <code class="docutils literal notranslate"><span class="pre">height</span></code> (in inches) and <code class="docutils literal notranslate"><span class="pre">weight</span></code> (in pounds) as parameters. It should calculate the ‘’bmi’’ by squaring the height then dividing the weight by the height squared and then converting to metric by multiplying by 703.</p>

</div>
<textarea data-lang="python" id="functEx2muc_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def bmi(height, weight):

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertAlmostEqual(bmi(60,140),27.33888888888889,1,"bmi(60,140)")
        self.assertAlmostEqual(bmi(72,164),22.23996913580247,1,"bmi(72,164)")
        self.assertAlmostEqual(bmi(32,100),68.65234375,1,"bmi(32,100)")
        self.assertAlmostEqual(bmi(84,230),22.91524943310657,1,"bmi(84,230)")
        self.assertAlmostEqual(bmi(24,89),108.6232638888889,1,"bmi(24,89)")


myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="functEx3muc" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Write two functions. The first is <code class="docutils literal notranslate"><span class="pre">square(x)</span></code>, which returns <code class="docutils literal notranslate"><span class="pre">x</span></code> squared. The second function is <code class="docutils literal notranslate"><span class="pre">sum_of_squares(x,y,z)</span></code>, which returns the sum of the squares of three numbers <code class="docutils literal notranslate"><span class="pre">z</span></code>, <code class="docutils literal notranslate"><span class="pre">y</span></code>, and <code class="docutils literal notranslate"><span class="pre">z</span></code>.  For example, <code class="docutils literal notranslate"><span class="pre">square(2)</span></code> should return <code class="docutils literal notranslate"><span class="pre">4</span></code> and <code class="docutils literal notranslate"><span class="pre">sum_of_squares(2,</span> <span class="pre">4,</span> <span class="pre">6)</span></code> should return 56.
For example, <code class="docutils literal notranslate"><span class="pre">square(2)</span></code> should return <code class="docutils literal notranslate"><span class="pre">4</span></code> and <code class="docutils literal notranslate"><span class="pre">sum_of_squares(2,</span> <span class="pre">4,</span> <span class="pre">6)</span></code> should return 56.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="5"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def square(x):
---
    y = x * x
---
    return y
---
def sum_of_squares(x, y, z):
---
    a = square(x)
    b = square(y)
    c = square(z)
---
    return a + b + c
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=functEx3muc_ac data-question_label="6">
<div id=functEx3muc_ac_question class="ac_question col-md-12">
<p>Write two functions. The first is <code class="docutils literal notranslate"><span class="pre">square(x)</span></code>, which returns <code class="docutils literal notranslate"><span class="pre">x</span></code> squared. The second function is <code class="docutils literal notranslate"><span class="pre">sum_of_squares(x,y,z)</span></code>, which returns the sum of the squares of three numbers <code class="docutils literal notranslate"><span class="pre">z</span></code>, <code class="docutils literal notranslate"><span class="pre">y</span></code>, and <code class="docutils literal notranslate"><span class="pre">z</span></code>.  For example, <code class="docutils literal notranslate"><span class="pre">square(2)</span></code> should return <code class="docutils literal notranslate"><span class="pre">4</span></code> and <code class="docutils literal notranslate"><span class="pre">sum_of_squares(2,</span> <span class="pre">4,</span> <span class="pre">6)</span></code> should return 56.</p>

</div>
<textarea data-lang="python" id="functEx3muc_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def square(x):

def sum_of_squares(x, y, z):

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(square(2), 4, "square(2)")
        self.assertEqual(square(-3), 9, "square(-3)")
        self.assertEqual(sum_of_squares(2, 4, 6), 56, "sum_of_squares(2, 4, 6)")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="functEx4muc_v5" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">average(num1,</span> <span class="pre">num2)</span></code>, which finds the average of two numbers <code class="docutils literal notranslate"><span class="pre">num1</span></code> and <code class="docutils literal notranslate"><span class="pre">num2</span></code>.  For example, <code class="docutils literal notranslate"><span class="pre">average(10,4)</span></code> should return <code class="docutils literal notranslate"><span class="pre">7.0</span></code>.
Note: there are two extra code blocks, and watch your indentation!</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="7"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def average(num1, num2):
---
def average(num1, num2) #paired
---
    sum = num1 + num2
---
    sum = sum / 2
---
    sum = sum // 2 #paired
---
    return sum
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=functEx4muc_v5_ac data-question_label="8">
<div id=functEx4muc_v5_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">average(num1,</span> <span class="pre">num2)</span></code>, which finds the average of two numbers <code class="docutils literal notranslate"><span class="pre">num1</span></code> and <code class="docutils literal notranslate"><span class="pre">num2</span></code>.  For example, <code class="docutils literal notranslate"><span class="pre">average(10,4)</span></code> should return <code class="docutils literal notranslate"><span class="pre">7.0</span></code>.</p>

</div>
<textarea data-lang="python" id="functEx4muc_v5_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def average(num1, num2):

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertAlmostEqual(average(32,3),17.5,1, "average(32,3)")
        self.assertAlmostEqual(average(5,20),12.5,1, "average(5,20)")
        self.assertAlmostEqual(average(10,4),7,1, "average(10,4)")
        self.assertAlmostEqual(average(8,120),64,1, "average(8,120)")
        self.assertAlmostEqual(average(3,19),11,1, "average(3,19)")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="functEx5muc_v2" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">swapValues(val1,</span> <span class="pre">val2)</span></code>, which takes two values and swaps them.
It then returns val1. For example, <code class="docutils literal notranslate"><span class="pre">swapValues(8,2)</span></code> should return <code class="docutils literal notranslate"><span class="pre">2</span></code>.  Note: there is an extra code block, and
watch your indentation!</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="9"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def swapValues(val1, val2):
---
    val1 = val2 #distractor
    val2 = val1
---
    temp = val2
---
    val2 = val1
---
    val1 = temp
---
    return val1
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=functEx5muc_v2_ac data-question_label="10">
<div id=functEx5muc_v2_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">swapValues(val1,</span> <span class="pre">val2)</span></code> , which takes two values and swaps them. It then returns val1.  For example, <code class="docutils literal notranslate"><span class="pre">swapValues(8,2)</span></code> should return <code class="docutils literal notranslate"><span class="pre">2</span></code></p>

</div>
<textarea data-lang="python" id="functEx5muc_v2_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def swapValues(val1, val2):

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(swapValues(5,20), 20, "swapValues(5, 20)")
        self.assertEqual(swapValues(8,2), 2, "swapValues(8, 2)")
        self.assertEqual(swapValues(2,8), 8, "swapValues(2, 8)")
        self.assertEqual(swapValues(-1,3), 3, "swapValues(-1, 3)")
        self.assertEqual(swapValues(5,5), 5, "swapValues(5, 5)")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="functEx6muc" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">get_avg_drop_lowest(num_list)</span></code> which returns the average of a list of numbers, ignoring the lowest value. However, if the list only contains one value, then return that. For example, <code class="docutils literal notranslate"><span class="pre">get_avg_drop_lowest([90])</span></code> returns <code class="docutils literal notranslate"><span class="pre">90</span></code>, <code class="docutils literal notranslate"><span class="pre">get_avg_drop_lowest([90,</span> <span class="pre">10])</span></code> also returns <code class="docutils literal notranslate"><span class="pre">90</span></code>, and <code class="docutils literal notranslate"><span class="pre">get_avg_drop_lowest([90,</span> <span class="pre">10,</span> <span class="pre">0])</span></code> returns <code class="docutils literal notranslate"><span class="pre">50</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="11"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def get_avg_drop_lowest(num_list):
---
def get_avg_drop_lowest(num_list) #paired
---
    total = sum(num_list)
    lowest = min(num_list)
    num = len(num_list)
---
    if num == 1:
---
    if num = 1: #paired
---
        return total
---
    else:
---
        return (total - lowest) / (num - 1)
---
        return (total - lowest) / (num + 1) #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=functEx6muc_ac data-question_label="12">
<div id=functEx6muc_ac_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="functEx6muc_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
 Write a function called ``get_avg_drop_lowest(num_list)`` which returns the average of a list of numbers, ignoring the lowest value. However, if the list only contains one value, then return that. For example, ``get_avg_drop_lowest([90])`` returns ``90``, ``get_avg_drop_lowest([90, 10])`` also returns ``90``, and ``get_avg_drop_lowest([90, 10, 0])`` returns ``50``.
 ~~~~
 def get_avg_drop_lowest(num_list):


 ====
 from unittest.gui import TestCaseGui

 class myTests(TestCaseGui):

     def testOne(self):
         self.assertEqual(get_avg_drop_lowest([40]), 40, "get_avg_drop_lowest([40])")
         self.assertEqual(get_avg_drop_lowest([80,10]), 80, "get_avg_drop_lowest([80,10])")
         self.assertEqual(get_avg_drop_lowest([74,25,6]), 49.5, "get_avg_drop_lowest([74,25,6])")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="functEx7muc_v2" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">compare(p1,</span> <span class="pre">p2)</span></code> which returns whichever is greater of its two parameters. Note:
there are two unused code blocks.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="13"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def compare(p1, p2):
---
Def compare(p1, p2): #paired
---
    if p1 > p2:
---
    if p1 == p2: #paired
---
        return p1
---
    else:
---
        return p2
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=functEx7muc_v2_ac data-question_label="14">
<div id=functEx7muc_v2_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">compare(p1,</span> <span class="pre">p2)</span></code> which returns whichever is greater of its two parameters.</p>

</div>
<textarea data-lang="python" id="functEx7muc_v2_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def compare(p1, p2):


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(compare(20,15), 20, "compare(20,15)")
        self.assertEqual(compare(8,2), 8, "compare(8,2)")
        self.assertEqual(compare(2,8), 8, "compare(2,8)")
        self.assertEqual(compare(-1,1), 1, "compare(-1,1)")
        self.assertEqual(compare(-1,-1), -1, "compare(-1,-1)")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="functEx8muc" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>You are driving a little too fast, and a police officer stops you. Create a function
called <code class="docutils literal notranslate"><span class="pre">caught_speeding(speed,</span> <span class="pre">is_birthday)</span></code> which returns the type of ticket the police officer will give you.
If <code class="docutils literal notranslate"><span class="pre">speed</span></code> is 60 or less, the result is <code class="docutils literal notranslate"><span class="pre">&quot;no</span> <span class="pre">ticket&quot;</span></code>. If speed is between 61 and 80 inclusive,
the result is <code class="docutils literal notranslate"><span class="pre">&quot;minor</span> <span class="pre">ticket&quot;</span></code>. If speed is 81 or more, the result is <code class="docutils literal notranslate"><span class="pre">&quot;major</span> <span class="pre">ticket&quot;</span></code>. All this
is true, unless it is your birthday – on that day, your speed can be 5 higher in all cases.
Note: there are two extra code blocks and lots of indentation to watch out for!</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="15"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def caught_speeding(speed, is_birthday):
---
    if is_birthday:
---
        if speed <= 65:
---
        if speed < 65: #distractor
---
            return "no ticket"
---
        elif speed <= 85:
---
            return "minor ticket"
---
            return minor ticket #distractor
---
        else:
---
            return "major ticket"
---
    else:
---
        if speed <= 60:
---
            return "no ticket"
---
        elif speed <= 80:
---
            return "minor ticket"
---
        else:
---
            return "major ticket"
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=functEx8muc_ac data-question_label="16">
<div id=functEx8muc_ac_question class="ac_question col-md-12">
<p>You are driving a little too fast, and a police officer stops you. Write a function called <code class="docutils literal notranslate"><span class="pre">caught_speeding(speed,</span> <span class="pre">is_birthday)</span></code> which returns the type of ticket the police officer will give you. If <code class="docutils literal notranslate"><span class="pre">speed</span></code> is 60 or less, the result is <code class="docutils literal notranslate"><span class="pre">&quot;no</span> <span class="pre">ticket&quot;</span></code>. If speed is between 61 and 80 inclusive, the result is <code class="docutils literal notranslate"><span class="pre">&quot;minor</span> <span class="pre">ticket&quot;</span></code>. If speed is 81 or more, the result is <code class="docutils literal notranslate"><span class="pre">&quot;major</span> <span class="pre">ticket&quot;</span></code>. All this is true, unless it is your birthday – on that day, your speed can be 5 higher in all cases.</p>

</div>
<textarea data-lang="python" id="functEx8muc_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def caught_speeding(speed, is_birthday):

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(caught_speeding(50, False), "no ticket", "caught_speeding(50, False)")
        self.assertEqual(caught_speeding(60, False), "no ticket", "caught_speeding(60, False)")
        self.assertEqual(caught_speeding(65, True), "no ticket", "caught_speeding(65, True)")
        self.assertEqual(caught_speeding(61,False), "minor ticket", "caught_speeding(61,False)")
        self.assertEqual(caught_speeding(73,False), "minor ticket", "caught_speeding(73,False)")
        self.assertEqual(caught_speeding(85,True), "minor ticket", "caught_speeding(85,True)")
        self.assertEqual(caught_speeding(81,False), "major ticket", "caught_speeding(81,False)")
        self.assertEqual(caught_speeding(97,True), "major ticket", "caught_speeding(97,True)")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="functEx9muc" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a <code class="docutils literal notranslate"><span class="pre">check_guess(guess,</span> <span class="pre">target)</span></code> function which computes if a <code class="docutils literal notranslate"><span class="pre">guess</span></code> is too low, too high, or correct compared to the <code class="docutils literal notranslate"><span class="pre">target</span></code>.
Return <code class="docutils literal notranslate"><span class="pre">'too</span> <span class="pre">low'</span></code> if <code class="docutils literal notranslate"><span class="pre">guess</span></code> is less than <code class="docutils literal notranslate"><span class="pre">target</span></code>, <code class="docutils literal notranslate"><span class="pre">'correct'</span></code> if they are equal, and
<code class="docutils literal notranslate"><span class="pre">'too</span> <span class="pre">high'</span></code> if <code class="docutils literal notranslate"><span class="pre">guess</span></code> is greater than <code class="docutils literal notranslate"><span class="pre">target</span></code>.  For example, <code class="docutils literal notranslate"><span class="pre">check_guess(5,</span> <span class="pre">7)</span></code> returns
<code class="docutils literal notranslate"><span class="pre">'too</span> <span class="pre">low'</span></code>, <code class="docutils literal notranslate"><span class="pre">check_guess(7,</span> <span class="pre">7)</span></code> returns <code class="docutils literal notranslate"><span class="pre">'correct'</span></code>, and <code class="docutils literal notranslate"><span class="pre">check_guess(9,</span> <span class="pre">7)</span></code> returns <code class="docutils literal notranslate"><span class="pre">'too</span> <span class="pre">high'</span></code>.
Note: there are three extra code blocks, and watch your indentation!</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="17"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def check_guess(guess, target):
---
def check_guess(guess), (target): #paired
---
    if guess < target:
---
        return 'too low'
---
        return "too low' #paired
---
    elif guess == target:
---
        return 'correct'
---
        return correct #paired
---
    else:
---
        return 'too high'
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=functEx9muc_ac data-question_label="18">
<div id=functEx9muc_ac_question class="ac_question col-md-12">
<p>Write  a <code class="docutils literal notranslate"><span class="pre">check_guess(guess,</span> <span class="pre">target)</span></code> function which computes if a <code class="docutils literal notranslate"><span class="pre">guess</span></code> is too low, too high, or correct compared to the <code class="docutils literal notranslate"><span class="pre">target</span></code>.
Return <code class="docutils literal notranslate"><span class="pre">'too</span> <span class="pre">low'</span></code> if <code class="docutils literal notranslate"><span class="pre">guess</span></code> is less than <code class="docutils literal notranslate"><span class="pre">target</span></code>, <code class="docutils literal notranslate"><span class="pre">'correct'</span></code> if they are equal, and
<code class="docutils literal notranslate"><span class="pre">'too</span> <span class="pre">high'</span></code> if <code class="docutils literal notranslate"><span class="pre">guess</span></code> is greater than <code class="docutils literal notranslate"><span class="pre">target</span></code>.  For example, <code class="docutils literal notranslate"><span class="pre">check_guess(5,</span> <span class="pre">7)</span></code> returns
<code class="docutils literal notranslate"><span class="pre">'too</span> <span class="pre">low'</span></code>, <code class="docutils literal notranslate"><span class="pre">check_guess(7,</span> <span class="pre">7)</span></code> returns <code class="docutils literal notranslate"><span class="pre">'correct'</span></code>, and <code class="docutils literal notranslate"><span class="pre">check_guess(9,</span> <span class="pre">7)</span></code> returns <code class="docutils literal notranslate"><span class="pre">'too</span> <span class="pre">high'</span></code>.</p>

</div>
<textarea data-lang="python" id="functEx9muc_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def check_guess(guess, target):

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(check_guess(4, 6), 'too low', "check_guess(4,6)")
        self.assertEqual(check_guess(3,3), 'correct', "check_guess(3,3)")
        self.assertEqual(check_guess(7,5), 'too high', "check_guess(7,5)")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="functEx10muc" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Put the code blocks below to define the function <code class="docutils literal notranslate"><span class="pre">alarm_clock</span></code>. It will be given a day of the week encoded as 0 = Sun, 1 = Mon, 2 = Tue, …6 = Sat, and a
boolean indicating if we are on vacation, and will return a string indicating when the alarm
clock should ring. If we are on vacation and it is a weekend (0 = Saturday or 6 =
Sunday), it should return <code class="docutils literal notranslate"><span class="pre">&quot;off&quot;</span></code>, and otherwise return <code class="docutils literal notranslate"><span class="pre">&quot;10:00&quot;</span></code>. If we are not
on vacation and it is a weekend, it should return <code class="docutils literal notranslate"><span class="pre">&quot;10:00&quot;</span></code>, and otherwise return
<code class="docutils literal notranslate"><span class="pre">&quot;7:00&quot;</span></code>. Note: there are two extra code blocks, and watch your indentation!</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="19"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def alarm_clock(day, vacation):
---
def alarm_clock(day, vacation) #paired
---
    if vacation:
---
        if day == 0 or day == 6:
---
        if (day == 0) || (day == 6): #paired
---
            return 'off'
---
        else:
---
            return '10:00'
---
    else:
---
        if day == 0 or day == 6:
---
            return '10:00'
---
        else:
---
            return '7:00'
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=functEx10muc_ac data-question_label="20">
<div id=functEx10muc_ac_question class="ac_question col-md-12">
<p>Write the function <code class="docutils literal notranslate"><span class="pre">alarm_clock</span></code>. It will be given a day of the week encoded as 0 = Sun, 1 = Mon, 2 = Tue, …6 = Sat, and a boolean indicating if we are on vacation, and will return a string indicating when the alarm clock should ring. If we are on vacation and it is a weekend (0 = Saturday or 6 = Sunday), it should return <code class="docutils literal notranslate"><span class="pre">&quot;off&quot;</span></code>, and otherwise return <code class="docutils literal notranslate"><span class="pre">&quot;10:00&quot;</span></code>. If we are not on vacation and it is a weekend, it should return <code class="docutils literal notranslate"><span class="pre">&quot;10:00&quot;</span></code>, and otherwise return <code class="docutils literal notranslate"><span class="pre">&quot;7:00&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="functEx10muc_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def alarm_clock(day, vacation):

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(alarm_clock(2,False), '7:00', "alarm_clock(2,False)")
        self.assertEqual(alarm_clock(0,True), 'off', "alarm_clock(0,True)")
        self.assertEqual(alarm_clock(6,True), 'off', "alarm_clock(6,True)")
        self.assertEqual(alarm_clock(6,False), '10:00', "alarm_clock(6,False)")
        self.assertEqual(alarm_clock(0,False), '10:00', "alarm_clock(6,False)")
        self.assertEqual(alarm_clock(1,False), '7:00', "alarm_clock(6,False)")
        self.assertEqual(alarm_clock(1,True), '10:00', "alarm_clock(6,False)")
        self.assertEqual(alarm_clock(3,True), '10:00', "alarm_clock(6,False)")
        self.assertEqual(alarm_clock(5,True), '10:00', "alarm_clock(6,False)")



myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="functEx11muc" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>First create a function called <code class="docutils literal notranslate"><span class="pre">square_it</span></code> which
squares the parameter <code class="docutils literal notranslate"><span class="pre">n</span></code> and returns the result. Then, create a function called <code class="docutils literal notranslate"><span class="pre">cube_it</span></code> which cubes
the parameter <code class="docutils literal notranslate"><span class="pre">n</span></code> and returns the result. Note : there are three extra code blocks, and watch your indentation!</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="21"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def square_it(n):
---
    return n * n
---
    return n * 2 #distractor
---
def cube_it(n):
---
def cube_it(n) #distractor
---
    return n * n * n
---
    return n * n * 3 #distractor
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=functEx11muc_ac data-question_label="22">
<div id=functEx11muc_ac_question class="ac_question col-md-12">
<p>First write a function called  <code class="docutils literal notranslate"><span class="pre">square_it</span></code> which squares the parameter <code class="docutils literal notranslate"><span class="pre">n</span></code> and returns the result. Then, write a function called <code class="docutils literal notranslate"><span class="pre">cube_it</span></code> which cubes the parameter <code class="docutils literal notranslate"><span class="pre">n</span></code> and returns the result.</p>

</div>
<textarea data-lang="python" id="functEx11muc_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(square_it(4),16, "square_it(4)")
        self.assertEqual(square_it(2),4, "square_it(2)")
        self.assertEqual(cube_it(5),125, "cube_it(5)")
        self.assertEqual(cube_it(3),27, "cube_it(3)")



myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="functEx12muc_v2" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">distance</span></code> which returns the distance between two coordinates using
the distance formula: d = √((x_2 - x_1)² + (y_2 - y_1)²).  Use two functions in Python’s math module (<code class="docutils literal notranslate"><span class="pre">math.pow</span></code>, <code class="docutils literal notranslate"><span class="pre">math.sqrt</span></code>).
The function <code class="docutils literal notranslate"><span class="pre">math.pow(a,b)</span></code> returns <code class="docutils literal notranslate"><span class="pre">a</span></code> raised to the <code class="docutils literal notranslate"><span class="pre">b</span></code> power.  The function <code class="docutils literal notranslate"><span class="pre">math.sqrt(a)</span></code> returns the square root of <code class="docutils literal notranslate"><span class="pre">a</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="23"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        import math
---
def distance(x1, y1, x2, y2):
---
def distance(x1, y1, x2, y2) #paired
---
    dx = x2 - x1
    dy = y2 - y1
---
    dsquared = math.pow(dx, 2) + math.pow(dy, 2)
---
    dsquared = math.pow(2, dx) + math.pow(2, dy) #paired
---
    result = math.sqrt(dsquared)
---
    result = math.square_root(dsquared) #paired
---
    return result
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=functEx12muc_v2_ac data-question_label="24">
<div id=functEx12muc_v2_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">distance</span></code> which returns the distance between two coordinates using
the distance formula: d = √((x_2 - x_1)² + (y_2 - y_1)²).  Use two functions in Python’s math module (<code class="docutils literal notranslate"><span class="pre">math.pow</span></code>, <code class="docutils literal notranslate"><span class="pre">math.sqrt</span></code>).
The function <code class="docutils literal notranslate"><span class="pre">math.pow(a,b)</span></code> returns <code class="docutils literal notranslate"><span class="pre">a</span></code> raised to the <code class="docutils literal notranslate"><span class="pre">b</span></code> power.  The function <code class="docutils literal notranslate"><span class="pre">math.sqrt(a)</span></code> returns the square root of <code class="docutils literal notranslate"><span class="pre">a</span></code>.</p>

</div>
<textarea data-lang="python" id="functEx12muc_v2_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
import math

def distance(x1, y1, x2, y2):


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertAlmostEqual(distance(1,1,2,2), 1.4142135624, 1, "distance(1, 1, 2, 2)")
        self.assertAlmostEqual(distance(3,2,0,0), 3.6055, 1, "distance(3, 2, 0, 0)")
        self.assertAlmostEqual(distance(3,2,-1,-1), 5.0, 1, "distance(3, 2, -1, -1)")
        self.assertAlmostEqual(distance(30,20,-1,-1), 37.44, 1, "distance(30, 20, -1, -1)")
        self.assertAlmostEqual(distance(2,5,3,2), 3.162, 1, "distance(2, 5, 3, 2)")


myTests().main()
</textarea>
</div>
</div>
</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>