<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <title>4.15. Write Code Questions &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=DB047029" />
    <link rel="stylesheet" type="text/css" href="https://cdn.jsdelivr.net/npm/handsontable@7.2.2/dist/handsontable.full.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-3.4.1/css/bootstrap.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/presenter_mode.css?v=9FF1320F" />
    <link rel="stylesheet" type="text/css" href="../_static/jquery-ui-1.10.3.custom.min.css?v=97E0178A" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/user-highlights.css?v=A2326E52" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone-custom-sphinx-bootstrap.css?v=5.0.8" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/language_data.js"></script>
    <script src="../_static/jquery.highlight.js?v=F1C496C8"></script>
    <script src="../_static/sql-wasm.js?v=44FE0709"></script>
    <script src="https://cdn.jsdelivr.net/npm/handsontable@7.2.2/dist/handsontable.full.js"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/pytutor-embed.bundle.js?v=268B5C04"></script>
    <script src="../_static/html4css1.css?v=4D8F3FB4"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/Blob.js?v=AEC2447"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.emitter.bidi.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.emitter.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.fallbacks.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.messagestore.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.parser.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.language.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega@4.0.0-rc.2/build/vega.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega-lite@2.5.0/build/vega-lite.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega-embed@3.14.0/build/vega-embed.js"></script>
    <script src="../_static/runestone.js?v=3EAD674C"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js?v=2B07FE4F"></script>
    <script src="../_static/bootstrap-3.4.1/js/bootstrap.min.js"></script>
    <script src="../_static/jquery-fix.js?v=1226591B"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/jquery.idle-timer.js?v=E942DAF1"></script>
    <script src="../_static/presenter_mode.js?v=2DF16FB4"></script>
    <script src="../_static/theme.js?v=D899FB97"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="4.16. POGIL Activity - Functions, Conditionals, and Strings" href="pogil.html" />
    <link rel="prev" title="4.14. Mixed-up Code Questions" href="mixedupcode.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'py4e-int';
    eBookConfig.basecourse = 'py4e-int';
    eBookConfig.isLoggedIn = false;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;


</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar" aria-label="navbar toggle" class="visuallyhidden">toggle navbar</span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
          <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          py4e-int
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span aria-label="Social" class="visuallyhidden">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span aria-label="Search" class="visuallyhidden">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span aria-label="User" class="visuallyhidden">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:ACFactory.toggleScratchActivecode()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span aria-label="Scratch Activecode" class="visuallyhidden">Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span aria-label="Help" class="visuallyhidden">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=WriteCode'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">4.14. </span>Mixed-up Code Questions" data-toggle="tooltip">
    <a href="mixedupcode.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">4.16. </span>POGIL Activity - Functions, Conditionals, and Strings' data-toggle="tooltip" >              
    <a href="pogil.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

      </div> <!-- navbar -->







    </div> <!-- container -->
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
  
  <div class="section" id="write-code-questions">
<h1><span class="section-number">4.15. </span>Write Code Questions<a class="headerlink" href="#write-code-questions" title="Permalink to this headline">¶</a></h1>
<ol class="arabic">
<li><div id='funct-ex-perimeter' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="funct-ex-perimeterq" class="runestone explainer ac_section alert alert-warning">
<p>Fix the 4 errors so the following code runs and returns the perimeter of a rectangle.</p>

<textarea data-component="activecode" id=funct-ex-perimeterq data-lang="python" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static>
def recPerimeter(length, width)
perimeter = 2 * (length + width)
    Return recPerimeter

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(recPerimeter(10, 20),60,"Tested recPerimeter on inputs 10 and 20")
        self.assertEqual(recPerimeter(1, 2),6,"Tested recPerimeter on inputs 1 and 2")
        self.assertEqual(recPerimeter(23.2, 12.3),71,"Tested recPerimeter on inputs 23.2 and 12.3")
        self.assertEqual(recPerimeter(3.0, 2),10.0,"Tested recPerimeter on inputs 3.0 and 2")

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="funct-ex-perimetera" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=funct-ex-perimetera data-lang="python" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static>
def recPerimeter(length, width):
    perimeter = 2 * (length + width)
    return perimeter
</textarea>
</div>
</div>
    </div>

</li>
<li>
<div data-childcomponent="funct_ex_squareq" class="runestone explainer ac_section alert alert-warning">
<p>Fix the 5 errors so the following code runs and returns the area of a square.</p>

<textarea data-component="activecode" id=funct_ex_squareq data-lang="python" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static>
Def squareArea(sideLength)
    area = length * length
    return squareArea


====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(squareArea(10),100,"Tested squareArea on input 10")
        self.assertEqual(squareArea(23),529,"Tested squareArea on input 23")

myTests().main()
</textarea>
</div>
</li>
<li><div id='funct_ex_triangle' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="funct_ex_triangleq" class="runestone explainer ac_section alert alert-warning">
<p>Change the code so the function takes parameters for the base and height of the triangle.
Then, write code to call the function and print the result.</p>

<textarea data-component="activecode" id=funct_ex_triangleq data-lang="python" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static>
def areaTriangle():
    base = 5
    height = 4
    return (5 * 4) / 2

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(areaTriangle(12,45),270,"Tested areaTriangle on inputs 12 and 45")
        self.assertEqual(areaTriangle(5,4),10,"Tested areaTriangle on inputs 5 and 4")

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="funct_ex_trianglea" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=funct_ex_trianglea data-lang="python" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static>
def areaTriangle(base, height):
    area = (base * height) / 2
    return area

print(areaTriangle(2, 6))
</textarea>
</div>
</div>
    </div>

</li>
<li>
<div data-childcomponent="funct_ex_tripq" class="runestone explainer ac_section alert alert-warning">
<p>Change the code below to create a function <code class="docutils literal notranslate"><span class="pre">tripCost</span></code> that calculates the cost of a trip.
It should take the <code class="docutils literal notranslate"><span class="pre">miles</span></code>, <code class="docutils literal notranslate"><span class="pre">milesPerGallon</span></code>, and <code class="docutils literal notranslate"><span class="pre">pricePerGallon</span></code> as parameters and
should return the cost of the trip.</p>

<textarea data-component="activecode" id=funct_ex_tripq data-lang="python" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static>
miles = 500
milesPerGallon = 26
numGallons = miles / milesPerGallon
pricePerGallon = 3.45
total = numGallons * pricePerGallon
print(total)

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(tripCost(100, 25, 2.24),8.96,"Tested tripCost on inputs 100, 25, and 2.24")
        self.assertEqual(tripCost(250, 20, 3.01),37.625,"Tested tripCost on inputs 250, 20, and 3.01")

myTests().main()
</textarea>
</div>
</li>
<li><div id='funct_ex_name' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="funct_ex_nameq" class="runestone explainer ac_section alert alert-warning">
<p>Fix the errors on line 2 so the function <code class="docutils literal notranslate"><span class="pre">nameAndAge</span></code> returns
the string “My name is <strong>name</strong> and I am <strong>age</strong> years old.” The function
call below should print “My name is John and I am 18 years old.”</p>

<textarea data-component="activecode" id=funct_ex_nameq data-lang="python" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static>
def nameAndAge(nameString, ageInt):
    return(My name is "nameString" and I am + "str(ageInt)" + years old.)

print(nameAndAge("John", 18))

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(nameAndAge("John", 18),"My name is John and I am 18 years old.","Checks if the returned strings are equal.")

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="funct_ex_namea" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=funct_ex_namea data-lang="python" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static>
def nameAndAge(nameString, ageInt):
    return("My name is "+ nameString + " and I am " + str(ageInt) + " years old")

print(nameAndAge("John", 18))
</textarea>
</div>
</div>
    </div>

</li>
<li>
<div data-childcomponent="funct_ex_gradeq" class="runestone explainer ac_section alert alert-warning">
<p>Rewrite the grade program from the previous chapter using a function called <code class="docutils literal notranslate"><span class="pre">computegrade</span></code>
that takes a score as its parameter and returns a grade as a string. If someone enters an
invalid score, return ‘Bad score’.</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">Score</span>    <span class="n">Grade</span>
<span class="o">&gt;=</span> <span class="mf">0.9</span>     <span class="n">A</span>
<span class="o">&gt;=</span> <span class="mf">0.8</span>     <span class="n">B</span>
<span class="o">&gt;=</span> <span class="mf">0.7</span>     <span class="n">C</span>
<span class="o">&gt;=</span> <span class="mf">0.6</span>     <span class="n">D</span>
<span class="o">&lt;</span> <span class="mf">0.6</span>      <span class="n">F</span>
</pre></div>
</div>

<textarea data-component="activecode" id=funct_ex_gradeq data-lang="python" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static>
def computegrade(r):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(computegrade(.95),'A',"Tested input: computegrade(.95)")
        self.assertEqual(computegrade('perfect'),'Bad score',"computegrade('perfect')")
        self.assertEqual(computegrade(10.0),'Bad score',"Tested input: computegrade(10.0)")
        self.assertEqual(computegrade(.75),'C',"Tested input: computegrade(.75)")
        self.assertEqual(computegrade(.5),'F',"Tested input: computegrade(.5)")


myTests().main()
</textarea>
</div>
</li>
<li><div id='funct_ex_sum' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="funct_ex_sumq" class="runestone explainer ac_section alert alert-warning">
<p>Write a fruitful function <code class="docutils literal notranslate"><span class="pre">sumTo(n)</span></code> that returns the sum of all integer numbers up to and
including <cite>n</cite>. So <code class="docutils literal notranslate"><span class="pre">sumTo(10)</span></code> would be <code class="docutils literal notranslate"><span class="pre">1+2+3...+10</span></code> which would return the value 55.
Use the equation  (n * (n + 1)) / 2.</p>

<textarea data-component="activecode" id=funct_ex_sumq data-lang="python" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static>
def sumTo(n):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertAlmostEqual(sumTo(15),120.0,0,"Tested sumTo on input 15")
        self.assertAlmostEqual(sumTo(0),0.0,0,"Tested sumTo on input 0")
        self.assertAlmostEqual(sumTo(25),325.0,0,"Tested sumTo on input 25")
        self.assertAlmostEqual(sumTo(7),28.0,0,"Tested sumTo on input 7")

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="funct_ex_suma" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=funct_ex_suma data-lang="python" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static>
def sumTo(n):
    result = (n * (n + 1)) / 2
    return result
</textarea>
</div>
</div>
    </div>

</li>
<li>
<div data-childcomponent="funct_ex_addq" class="runestone explainer ac_section alert alert-warning">
<p>Rewrite the function <code class="docutils literal notranslate"><span class="pre">sumTo(n)</span></code> that returns the sum of all integer
numbers up to and including <cite>n</cite>. This time use the accumulator pattern.</p>

<textarea data-component="activecode" id=funct_ex_addq data-lang="python" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static>
def sumTo(n):
    # your code here

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sumTo(15),120,"Tested sumTo on input 15")
        self.assertEqual(sumTo(0),0,"Tested sumTo on input 0")
        self.assertEqual(sumTo(25),325,"Tested sumTo on input 25")
        self.assertEqual(sumTo(7),28,"Tested sumTo on input 7")

myTests().main()
</textarea>
</div>
</li>
<li><div id='funct_ex_circle' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="funct_ex_circleq" class="runestone explainer ac_section alert alert-warning">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">areaOfCircle(r)`</span></code> which returns the area of a circle of radius <cite>r</cite>.
Make sure you import the math module in your solution.</p>

<textarea data-component="activecode" id=funct_ex_circleq data-lang="python" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static>
def areaOfCircle(r):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertAlmostEqual(areaOfCircle(5.0),78.53981633974483,5,"Tested input: areaOfCircle(5.0)")
        self.assertEqual(areaOfCircle(5.0),78.53981633974483,"Tested input: areaOfCircle(5.0)")
        self.assertEqual(areaOfCircle(0),0.0,"Tested input: areaOfCircle(0)")
        self.assertAlmostEqual(areaOfCircle(31415.926535897932),3100627668.0299816,5,"Tested input: areaOfCircle(31415.926535897932)")


myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="funct_ex_circlea" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=funct_ex_circlea data-lang="python" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static>
import math

def areaOfCircle(r):
    a = r**2 * math.pi
    return a
</textarea>
</div>
</div>
    </div>

</li>
<li>
<div data-childcomponent="funct_avg_drop_lowest" class="runestone explainer ac_section alert alert-warning">
<p>Finish the function to return the average of a list of numbers, but drop the lowest value. However,
if the list only has one value then return that.  For example, get_avg_drop_lowest([90]) returns 90 and
get_avg_drop_lowest([90, 10]) also returns 90.</p>

<textarea data-component="activecode" id=funct_avg_drop_lowest data-lang="python" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static>
def get_avg_drop_lowest(num_list):

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(get_avg_drop_lowest([90]), 90, "get_avg_drop_lowest([90])")
        self.assertEqual(get_avg_drop_lowest([90, 10]), 90, "get_avg_drop_lowest([90, 10])")
        self.assertEqual(get_avg_drop_lowest([20, -20, 20]), 20, "get_avg_drop_lowest([20, -20, 20])")
        self.assertEqual(get_avg_drop_lowest([70, 80, 100]), 90, "get_avg_drop_lowest([70, 80, 100])")
        self.assertEqual(get_avg_drop_lowest([75, 83, 90]), 86.5, "get_avg_drop_lowest([75, 83, 90])")

myTests().main()
</textarea>
</div>
</li>
<li>
<div data-childcomponent="funct_get_middle" class="runestone explainer ac_section alert alert-warning">
<p>Finish the function below to return the middle characters from the passed string.
If the string has less than 3 characters then return the passed string.
If the string has an odd length then return the middle character.  If the string has an even
length return the two middle characters.  For example, get_middle(‘abc’) returns
‘b’ and get_middle(‘abcd’) returns ‘bc’.</p>

<textarea data-component="activecode" id=funct_get_middle data-lang="python" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static>
def get_middle(str):

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(get_middle('abc'), 'b', "get_middle('abc')")
        self.assertEqual(get_middle('abcd'), 'bc', "get_middle('abcd')")
        self.assertEqual(get_middle('12345'), '3', "get_middle('12345')")
        self.assertEqual(get_middle('123456'), '34', "get_middle('123456')")
        self.assertEqual(get_middle('ab'), 'ab', "get_middle('ab')")
        self.assertEqual(get_middle('a'), 'a', "get_middle('a')")
        self.assertEqual(get_middle(''), '', "get_middle('')")


myTests().main()
</textarea>
</div>
</li>
<li>
<div data-childcomponent="funct_speeding" class="runestone explainer ac_section alert alert-warning">
<p>You are driving a little too fast, and a police officer stops you. Write code to
compute the result, encoded as an int value: 0=no ticket, 1=small ticket,
2=big ticket. If speed is 60 or less, the result is 0. If speed is between 61 and
80 inclusive, the result is 1. If speed is 81 or more, the result is 2. Unless it is
your birthday – on that day, your speed can be 5 higher in all cases.</p>

<textarea data-component="activecode" id=funct_speeding data-lang="python" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static>
def caught_speeding(speed, is_birthday):

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(caught_speeding(60,False), 0, "caught_speeding(60,False)")
        self.assertEqual(caught_speeding(65,False), 1, "caught_speeding(65,False)")
        self.assertEqual(caught_speeding(65,True), 0, "caught_speeding(65,True)")
        self.assertEqual(caught_speeding(80,False), 1, "caught_speeding(80,False)")
        self.assertEqual(caught_speeding(85,False), 2, "caught_speeding(85,False)")
        self.assertEqual(caught_speeding(85,True), 1, "caught_speeding(85,True)")
        self.assertEqual(caught_speeding(70,False), 1, "caught_speeding(70,False)")
        self.assertEqual(caught_speeding(75,False), 1, "caught_speeding(75,False)")
        self.assertEqual(caught_speeding(75,True), 1, "caught_speeding(75,True)")
        self.assertEqual(caught_speeding(40, False), 0, "caught_speeding(40, False)")
        self.assertEqual(caught_speeding(40, True), 0, "caught_speeding(40, True)")
        self.assertEqual(caught_speeding(90, False), 2, "caught_speeding(90, False)")

myTests().main()
</textarea>
</div>
</li>
<li>
<div data-childcomponent="funct_check_guess" class="runestone explainer ac_section alert alert-warning">
<p>Finish the function below to return ‘too low’ if the guess is less than the passed
target, ‘correct’ if they are equal, and ‘too high’ if the guess is greater than the
passed target.  For example, check_guess(5,7) returns ‘too low’,
check_guess(7,7) returns ‘correct’, and check_guess(9,7) returns ‘too high’.</p>

<textarea data-component="activecode" id=funct_check_guess data-lang="python" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static>
def check_guess(guess, target):

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(check_guess(5, 7), 'too low', "check_guess(5, 7)")
        self.assertEqual(check_guess(7, 7), 'correct', "check_guess(7, 7)")
        self.assertEqual(check_guess(9, 7), 'too high', "check_guess(9, 7)")
        self.assertEqual(check_guess(3, 9), 'too low', "check_guess(3, 9)")
        self.assertEqual(check_guess(3, 3), 'correct', "check_guess(3, 3)")
        self.assertEqual(check_guess(20, 9), 'too high', "check_guess(20, 9)")
        self.assertEqual(check_guess(-5, 7), 'too low', "check_guess(-5, 7)")


myTests().main()
</textarea>
</div>
</li>
<li>
<div data-childcomponent="funct_ex_weekq" class="runestone explainer ac_section alert alert-warning">
<p>Given a day of the week encoded as 0=Sun, 1=Mon, 2=Tue, …6=Sat, and a
boolean indicating if we are on vacation, return a string of the form “7:00”
indicating when the alarm clock should ring. If we are on vacation and it is
a weekend (0 - Saturday or 6 - Sunday) it should return “off” and otherwise return
“10:00”. If we are not on vacation and it is a weekend it should return
“10:00” and otherwise “7:00”</p>

<textarea data-component="activecode" id=funct_ex_weekq data-lang="python" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static>
def alarm_clock(day, vacation):
    # your code here

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(alarm_clock(1, False),'7:00',"Tested alarm_clock on input 1 and False")
        self.assertEqual(alarm_clock(5, False),'7:00',"Tested alarm_clock on input 5 and False")
        self.assertEqual(alarm_clock(0, False),'10:00',"Tested alarm_clock on input 0 and False")
        self.assertEqual(alarm_clock(6, False),'10:00',"Tested alarm_clock on input 6 and False")
        self.assertEqual(alarm_clock(0, True),'off',"Tested alarm_clock on input 0 and True")
        self.assertEqual(alarm_clock(6, True),'off',"Tested alarm_clock on input 6 and True")
        self.assertEqual(alarm_clock(1, True),'10:00',"Tested alarm_clock on input 1 and True")

myTests().main()
</textarea>
</div>
</li>
</ol>
</div>


  <div id=scprogresscontainer>
    You have attempted <span id=scprogresstotal></span> of <span id=scprogressposs></span> activities on this page <div id="subchapterprogress"></div>
  </div>
  
      
<ul>
  <li id="relations-prev" style="opacity:0.4" onmouseover="mouseOverPrev()" onmouseout="mouseOutPrev()"class="navLink" title='Previous Section - 4.14. Mixed-up Code Questions' data-toggle="tooltip">
    <a href="mixedupcode.html" aria-label="4.14. Mixed-up Code Questions">
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">
            <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
        </div>
    </a>
  </li>
</ul>

  
<ul>
  <li id="relations-next" style="opacity:0.4" onmouseover="mouseOverNext()" onmouseout="mouseOutNext()" class="navLink" title='Next Section - 4.16. POGIL Activity - Functions, Conditionals, and Strings' data-toggle="tooltip" >
    <a href="pogil.html" aria-label="4.16. POGIL Activity - Functions, Conditionals, and Strings">
      <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">
          <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
      </div>
    </a>
  </li>
</ul>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});

</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>




<script> getSwitch(); </script>


  </body>
</html>