<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Write Code Questions &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/637.css?v=bddd2011b20db76a20ea" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.css?v=7c5388ad4389eae1a318" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/runtime.bundle.js?v=9b07575de0497b0a2725"></script>
    <script src="../_static/637.bundle.js?v=317731c6c17dc7e280ff"></script>
    <script src="../_static/runestone.bundle.js?v=1980ca51def921204b31"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'py4e-int';
    eBookConfig.basecourse = 'py4e-int';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          py4e-int
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=writecode_orig'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
  
<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="write-code-questions">
<h1>Write Code Questions<a class="headerlink" href="#write-code-questions" title="Permalink to this headline">¶</a></h1>
<ol class="arabic">
<li><div id='bsoup_writecode1' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=bsoup_writecode1q data-question_label="1">
<div id=bsoup_writecode1q_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">cur_stud_link</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">url</span></code>
and uses BeautifulSoup to return the URL that links to Current Students in the Navigation Bar.
For example, <code class="docutils literal notranslate"><span class="pre">cur_stud_link('https://umich.edu/')</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;https://umich.edu/current-students/&quot;</span></code>.</p>

</div>
<textarea data-lang="python3" id="bsoup_writecode1q_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
from bs4 import BeautifulSoup
import requests

def cur_stud_link(url):

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(cur_stud_link('https://umich.edu/'), "https://umich.edu/current-students/", "cur_stud_link('https://umich.edu/')")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=bsoup_writecode1a data-question_label="2">
<div id=bsoup_writecode1a_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">cur_stud_link</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">url</span></code>
and uses BeautifulSoup to return the URL that links to Current Students in the Navigation Bar.
For example, <code class="docutils literal notranslate"><span class="pre">cur_stud_link('https://umich.edu/')</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;https://umich.edu/current-students/&quot;</span></code>.</p>

</div>
<textarea data-lang="python3" id="bsoup_writecode1a_editor" 
      data-timelimit=25000     data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
from bs4 import BeautifulSoup
import requests

def cur_stud_link(url):
    r = requests.get(url)
    soup = BeautifulSoup(r.text, 'html.parser')

    # inspect Current Students to get the details
    text = soup.find('a', text = "Current Students")
    info = text.get('href')
    all_info = "https://umich.edu" + info
    return all_info

print(cur_stud_link('https://umich.edu/'))
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=bsoup_writecode2q data-question_label="3">
<div id=bsoup_writecode2q_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">headings</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">url</span></code> and uses BeautifulSoup to return a list of all heading names (not subheadings) from the
Contents box (History, Emergence of MOOC providers, Student experience and pedagogy, etc.). For example, <code class="docutils literal notranslate"><span class="pre">headings('https://en.wikipedia.org/wiki/Massive_open_online_course')</span></code>
should return <code class="docutils literal notranslate"><span class="pre">['History',</span> <span class="pre">'Emergence</span> <span class="pre">of</span> <span class="pre">MOOC</span> <span class="pre">providers',</span> <span class="pre">'Student</span> <span class="pre">experience</span> <span class="pre">and</span> <span class="pre">pedagogy',</span> <span class="pre">'Information</span> <span class="pre">architecture',</span> <span class="pre">'Industry',</span> <span class="pre">'Benefits',</span> <span class="pre">'Challenges</span> <span class="pre">and</span> <span class="pre">criticisms',</span> <span class="pre">'See</span> <span class="pre">also',</span> <span class="pre">'References',</span> <span class="pre">'Sources',</span> <span class="pre">'Further</span> <span class="pre">reading']</span></code>.</p>

</div>
<textarea data-lang="python3" id="bsoup_writecode2q_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
from bs4 import BeautifulSoup
import requests

def headings(url):

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(headings('https://en.wikipedia.org/wiki/Massive_open_online_course'), ['History', 'Emergence of MOOC providers', 'Student experience and pedagogy', 'Information architecture', 'Industry', 'Benefits', 'Challenges and criticisms', 'See also', 'References', 'Sources', 'Further reading'], "headings('https://en.wikipedia.org/wiki/Massive_open_online_course')")

myTests().main()
</textarea>
</div>
</div>
</li>
<li><div id='bsoup_writecode3' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/Users/barbarer/Documents/GitHub/py4e-int/_sources/bsoup/writecode_orig.rst</span>, line 91)</p>
<p>Duplicate ID – see bsoup/writecode, line 104</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: bsoup_writecode3q
    :practice: T
    :nocodelens:
    :language: python3
    :autograde: unittest

    Write a function called ``coursera_dict`` that takes in a parameter ``url``, uses BeautifulSoup to web scrape the Coursera enrollees table,
    and returns a dictionary with Country as keys and Percentage as values. For example, ``coursera_dict(&#39;https://en.wikipedia.org/wiki/Massive_open_online_course&#39;)``
    should return ``{&#39;United States&#39;: &#39;27.7%&#39;, &#39;India&#39;: &#39;8.8%&#39;, &#39;Brazil&#39;: &#39;5.1%&#39;, &#39;United Kingdom&#39;: &#39;4.4%&#39;, &#39;Spain&#39;: &#39;4.0%&#39;, &#39;Canada&#39;: &#39;3.6%&#39;, &#39;Australia&#39;: &#39;2.3%&#39;, &#39;Russia&#39;: &#39;2.2%&#39;, &#39;Rest of world&#39;: &#39;41.9%&#39;}``.
    ~~~~
    from bs4 import BeautifulSoup
    import requests

    def coursera_dict(url):

    ====
    from unittest.gui import TestCaseGui

    class myTests(TestCaseGui):

        def testOne(self):
            self.assertEqual(coursera_dict(&#39;https://en.wikipedia.org/wiki/Massive_open_online_course&#39;), {&#39;United States&#39;: &#39;27.7%&#39;, &#39;India&#39;: &#39;8.8%&#39;, &#39;Brazil&#39;: &#39;5.1%&#39;, &#39;United Kingdom&#39;: &#39;4.4%&#39;, &#39;Spain&#39;: &#39;4.0%&#39;, &#39;Canada&#39;: &#39;3.6%&#39;, &#39;Australia&#39;: &#39;2.3%&#39;, &#39;Russia&#39;: &#39;2.2%&#39;, &#39;Rest of world&#39;: &#39;41.9%&#39;}, &quot;coursera_dict(&#39;https://en.wikipedia.org/wiki/Massive_open_online_course&#39;)&quot;)

    myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/Users/barbarer/Documents/GitHub/py4e-int/_sources/bsoup/writecode_orig.rst</span>, line 118)</p>
<p>Duplicate ID – see bsoup/writecode, line 130</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: bsoup_writecode3a
    :optional:
    :language: python3
    :autograde: unittest

    Write a function called ``coursera_dict`` that takes in a parameter ``url``, uses BeautifulSoup to web scrape the Coursera enrollees table,
    and returns a dictionary with Country as keys and Percentage as values. For example, ``coursera_dict(&#39;https://en.wikipedia.org/wiki/Massive_open_online_course&#39;)``
    should return ``{&#39;United States&#39;: &#39;27.7%&#39;, &#39;India&#39;: &#39;8.8%&#39;, &#39;Brazil&#39;: &#39;5.1%&#39;, &#39;United Kingdom&#39;: &#39;4.4%&#39;, &#39;Spain&#39;: &#39;4.0%&#39;, &#39;Canada&#39;: &#39;3.6%&#39;, &#39;Australia&#39;: &#39;2.3%&#39;, &#39;Russia&#39;: &#39;2.2%&#39;, &#39;Rest of world&#39;: &#39;41.9%&#39;}``.
    ~~~~
    from bs4 import BeautifulSoup
    import requests

    def coursera_dict(url):
        r = requests.get(url)
        soup = BeautifulSoup(r.text, &#39;html.parser&#39;)

        coursera_enrollees = {}

        table = soup.find(&#39;table&#39;, class_ = &#39;wikitable&#39;)
        rows = table.find_all(&#39;tr&#39;)

        # The first tr (table row) tag with th (table heading) tags, index 0, doesn&#39;t have td tags, so it&#39;d return an empty list. That&#39;s why we do rows[1:].
        for row in rows[1:]:
            both_columns = row.find_all(&quot;td&quot;)
            key = both_columns[0].text
            # Remember to strip the newlines from the percentages
            value = both_columns[1].text.rstrip(&quot;\n&quot;)
            coursera_enrollees[key] = value

        return coursera_enrollees

    print(coursera_dict(&#39;https://en.wikipedia.org/wiki/Massive_open_online_course&#39;))
</pre></div>
</div>
</div>
</div>
    </div>

</li>
<li><div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/Users/barbarer/Documents/GitHub/py4e-int/_sources/bsoup/writecode_orig.rst</span>, line 152)</p>
<p>Duplicate ID – see bsoup/writecode, line 164</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: bsoup_writecode4q
    :practice: T
    :nocodelens:
    :language: python3
    :autograde: unittest

    Write a function called ``types`` that takes in a parameter ``url``, uses BeautifulSoup to web scrape the Summary of Python 3&#39;s built-in types table,
    and returns a list of the types. For example, ``types(&#39;https://en.wikipedia.org/wiki/Python_(programming_language)&#39;)``
    should return ``[&#39;bool&#39;, &#39;bytearray&#39;, &#39;bytes&#39;, &#39;complex&#39;, &#39;dict&#39;, &#39;ellipsis&#39;, &#39;float&#39;, &#39;frozenset&#39;, &#39;int&#39;, &#39;list&#39;, &#39;NoneType&#39;, &#39;NotImplementedType&#39;, &#39;range&#39;, &#39;set&#39;, &#39;str&#39;, &#39;tuple&#39;]``.
    ~~~~
    from bs4 import BeautifulSoup
    import requests

    def types(url):

    ====
    from unittest.gui import TestCaseGui

    class myTests(TestCaseGui):

        def testOne(self):
            self.assertEqual(types(&#39;https://en.wikipedia.org/wiki/Python_(programming_language)&#39;), [&#39;bool&#39;, &#39;bytearray&#39;, &#39;bytes&#39;, &#39;complex&#39;, &#39;dict&#39;, &#39;ellipsis&#39;, &#39;float&#39;, &#39;frozenset&#39;, &#39;int&#39;, &#39;list&#39;, &#39;NoneType&#39;, &#39;NotImplementedType&#39;, &#39;range&#39;, &#39;set&#39;, &#39;str&#39;, &#39;tuple&#39;], &quot;types(&#39;https://en.wikipedia.org/wiki/Python_(programming_language)&#39;)&quot;)

    myTests().main()

</pre></div>
</div>
</div>
</li>
<li><div id='bsoup_writecode5' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/Users/barbarer/Documents/GitHub/py4e-int/_sources/bsoup/writecode_orig.rst</span>, line 183)</p>
<p>Duplicate ID – see bsoup/writecode, line 194</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: bsoup_writecode5q
    :practice: T
    :nocodelens:
    :language: python3
    :autograde: unittest

    Write a function called ``restaurants_list`` that takes in a parameter ``url``, uses BeautifulSoup to web scrape the 12 essential Ann Arbor restaurants&#39; names,
    and returns a list of the 12 best Ann Arbor restaurants in 2019. For example, ``restaurants_list(&#39;https://detroit.eater.com/maps/best-ann-arbor-restaurants&#39;)``
    should return ``[&quot;1. Knight&#39;s Steakhouse&quot;, &#39;2. The Last Word&#39;, &quot;3. Krazy Jim&#39;s Blimpy Burger&quot;, &#39;4. Blue Llama Jazz Club&#39;, &#39;5. Spencer&#39;, &#39;6. Miss Kim&#39;, &quot;7. Zingerman&#39;s Delicatessen&quot;, &#39;8. Tomukun Noodle Bar&#39;, &#39;9. Seoul Street&#39;, &#39;10. Ricewood&#39;, &#39;11. Bellflower&#39;, &quot;12. Ma Lou&#39;s Fried Chicken&quot;]``.
    ~~~~
    from bs4 import BeautifulSoup
    import requests

    def restaurants_list(url):

    ====
    from unittest.gui import TestCaseGui

    class myTests(TestCaseGui):

        def testOne(self):
            self.assertEqual(restaurants_list(&#39;https://detroit.eater.com/maps/best-ann-arbor-restaurants&#39;), [&quot;1. Knight&#39;s Steakhouse&quot;, &#39;2. The Last Word&#39;, &quot;3. Krazy Jim&#39;s Blimpy Burger&quot;, &#39;4. Blue Llama Jazz Club&#39;, &#39;5. Spencer&#39;, &#39;6. Miss Kim&#39;, &quot;7. Zingerman&#39;s Delicatessen&quot;, &#39;8. Tomukun Noodle Bar&#39;, &#39;9. Seoul Street&#39;, &#39;10. Ricewood&#39;, &#39;11. Bellflower&#39;, &quot;12. Ma Lou&#39;s Fried Chicken&quot;], &quot;restaurants_list(&#39;https://detroit.eater.com/maps/best-ann-arbor-restaurants&#39;)&quot;)

    myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/Users/barbarer/Documents/GitHub/py4e-int/_sources/bsoup/writecode_orig.rst</span>, line 210)</p>
<p>Duplicate ID – see bsoup/writecode, line 220</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: bsoup_writecode5a
    :optional:
    :language: python3
    :autograde: unittest

    Write a function called ``restaurants_list`` that takes in a parameter ``url``, uses BeautifulSoup to web scrape the 12 essential Ann Arbor restaurants&#39; names,
    and returns a list of the 12 best Ann Arbor restaurants in 2019. For example, ``restaurants_list(&#39;https://detroit.eater.com/maps/best-ann-arbor-restaurants&#39;)``
    should return ``[&quot;1. Knight&#39;s Steakhouse&quot;, &#39;2. The Last Word&#39;, &quot;3. Krazy Jim&#39;s Blimpy Burger&quot;, &#39;4. Blue Llama Jazz Club&#39;, &#39;5. Spencer&#39;, &#39;6. Miss Kim&#39;, &quot;7. Zingerman&#39;s Delicatessen&quot;, &#39;8. Tomukun Noodle Bar&#39;, &#39;9. Seoul Street&#39;, &#39;10. Ricewood&#39;, &#39;11. Bellflower&#39;, &quot;12. Ma Lou&#39;s Fried Chicken&quot;]``.
    ~~~~
    from bs4 import BeautifulSoup
    import requests

    def restaurants_list(url):
        r = requests.get(url)
        soup = BeautifulSoup(r.text, &#39;html.parser&#39;)

        ann_arbor_2019_restaurants_list = []

        restaurants = soup.find_all(&#39;div&#39;, class_ = &quot;c-mapstack__card-hed&quot;)

        for restaurant in restaurants:
            # Be careful not to strip the n off Chicke(n) for 12. Ma Lou&#39;s Fried Chicken.
            ann_arbor_2019_restaurants_list.append(restaurant.text.strip(&quot;\n\n\n&quot;).rstrip(&quot;Copy Link&quot;).rstrip(&quot;\n\n\n&quot;))

        return ann_arbor_2019_restaurants_list

    print(restaurants_list(&#39;https://detroit.eater.com/maps/best-ann-arbor-restaurants&#39;))
</pre></div>
</div>
</div>
</div>
    </div>

</li>
<li><div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/Users/barbarer/Documents/GitHub/py4e-int/_sources/bsoup/writecode_orig.rst</span>, line 239)</p>
<p>Duplicate ID – see bsoup/writecode, line 249</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: bsoup_writecode6q
    :practice: T
    :nocodelens:
    :language: python3
    :autograde: unittest

    Write a function called ``bsi_list`` that takes in a parameter ``url``, uses BeautifulSoup to web scrape the section names and subsection names under Programs Bachelor of Science in Information,
    and returns a list of the section names and subsection names under Programs Bachelor of Science in Information. For example, ``bsi_list(&#39;https://www.si.umich.edu/programs/bachelor-science-information&#39;)``
    should return ``[&#39;How do I apply?&#39;, &#39;Current U-M students&#39;, &#39;Transfer students&#39;, &#39;High school students&#39;, &#39;New transfers&#39;, &#39;UMSI Community College Summer Institute&#39;, &#39;Curriculum&#39;, &#39;Career outcomes&#39;, &#39;Internships&#39;, &#39;F-1 visa holders&#39;, &#39;Career development resources&#39;, &#39;Engaged learning&#39;, &#39;Tuition and funding&#39;, &#39;Connect with us&#39;, &#39;Study abroad for undergraduates&#39;, &#39;Exchange student information&#39;]``.
    ~~~~
    from bs4 import BeautifulSoup
    import requests
    import re

    def bsi_list(url):

    ====
    from unittest.gui import TestCaseGui

    class myTests(TestCaseGui):

        def testOne(self):
            self.assertEqual(bsi_list(&#39;https://www.si.umich.edu/programs/bachelor-science-information&#39;), [&#39;How do I apply?&#39;, &#39;Current U-M students&#39;, &#39;Transfer students&#39;, &#39;High school students&#39;, &#39;New transfers&#39;, &#39;UMSI Community College Summer Institute&#39;, &#39;Curriculum&#39;, &#39;Career outcomes&#39;, &#39;Internships&#39;, &#39;F-1 visa holders&#39;, &#39;Career development resources&#39;, &#39;Engaged learning&#39;, &#39;Tuition and funding&#39;, &#39;Connect with us&#39;, &#39;Study abroad for undergraduates&#39;, &#39;Exchange student information&#39;], &quot;bsi_list(&#39;https://www.si.umich.edu/programs/bachelor-science-information&#39;)&quot;)

    myTests().main()

</pre></div>
</div>
</div>
</li>
<li><div id='bsoup_writecode7' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/Users/barbarer/Documents/GitHub/py4e-int/_sources/bsoup/writecode_orig.rst</span>, line 271)</p>
<p>Duplicate ID – see bsoup/writecode, line 280</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: bsoup_writecode7q
    :practice: T
    :nocodelens:
    :language: python3
    :autograde: unittest

    Write a function called ``bsi_dict`` that takes in a parameter ``url``, uses BeautifulSoup to web scrape and create a list of the section names and subsection names under Programs Bachelor of Science in Information,
    and creates another list that contains the full urls of the named sections and subsections. Then, return a dictionary that contains the names as keys and their full urls as values. For example, ``bsi_dict(&#39;https://www.si.umich.edu/programs/bachelor-science-information&#39;)``
    should return ``{&#39;How do I apply?&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/how-do-i-apply&#39;, &#39;Current U-M students&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/how-do-i-apply/cross-campus-transfer&#39;, &#39;Transfer students&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/how-do-i-apply/new-transfer-students-bsi&#39;, &#39;High school students&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/how-do-i-apply/preferred-admissions&#39;, &#39;New transfers&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/new-transfers&#39;, &#39;UMSI Community College Summer Institute&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/new-transfers/umsi-community-college-summer-institute&#39;, &#39;Curriculum&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/curriculum&#39;, &#39;Career outcomes&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/career-outcomes&#39;, &#39;Internships&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/internships&#39;, &#39;F-1 visa holders&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/internships/f-1-visa-holders&#39;, &#39;Career development resources&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/career-development-resources&#39;, &#39;Engaged learning&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/engaged-learning&#39;, &#39;Tuition and funding&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/tuition-and-funding&#39;, &#39;Connect with us&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/connect-us&#39;, &#39;Study abroad for undergraduates&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/study-abroad-undergraduates&#39;, &#39;Exchange student information&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/exchange-student-information&#39;}``.
    ~~~~
    from bs4 import BeautifulSoup
    import requests
    import re

    def bsi_dict(url):

    ====
    from unittest.gui import TestCaseGui

    class myTests(TestCaseGui):

        def testOne(self):
            self.assertEqual(bsi_dict(&#39;https://www.si.umich.edu/programs/bachelor-science-information&#39;), {&#39;How do I apply?&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/how-do-i-apply&#39;, &#39;Current U-M students&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/how-do-i-apply/cross-campus-transfer&#39;, &#39;Transfer students&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/how-do-i-apply/new-transfer-students-bsi&#39;, &#39;High school students&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/how-do-i-apply/preferred-admissions&#39;, &#39;New transfers&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/new-transfers&#39;, &#39;UMSI Community College Summer Institute&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/new-transfers/umsi-community-college-summer-institute&#39;, &#39;Curriculum&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/curriculum&#39;, &#39;Career outcomes&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/career-outcomes&#39;, &#39;Internships&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/internships&#39;, &#39;F-1 visa holders&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/internships/f-1-visa-holders&#39;, &#39;Career development resources&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/career-development-resources&#39;, &#39;Engaged learning&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/engaged-learning&#39;, &#39;Tuition and funding&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/tuition-and-funding&#39;, &#39;Connect with us&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/connect-us&#39;, &#39;Study abroad for undergraduates&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/study-abroad-undergraduates&#39;, &#39;Exchange student information&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/exchange-student-information&#39;}, &quot;bsi_dict(&#39;https://www.si.umich.edu/programs/bachelor-science-information&#39;)&quot;)

    myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/Users/barbarer/Documents/GitHub/py4e-int/_sources/bsoup/writecode_orig.rst</span>, line 299)</p>
<p>Duplicate ID – see bsoup/writecode, line 307</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: bsoup_writecode7a
    :optional:
    :language: python3
    :autograde: unittest

    Write a function called ``bsi_dict`` that takes in a parameter ``url``, uses BeautifulSoup to web scrape and create a list of the section names and subsection names under Programs Bachelor of Science in Information,
    and creates another list that contains the full urls of the named sections and subsections. Then, return a dictionary that contains the names as keys and their full urls as values. For example, ``bsi_dict(&#39;https://www.si.umich.edu/programs/bachelor-science-information&#39;)``
    should return ``{&#39;How do I apply?&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/how-do-i-apply&#39;, &#39;Current U-M students&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/how-do-i-apply/cross-campus-transfer&#39;, &#39;Transfer students&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/how-do-i-apply/new-transfer-students-bsi&#39;, &#39;High school students&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/how-do-i-apply/preferred-admissions&#39;, &#39;New transfers&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/new-transfers&#39;, &#39;UMSI Community College Summer Institute&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/new-transfers/umsi-community-college-summer-institute&#39;, &#39;Curriculum&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/curriculum&#39;, &#39;Career outcomes&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/career-outcomes&#39;, &#39;Internships&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/internships&#39;, &#39;F-1 visa holders&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/internships/f-1-visa-holders&#39;, &#39;Career development resources&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/career-development-resources&#39;, &#39;Engaged learning&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/engaged-learning&#39;, &#39;Tuition and funding&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/tuition-and-funding&#39;, &#39;Connect with us&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/connect-us&#39;, &#39;Study abroad for undergraduates&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/study-abroad-undergraduates&#39;, &#39;Exchange student information&#39;: &#39;https://www.si.umich.edu/programs/bachelor-science-information/exchange-student-information&#39;}``.
    ~~~~
    from bs4 import BeautifulSoup
    import requests
    import re

    def bsi_dict(url):
        r = requests.get(url)
        soup = BeautifulSoup(r.text, &#39;html.parser&#39;)

        bsi_list = []
        href_list = []
        bsi_dict = {}

        sections = soup.find_all(&#39;a&#39;, href = re.compile(&quot;/programs/bachelor-science-information/&quot;))

        for section in sections[:-2]:
            bsi_list.append(section.text.strip(&quot;\n&quot;).strip())
            url_end = section.get(&#39;href&#39;)
            full_url = &quot;https://www.si.umich.edu&quot; + url_end
            href_list.append(full_url)

        for i in range(len(bsi_list)):
            bsi_dict[bsi_list[i]] = href_list[i]

        return bsi_dict

    print(bsi_dict(&#39;https://www.si.umich.edu/programs/bachelor-science-information&#39;))
</pre></div>
</div>
</div>
</div>
    </div>

</li>
<li><div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/Users/barbarer/Documents/GitHub/py4e-int/_sources/bsoup/writecode_orig.rst</span>, line 336)</p>
<p>Duplicate ID – see bsoup/writecode, line 344</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: bsoup_writecode8q
    :practice: T
    :nocodelens:
    :language: python3
    :autograde: unittest

    Write a function called ``envelope_address`` that takes in a parameter ``url`` and uses BeautifulSoup to web scrape and return the address text in the footer with the newlines stripped. For example, ``envelope_address(&#39;https://www.si.umich.edu/programs/bachelor-science-information&#39;)``
    should return first line: ``&#39;School of Information&#39;``, second line: ``&#39;University of Michigan&#39;``, third line: ``&#39;105 S State St.&#39;``, and last line: ``&#39;Ann Arbor, MI 48109-1285&#39;``.
    ~~~~
    from bs4 import BeautifulSoup
    import requests

    def envelope_address(url):

    ====
    from unittest.gui import TestCaseGui

    class myTests(TestCaseGui):

        def testOne(self):
            r = requests.get(&#39;https://www.si.umich.edu/programs/bachelor-science-information&#39;)
            soup = BeautifulSoup(r.text, &#39;html.parser&#39;)
            address = soup.find(&#39;div&#39;, class_ = &#39;footer--address&#39;).text.strip(&quot;\n&quot;)
            self.assertEqual(envelope_address(&#39;https://www.si.umich.edu/programs/bachelor-science-information&#39;), address, &quot;envelope_address(&#39;https://www.si.umich.edu/programs/bachelor-science-information&#39;)&quot;)

    myTests().main()

</pre></div>
</div>
</div>
</li>
<li><div id='bsoup_writecode9' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/Users/barbarer/Documents/GitHub/py4e-int/_sources/bsoup/writecode_orig.rst</span>, line 369)</p>
<p>Duplicate ID – see bsoup/writecode, line 376</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: bsoup_writecode9q
    :practice: T
    :nocodelens:
    :language: python3
    :autograde: unittest

    Write a function called ``name_email`` that takes in a parameter ``url``, uses BeautifulSoup to web scrape the names and email addresses,
    and returns a dictionary with the names as keys and the email addresses as values. Use the string replace method while web scraping. For example, ``name_email(&#39;https://www.si.umich.edu/people/directory/faculty/e&#39;)``
    should return ``{&#39;Paul Edwards&#39;: &#39;pne@umich.edu&#39;, &#39;Ron Eglash&#39;: &#39;eglash@umich.edu&#39;, &#39;Nicole Ellison&#39;: &#39;enicole@umich.edu&#39;, &#39;Barbara Ericson&#39;: &#39;barbarer@umich.edu&#39;}``.
    ~~~~
    from bs4 import BeautifulSoup
    import requests

    def name_email(url):

    ====
    from unittest.gui import TestCaseGui

    class myTests(TestCaseGui):

        def testOne(self):
            self.assertEqual(name_email(&#39;https://www.si.umich.edu/people/directory/faculty/e&#39;), {&#39;Paul Edwards&#39;: &#39;pne@umich.edu&#39;, &#39;Ron Eglash&#39;: &#39;eglash@umich.edu&#39;, &#39;Nicole Ellison&#39;: &#39;enicole@umich.edu&#39;, &#39;Barbara Ericson&#39;: &#39;barbarer@umich.edu&#39;}, &quot;name_email(&#39;https://www.si.umich.edu/people/directory/faculty/e&#39;)&quot;)

    myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/Users/barbarer/Documents/GitHub/py4e-int/_sources/bsoup/writecode_orig.rst</span>, line 396)</p>
<p>Duplicate ID – see bsoup/writecode, line 402</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: bsoup_writecode9a
    :optional:
    :language: python3
    :autograde: unittest

    Write a function called ``name_email`` that takes in a parameter ``url``, uses BeautifulSoup to web scrape the names and email addresses,
    and returns a dictionary with the names as keys and the email addresses as values. Use the string replace method while web scraping. For example, ``name_email(&#39;https://www.si.umich.edu/people/directory/faculty/e&#39;)``
    should return ``{&#39;Paul Edwards&#39;: &#39;pne@umich.edu&#39;, &#39;Ron Eglash&#39;: &#39;eglash@umich.edu&#39;, &#39;Nicole Ellison&#39;: &#39;enicole@umich.edu&#39;, &#39;Barbara Ericson&#39;: &#39;barbarer@umich.edu&#39;}``.
    ~~~~
    from bs4 import BeautifulSoup
    import requests

    def name_email(url):
        r = requests.get(url)
        soup = BeautifulSoup(r.text, &#39;html.parser&#39;)

        names_list = []
        emails_list = []
        name_email_dict = {}

        names = soup.find_all(class_=&quot;research-person-profile__name&quot;)
        emails = soup.find_all(&quot;a&quot;, target=&quot;_blank&quot;)

        for name in names:
            profile_name = name.text.strip().replace(&#39;\n \n &#39;, &#39;&#39;)
            names_list.append(profile_name)

        for email in emails:
            profile_email = email.get(&quot;href&quot;)
            profile_email = profile_email.replace(&#39;mailto:&#39;, &#39;&#39;)
            emails_list.append(profile_email)

        for i in range(len(names_list)):
            name_email_dict[names_list[i]] = emails_list[i]

        return name_email_dict

    print(name_email(&#39;https://www.si.umich.edu/people/directory/faculty/e&#39;))
</pre></div>
</div>
</div>
</div>
    </div>

</li>
<li><div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/Users/barbarer/Documents/GitHub/py4e-int/_sources/bsoup/writecode_orig.rst</span>, line 436)</p>
<p>Duplicate ID – see bsoup/writecode, line 442</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: bsoup_writecode10q
    :practice: T
    :nocodelens:
    :language: python3
    :autograde: unittest

    Write a function called ``program_email`` that takes in a parameter ``url``, uses BeautifulSoup to web scrape the 5 program names under Email Addresses
    and Admissions (BSI program, MSI program, etc.) and their associated email addresses, and returns a dictionary with the program names as keys and the email addresses as values.
    For example, ``program_email(&#39;https://www.si.umich.edu/about-umsi/contact-us&#39;)`` should return ``{&#39;BSI program&#39;: &#39;umsi.undergrad@umich.edu&#39;, &#39;MSI program&#39;: &#39;umsi.admissions@umich.edu&#39;, &#39;MHI program&#39;: &#39;hi.admissions@umich.edu&#39;, &#39;MADS program&#39;: &#39;umsi.mads@umich.edu&#39;, &#39;Doctoral program&#39;: &#39;umsi.phd.admissions@umich.edu&#39;}``.
    ~~~~
    from bs4 import BeautifulSoup
    import requests

    def program_email(url):

    ====
    from unittest.gui import TestCaseGui

    class myTests(TestCaseGui):

        def testOne(self):
            self.assertEqual(program_email(&#39;https://www.si.umich.edu/about-umsi/contact-us&#39;), {&#39;BSI program&#39;: &#39;umsi.undergrad@umich.edu&#39;, &#39;MSI program&#39;: &#39;umsi.admissions@umich.edu&#39;, &#39;MHI program&#39;: &#39;hi.admissions@umich.edu&#39;, &#39;MADS program&#39;: &#39;umsi.mads@umich.edu&#39;, &#39;Doctoral program&#39;: &#39;umsi.phd.admissions@umich.edu&#39;}, &quot;program_email(&#39;https://www.si.umich.edu/about-umsi/contact-us&#39;)&quot;)

    myTests().main()
</pre></div>
</div>
</div>
</li>
</ol>
</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>