<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Mixed-Up Code Questions &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/637.css?v=bddd2011b20db76a20ea" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.css?v=7c5388ad4389eae1a318" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/runtime.bundle.js?v=9b07575de0497b0a2725"></script>
    <script src="../_static/637.bundle.js?v=317731c6c17dc7e280ff"></script>
    <script src="../_static/runestone.bundle.js?v=1980ca51def921204b31"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'py4e-int';
    eBookConfig.basecourse = 'py4e-int';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          py4e-int
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=mixedupcode_orig'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
  
<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="mixed-up-code-questions">
<h1>Mixed-Up Code Questions<a class="headerlink" href="#mixed-up-code-questions" title="Permalink to this headline">¶</a></h1>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="inheritance_mixed1" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class called <code class="docutils literal notranslate"><span class="pre">Fruit</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">size</span></code>. Next, create a class called <code class="docutils literal notranslate"><span class="pre">Orange</span></code> with a
<code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits from the <code class="docutils literal notranslate"><span class="pre">Fruit</span></code> constructor and additionally has a parameter <code class="docutils literal notranslate"><span class="pre">flavor</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">Fruit(10).size</span></code> would be <code class="docutils literal notranslate"><span class="pre">10</span></code>,
and <code class="docutils literal notranslate"><span class="pre">Orange(10,</span> <span class="pre">'sweet').flavor</span></code> would be <code class="docutils literal notranslate"><span class="pre">&quot;sweet&quot;</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="1"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class Fruit:
---
    def __init__(self, size):
---
    def init(self, size): #paired
---
        self.size = size
---
class Orange(Fruit):
---
class Orange(): #paired
---
    def __init__(self, size, flavor):
---
        super().__init__(size)
---
        self.flavor = flavor
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=inheritance_mixed1_ac data-question_label="2">
<div id=inheritance_mixed1_ac_question class="ac_question col-md-12">
<p>Write a class called <code class="docutils literal notranslate"><span class="pre">Fruit</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">size</span></code>. Next, create a class called <code class="docutils literal notranslate"><span class="pre">Orange</span></code> with a
<code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits from the <code class="docutils literal notranslate"><span class="pre">Fruit</span></code> constructor and additionally has a parameter <code class="docutils literal notranslate"><span class="pre">flavor</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">Fruit(10).size</span></code> would be <code class="docutils literal notranslate"><span class="pre">10</span></code>,
and <code class="docutils literal notranslate"><span class="pre">Orange(10,</span> <span class="pre">'sweet').flavor</span></code> would be <code class="docutils literal notranslate"><span class="pre">&quot;sweet&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="inheritance_mixed1_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        size = 10
        flavor = 'sweet'
        f = Fruit(10)
        f2 = Orange(10, 'sweet')
        self.assertEqual(f.size, size, "Fruit(10).size")
        self.assertEqual(f2.size, size, "Orange(10, 'sweet').size")
        self.assertEqual(f2.flavor, flavor, "Orange(10, 'sweet').flavor")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="inheritance_mixed2" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class called <code class="docutils literal notranslate"><span class="pre">Person</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">name</span></code> and <code class="docutils literal notranslate"><span class="pre">age</span></code>. Next, create a class called <code class="docutils literal notranslate"><span class="pre">Student</span></code> with a
<code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits from the <code class="docutils literal notranslate"><span class="pre">Person</span></code> constructor and additionally has parameters <code class="docutils literal notranslate"><span class="pre">grade</span></code> and <code class="docutils literal notranslate"><span class="pre">gpa</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">Person('Susan',</span> <span class="pre">13).name</span></code> would be <code class="docutils literal notranslate"><span class="pre">&quot;Susan&quot;</span></code>
and <code class="docutils literal notranslate"><span class="pre">Student('Sally',</span> <span class="pre">16,</span> <span class="pre">11,</span> <span class="pre">3.7).gpa</span></code> would be <code class="docutils literal notranslate"><span class="pre">3.7</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="3"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class Person:
---
    def __init__(self, name, age):
---
    def init(self, name, age): #paired
---
        self.name = name
        self.age = age
---
class Student(Person):
---
class Student(): #paired
---
    def __init__(self, name, age, grade, gpa):
---
    def __init__(name, age, grade, gpa): #paired
---
        super().__init__(name, age)
---
        super().__init__(age) #paired
        super().__init__(name)
---
        self.grade = grade
---
        self.gpa = gpa
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=inheritance_mixed2_ac data-question_label="4">
<div id=inheritance_mixed2_ac_question class="ac_question col-md-12">
<p>Write a class called <code class="docutils literal notranslate"><span class="pre">Person</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">name</span></code> and <code class="docutils literal notranslate"><span class="pre">age</span></code>. Next, create a class called <code class="docutils literal notranslate"><span class="pre">Student</span></code> with a
<code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits from the <code class="docutils literal notranslate"><span class="pre">Person</span></code> constructor and additionally has parameters <code class="docutils literal notranslate"><span class="pre">grade</span></code> and <code class="docutils literal notranslate"><span class="pre">gpa</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">Person('Susan',</span> <span class="pre">13).name</span></code> would be <code class="docutils literal notranslate"><span class="pre">&quot;Susan&quot;</span></code>
and <code class="docutils literal notranslate"><span class="pre">Student('Sally',</span> <span class="pre">16,</span> <span class="pre">11,</span> <span class="pre">3.7).gpa</span></code> would be <code class="docutils literal notranslate"><span class="pre">3.7</span></code>.</p>

</div>
<textarea data-lang="python" id="inheritance_mixed2_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        p = Person('Susan', 13)
        s = Student('Sally', 16, 11, 3.7)
        self.assertEqual(p.name, 'Susan', "Person('Susan', 13).name")
        self.assertEqual(p.age, 13, "Person('Susan', 13).age")
        self.assertEqual(s.name, 'Sally', "Student('Sally', 16, 11, 3.7).name")
        self.assertEqual(s.age, 16, "Student('Sally', 16, 11, 3.7).age")
        self.assertEqual(s.grade, 11, "Student('Sally', 16, 11, 3.7).grade")
        self.assertEqual(s.gpa, 3.7, "Student('Sally', 16, 11, 3.7).gpa")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="inheritance_mixed3" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class called <code class="docutils literal notranslate"><span class="pre">Building</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">floors</span></code>, <code class="docutils literal notranslate"><span class="pre">sqft</span></code>, and <code class="docutils literal notranslate"><span class="pre">num_doors</span></code>. Next, create a class called <code class="docutils literal notranslate"><span class="pre">House</span></code> with a
<code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits from the <code class="docutils literal notranslate"><span class="pre">Building</span></code> constructor and additionally has parameters <code class="docutils literal notranslate"><span class="pre">bedrooms</span></code>, <code class="docutils literal notranslate"><span class="pre">bathrooms</span></code>, and <code class="docutils literal notranslate"><span class="pre">acres</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">Building(15,</span> <span class="pre">10000,</span> <span class="pre">70).num_doors</span></code> would be <code class="docutils literal notranslate"><span class="pre">70</span></code>
and <code class="docutils literal notranslate"><span class="pre">House(3,</span> <span class="pre">2000,</span> <span class="pre">10,</span> <span class="pre">4,</span> <span class="pre">4,</span> <span class="pre">0.5).acres</span></code> would be <code class="docutils literal notranslate"><span class="pre">0.5</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="5"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class Building:
---
    def __init__(self, floors, sqft, num_doors):
---
    def __init__(floors, sqft, num_doors):
---
        self.floors = floors
        self.sqft = sqft
        self.num_doors = num_doors
---
class House(Building):
---
class House: #paired
---
    def __init__(self, floors, sqft, num_doors, bedrooms, bathrooms, acres):
---
    def __init__(self, floors, sqft, bedrooms, bathrooms, acres): #paired
---
        super().__init__(floors, sqft, num_doors)
---
        self.bedrooms = bedrooms
        self.bathrooms = bathrooms
        self.acres = acres
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=inheritance_mixed3_ac data-question_label="6">
<div id=inheritance_mixed3_ac_question class="ac_question col-md-12">
<p>Write a class called <code class="docutils literal notranslate"><span class="pre">Building</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">floors</span></code>, <code class="docutils literal notranslate"><span class="pre">sqft</span></code>, and <code class="docutils literal notranslate"><span class="pre">num_doors</span></code>. Next, create a class called <code class="docutils literal notranslate"><span class="pre">House</span></code> with a
<code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits from the <code class="docutils literal notranslate"><span class="pre">Building</span></code> constructor and additionally has parameters <code class="docutils literal notranslate"><span class="pre">bedrooms</span></code>, <code class="docutils literal notranslate"><span class="pre">bathrooms</span></code>, and <code class="docutils literal notranslate"><span class="pre">acres</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">Building(15,</span> <span class="pre">10000,</span> <span class="pre">70).num_doors</span></code> would be <code class="docutils literal notranslate"><span class="pre">70</span></code>
and <code class="docutils literal notranslate"><span class="pre">House(3,</span> <span class="pre">2000,</span> <span class="pre">10,</span> <span class="pre">4,</span> <span class="pre">4,</span> <span class="pre">0.5).acres</span></code> would be <code class="docutils literal notranslate"><span class="pre">0.5</span></code>.</p>

</div>
<textarea data-lang="python" id="inheritance_mixed3_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        b = Building(15, 10000, 70)
        h = House(3, 2000, 10, 4, 4, 0.5)

        self.assertEqual(b.floors, 15, "Building(15, 10000, 70).floors")
        self.assertEqual(b.sqft, 10000, "Building(15, 10000, 70).sqft")
        self.assertEqual(b.num_doors, 70, "Building(15, 10000, 70).num_doors")

        self.assertEqual(h.floors, 3, "House(3, 2000, 10, 4, 4, 0.5).floors")
        self.assertEqual(h.sqft, 2000, "House(3, 2000, 10, 4, 4, 0.5).sqft")
        self.assertEqual(h.num_doors, 10, "House(3, 2000, 10, 4, 4, 0.5).num_doors")
        self.assertEqual(h.bedrooms, 4, "House(3, 2000, 10, 4, 4, 0.5).bedrooms")
        self.assertEqual(h.bathrooms, 4, "House(3, 2000, 10, 4, 4, 0.5).bathrooms")
        self.assertEqual(h.acres, 0.5, "House(3, 2000, 10, 4, 4, 0.5).acres")


myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="inheritance_mixed4" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class called <code class="docutils literal notranslate"><span class="pre">Sport</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">teamsize</span></code> and <code class="docutils literal notranslate"><span class="pre">equipment</span></code> and a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns
<code class="docutils literal notranslate"><span class="pre">&quot;teamsize:</span> <span class="pre">(teamsize),</span> <span class="pre">equipment:</span> <span class="pre">(equipment)&quot;</span></code>. Next, create a class called <code class="docutils literal notranslate"><span class="pre">Tennis</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits from the <code class="docutils literal notranslate"><span class="pre">Sport</span></code>
constructor and additionally has a parameter <code class="docutils literal notranslate"><span class="pre">type_court</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">Sport(20,</span> <span class="pre">'football').__str__()</span></code> would return
<code class="docutils literal notranslate"><span class="pre">&quot;teamsize:</span> <span class="pre">20,</span> <span class="pre">equipment:</span> <span class="pre">football&quot;</span></code>, and <code class="docutils literal notranslate"><span class="pre">Tennis(2,</span> <span class="pre">'racquet',</span> <span class="pre">'clay').type_court</span></code> would be <code class="docutils literal notranslate"><span class="pre">'clay'</span></code>.  In the constructor for <code class="docutils literal notranslate"><span class="pre">Tennis</span></code> first call the parent constructor to set the <code class="docutils literal notranslate"><span class="pre">teamsize</span></code> and <code class="docutils literal notranslate"><span class="pre">equipment</span></code>.  Then set the <code class="docutils literal notranslate"><span class="pre">type_court</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="7"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class Sport:
---
    def __init__(self, teamsize, equipment):
---
        self.teamsize = teamsize
        self.equipment = equipment
---
    def __str__(self):
---
        return "teamsize: " + str(self.teamsize) + ", equipment: " + str(self.equipment)
---
class Tennis(Sport):
---
    def __init__(self, teamsize, equipment, type_court):
---
        super().__init__(teamsize, equipment)
---
        self.type_court = type_court
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=inheritance_mixed4_ac data-question_label="8">
<div id=inheritance_mixed4_ac_question class="ac_question col-md-12">
<p>Write a class called <code class="docutils literal notranslate"><span class="pre">Sport</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">teamsize</span></code> and <code class="docutils literal notranslate"><span class="pre">equipment</span></code> and a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns
<code class="docutils literal notranslate"><span class="pre">&quot;teamsize:</span> <span class="pre">(teamsize),</span> <span class="pre">equipment:</span> <span class="pre">(equipment)&quot;</span></code>. Next, create a class called <code class="docutils literal notranslate"><span class="pre">Tennis</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits from the <code class="docutils literal notranslate"><span class="pre">Sport</span></code>
constructor and additionally has a parameter <code class="docutils literal notranslate"><span class="pre">type_court</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">Sport(20,</span> <span class="pre">'football').__str__()</span></code> would return
<code class="docutils literal notranslate"><span class="pre">&quot;teamsize:</span> <span class="pre">20,</span> <span class="pre">equipment:</span> <span class="pre">football&quot;</span></code>, and <code class="docutils literal notranslate"><span class="pre">Tennis(2,</span> <span class="pre">'racquet',</span> <span class="pre">'clay').type_court</span></code> would be <code class="docutils literal notranslate"><span class="pre">'clay'</span></code>.</p>

</div>
<textarea data-lang="python" id="inheritance_mixed4_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        s = Sport(20, 'football')
        t = Tennis(2, 'racquet', 'clay')

        self.assertEqual(s.teamsize, 20, "Sport(20, 'football').teamsize")
        self.assertEqual(s.equipment, "football", "Sport(20, 'football').equipment")
        self.assertEqual(s.__str__(), "teamsize: 20, equipment: football", "Sport(20, 'football').__str__()")

        self.assertEqual(t.teamsize, 2, "Tennis(2, 'racquet', 'clay').teamsize")
        self.assertEqual(t.equipment, 'racquet', "Tennis(2, 'racquet', 'clay').equipment")
        self.assertEqual(t.type_court, 'clay', "Tennis(2, 'racquet', 'clay').type_court")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="inheritance_mixed5" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class called <code class="docutils literal notranslate"><span class="pre">Animal</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">name</span></code> and <code class="docutils literal notranslate"><span class="pre">age</span></code> and a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns
<code class="docutils literal notranslate"><span class="pre">&quot;name:</span> <span class="pre">(name),</span> <span class="pre">age:</span> <span class="pre">(age)&quot;</span></code>. Next, create a class called <code class="docutils literal notranslate"><span class="pre">Dog</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits from the <code class="docutils literal notranslate"><span class="pre">Animal</span></code>
<code class="docutils literal notranslate"><span class="pre">constructor</span></code> and additionally has parameters <code class="docutils literal notranslate"><span class="pre">swim</span></code>, which is assigned to <code class="docutils literal notranslate"><span class="pre">True</span></code> if the <code class="docutils literal notranslate"><span class="pre">Dog</span></code> can swim and <code class="docutils literal notranslate"><span class="pre">False</span></code> otherwise, and <code class="docutils literal notranslate"><span class="pre">trained</span></code>,
which is assigned to an integer that rates how trained the <code class="docutils literal notranslate"><span class="pre">Dog</span></code> is on a scale from 1 to 10. Additionally, the <code class="docutils literal notranslate"><span class="pre">Dog</span></code> class has a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that
inherits from the <code class="docutils literal notranslate"><span class="pre">Animal</span></code> <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method and returns “name: (name), age: (age), swim: (swim), trained: (trained)”. For example,
<code class="docutils literal notranslate"><span class="pre">Animal('charlie',</span> <span class="pre">4).__str__()</span></code> would return <code class="docutils literal notranslate"><span class="pre">&quot;name:</span> <span class="pre">charlie,</span> <span class="pre">age:</span> <span class="pre">4&quot;</span></code>, and <code class="docutils literal notranslate"><span class="pre">Dog('coco',</span> <span class="pre">2,</span> <span class="pre">True,</span> <span class="pre">7).trained</span></code> would be <code class="docutils literal notranslate"><span class="pre">7</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="9"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class Animal:
---
    def __init__(self, name, age):
---
        self.name = name
        self.age = age
---
    def __str__(self):
---
        return "name: " + str(self.name) + ", age: " + str(self.age)
---
class Dog(Animal):
---
    def __init__(self, name, age, swim, trained):
---
        super().__init__(name, age)
---
        self.swim = swim
        self.trained = trained
---
    def __str__(self):
---
        return super().__str__() + ", swim: " + str(self.swim) + ", trained: " + str(self.trained)
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=inheritance_mixed5_ac data-question_label="10">
<div id=inheritance_mixed5_ac_question class="ac_question col-md-12">
<p>Write a class called <code class="docutils literal notranslate"><span class="pre">Animal</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">name</span></code> and <code class="docutils literal notranslate"><span class="pre">age</span></code> and a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns
<code class="docutils literal notranslate"><span class="pre">&quot;name:</span> <span class="pre">(name),</span> <span class="pre">age:</span> <span class="pre">(age)&quot;</span></code>. Next, create a class called <code class="docutils literal notranslate"><span class="pre">Dog</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits from the <code class="docutils literal notranslate"><span class="pre">Animal</span></code>
<code class="docutils literal notranslate"><span class="pre">constructor</span></code> and additionally has parameters <code class="docutils literal notranslate"><span class="pre">swim</span></code>, which is assigned to <code class="docutils literal notranslate"><span class="pre">True</span></code> if the <code class="docutils literal notranslate"><span class="pre">Dog</span></code> can swim and <code class="docutils literal notranslate"><span class="pre">False</span></code> otherwise, and <code class="docutils literal notranslate"><span class="pre">trained</span></code>,
which is assigned to an integer that rates how trained the <code class="docutils literal notranslate"><span class="pre">Dog</span></code> is on a scale from 1 to 10. Additionally, the <code class="docutils literal notranslate"><span class="pre">Dog</span></code> class has a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that
inherits from the <code class="docutils literal notranslate"><span class="pre">Animal</span></code> <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method and returns “name: (name), age: (age), swim: (swim), trained: (trained)”. For example,
<code class="docutils literal notranslate"><span class="pre">Animal('charlie',</span> <span class="pre">4).__str__()</span></code> would return <code class="docutils literal notranslate"><span class="pre">&quot;name:</span> <span class="pre">charlie,</span> <span class="pre">age:</span> <span class="pre">4&quot;</span></code>, and <code class="docutils literal notranslate"><span class="pre">Dog('coco',</span> <span class="pre">2,</span> <span class="pre">True,</span> <span class="pre">7).trained</span></code> would be <code class="docutils literal notranslate"><span class="pre">7</span></code>.</p>

</div>
<textarea data-lang="python" id="inheritance_mixed5_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        a = Animal('charlie', 4)
        d = Dog('coco', 2, True, 7)

        self.assertEqual(a.name, 'charlie', "Animal('charlie', 4).name")
        self.assertEqual(a.age, 4, "Animal('charlie', 4).age")
        self.assertEqual(a.__str__(), "name: charlie, age: 4", "Animal('charlie', 4).__str__()")

        self.assertEqual(d.name, 'coco', "Dog('coco', 2, True, 7).name")
        self.assertEqual(d.age, 2, "Dog('coco', 2, True, 7).age")
        self.assertEqual(d.swim, True, "Dog('coco', 2, True, 7).swim")
        self.assertEqual(d.trained, 7, "Dog('coco', 2, True, 7).trained")
        self.assertEqual(d.__str__(), "name: coco, age: 2, swim: True, trained: 7", "Dog('coco', 2, True, 7).__str__()")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="inheritance_mixed6" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class called <code class="docutils literal notranslate"><span class="pre">Book</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">pages</span></code>, <code class="docutils literal notranslate"><span class="pre">table_contents</span></code>, and <code class="docutils literal notranslate"><span class="pre">author</span></code>, and a <code class="docutils literal notranslate"><span class="pre">__str__</span></code>
method that returns <code class="docutils literal notranslate"><span class="pre">&quot;pages:</span> <span class="pre">(pages),</span> <span class="pre">table_contents:</span> <span class="pre">(table_contents),</span> <span class="pre">author:</span> <span class="pre">(author)&quot;</span></code>. Next, create a class called <code class="docutils literal notranslate"><span class="pre">Encyclopedia</span></code> with
a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits from the <code class="docutils literal notranslate"><span class="pre">Book</span></code> constructor and additionally has parameters <code class="docutils literal notranslate"><span class="pre">volumes</span></code> and <code class="docutils literal notranslate"><span class="pre">subject</span></code>. For example,
<code class="docutils literal notranslate"><span class="pre">Book(200,</span> <span class="pre">True,</span> <span class="pre">'JK</span> <span class="pre">Rowling').__str__()</span></code> would return <code class="docutils literal notranslate"><span class="pre">&quot;pages:</span> <span class="pre">200,</span> <span class="pre">table_contents:</span> <span class="pre">True,</span> <span class="pre">author:</span> <span class="pre">JK</span> <span class="pre">Rowling&quot;</span></code>, and
<code class="docutils literal notranslate"><span class="pre">Encyclopedia(800,</span> <span class="pre">True,</span> <span class="pre">'Denis</span> <span class="pre">Diderot',</span> <span class="pre">28,</span> <span class="pre">'science').subject</span></code> would be <code class="docutils literal notranslate"><span class="pre">'science'</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="11"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class Book:
---
    def __init__(self, pages, table_contents, author):
---
        self.pages = pages
        self.table_contents = table_contents
        self.author = author
---
    def __str__(self):
---
        return "pages: " + str(self.pages) + ", table_contents: " + str(self.table_contents) + ", author: " + str(self.author)
---
class Encyclopedia(Book):
---
    def __init__(self, pages, table_contents, author, volumes, subject):
---
        super().__init__(pages, table_contents, author)
---
        self.volumes = volumes
        self.subject = subject
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=inheritance_mixed6_ac data-question_label="12">
<div id=inheritance_mixed6_ac_question class="ac_question col-md-12">
<p>Write a class called <code class="docutils literal notranslate"><span class="pre">Book</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">pages</span></code>, <code class="docutils literal notranslate"><span class="pre">table_contents</span></code>, and <code class="docutils literal notranslate"><span class="pre">author</span></code>, and a <code class="docutils literal notranslate"><span class="pre">__str__</span></code>
method that returns <code class="docutils literal notranslate"><span class="pre">&quot;pages:</span> <span class="pre">(pages),</span> <span class="pre">table_contents:</span> <span class="pre">(table_contents),</span> <span class="pre">author:</span> <span class="pre">(author)&quot;</span></code>. Next, create a class called <code class="docutils literal notranslate"><span class="pre">Encyclopedia</span></code> with
a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits from the <code class="docutils literal notranslate"><span class="pre">Book</span></code> constructor and additionally has parameters <code class="docutils literal notranslate"><span class="pre">volumes</span></code> and <code class="docutils literal notranslate"><span class="pre">subject</span></code>. For example,
<code class="docutils literal notranslate"><span class="pre">Book(200,</span> <span class="pre">True,</span> <span class="pre">'JK</span> <span class="pre">Rowling').__str__()</span></code> would return <code class="docutils literal notranslate"><span class="pre">&quot;pages:</span> <span class="pre">200,</span> <span class="pre">table_contents:</span> <span class="pre">True,</span> <span class="pre">author:</span> <span class="pre">JK</span> <span class="pre">Rowling&quot;</span></code>, and
<code class="docutils literal notranslate"><span class="pre">Encyclopedia(800,</span> <span class="pre">True,</span> <span class="pre">'Denis</span> <span class="pre">Diderot',</span> <span class="pre">28,</span> <span class="pre">'science').subject</span></code> would be <code class="docutils literal notranslate"><span class="pre">'science'</span></code>.</p>

</div>
<textarea data-lang="python" id="inheritance_mixed6_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        b = Book(200, True, 'JK Rowling')
        e = Encyclopedia(800, True, 'Denis Diderot', 28, 'science')

        self.assertEqual(b.pages, 200, "Book(200, True, 'JK Rowling').pages")
        self.assertEqual(b.table_contents, True, "Book(200, True, 'JK Rowling').table_contents")
        self.assertEqual(b.author, 'JK Rowling', "Book(200, True, 'JK Rowling').author")
        self.assertEqual(b.__str__(), "pages: 200, table_contents: True, author: JK Rowling", "Book(200, True, 'JK Rowling').__str__()")

        self.assertEqual(e.pages, 800, "Encyclopedia(800, True, 'Denis Diderot', 28, 'science').pages")
        self.assertEqual(e.table_contents, True, "Encyclopedia(800, True, 'Denis Diderot', 28, 'science').table_contents")
        self.assertEqual(e.author, 'Denis Diderot', "Encyclopedia(800, True, 'Denis Diderot', 28, 'science').author")
        self.assertEqual(e.volumes, 28, "Encyclopedia(800, True, 'Denis Diderot', 28, 'science').volumes")
        self.assertEqual(e.subject, 'science', "Encyclopedia(800, True, 'Denis Diderot', 28, 'science').subject")


myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="inheritance_mixed7" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class called <code class="docutils literal notranslate"><span class="pre">Hat</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">size</span></code>, <code class="docutils literal notranslate"><span class="pre">cost</span></code>, and <code class="docutils literal notranslate"><span class="pre">material</span></code>, and a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns
<code class="docutils literal notranslate"><span class="pre">&quot;size:</span> <span class="pre">(size),</span> <span class="pre">cost:</span> <span class="pre">(cost),</span> <span class="pre">material:</span> <span class="pre">(material)&quot;</span></code>. Next, create a class called <code class="docutils literal notranslate"><span class="pre">Fedora</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits from the <code class="docutils literal notranslate"><span class="pre">Hat</span></code>
<code class="docutils literal notranslate"><span class="pre">constructor</span></code> and additionally has a parameter <code class="docutils literal notranslate"><span class="pre">basketweave</span></code>, which is assigned to <code class="docutils literal notranslate"><span class="pre">True</span></code> if the <code class="docutils literal notranslate"><span class="pre">Fedora</span></code> is made with a <code class="docutils literal notranslate"><span class="pre">basketweave`</span></code> and
<code class="docutils literal notranslate"><span class="pre">False</span></code> otherwise. Additionally, the <code class="docutils literal notranslate"><span class="pre">Fedora</span></code> class has a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that inherits from the <code class="docutils literal notranslate"><span class="pre">Hat</span></code> <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method and returns
<code class="docutils literal notranslate"><span class="pre">&quot;size:</span> <span class="pre">(size),</span> <span class="pre">cost:</span> <span class="pre">(cost),</span> <span class="pre">material:</span> <span class="pre">(material),</span> <span class="pre">basketweave:</span> <span class="pre">(basketweave)&quot;</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">Hat('small',</span> <span class="pre">30,</span> <span class="pre">'Cotton').__str__()</span></code> would
return <code class="docutils literal notranslate"><span class="pre">&quot;size:</span> <span class="pre">small,</span> <span class="pre">cost:</span> <span class="pre">30,</span> <span class="pre">material:</span> <span class="pre">Cotton&quot;</span></code>, and <code class="docutils literal notranslate"><span class="pre">Fedora('large',</span> <span class="pre">50,</span> <span class="pre">'Polyester',</span> <span class="pre">True).material</span></code> would be <code class="docutils literal notranslate"><span class="pre">'Polyester'</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="13"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class Hat:
---
    def __init__(self, size, cost, material):
---
        self.size = size
        self.cost = cost
        self.material = material
---
    def __str__(self):
---
        return "size: " + str(self.size) + ", cost: " + str(self.cost) + ", material: " + str(self.material)
---
class Fedora(Hat):
---
    def __init__(self, size, cost, material, basketweave):
---
        super().__init__(size, cost, material)
---
        self.basketweave = basketweave
---
    def __str__(self):
---
        return "size: " + str(self.size) + ", cost: " + str(self.cost) + ", material: " + str(self.material) + ", basketweave: " + str(self.basketweave)
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=inheritance_mixed7_ac data-question_label="14">
<div id=inheritance_mixed7_ac_question class="ac_question col-md-12">
<p>Write a class called <code class="docutils literal notranslate"><span class="pre">Hat</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">size</span></code>, <code class="docutils literal notranslate"><span class="pre">cost</span></code>, and <code class="docutils literal notranslate"><span class="pre">material</span></code>, and a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns
<code class="docutils literal notranslate"><span class="pre">&quot;size:</span> <span class="pre">(size),</span> <span class="pre">cost:</span> <span class="pre">(cost),</span> <span class="pre">material:</span> <span class="pre">(material)&quot;</span></code>. Next, create a class called <code class="docutils literal notranslate"><span class="pre">Fedora</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits from the <code class="docutils literal notranslate"><span class="pre">Hat</span></code>
<code class="docutils literal notranslate"><span class="pre">constructor</span></code> and additionally has a parameter <code class="docutils literal notranslate"><span class="pre">basketweave</span></code>, which is assigned to <code class="docutils literal notranslate"><span class="pre">True</span></code> if the <code class="docutils literal notranslate"><span class="pre">Fedora</span></code> is made with a <code class="docutils literal notranslate"><span class="pre">basketweave`</span></code> and
<code class="docutils literal notranslate"><span class="pre">False</span></code> otherwise. Additionally, the <code class="docutils literal notranslate"><span class="pre">Fedora</span></code> class has a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that inherits from the <code class="docutils literal notranslate"><span class="pre">Hat</span></code> <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method and returns
<code class="docutils literal notranslate"><span class="pre">&quot;size:</span> <span class="pre">(size),</span> <span class="pre">cost:</span> <span class="pre">(cost),</span> <span class="pre">material:</span> <span class="pre">(material),</span> <span class="pre">basketweave:</span> <span class="pre">(basketweave)&quot;</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">Hat('small',</span> <span class="pre">30,</span> <span class="pre">'Cotton').__str__()</span></code> would
return <code class="docutils literal notranslate"><span class="pre">&quot;size:</span> <span class="pre">small,</span> <span class="pre">cost:</span> <span class="pre">30,</span> <span class="pre">material:</span> <span class="pre">Cotton&quot;</span></code>, and <code class="docutils literal notranslate"><span class="pre">Fedora('large',</span> <span class="pre">50,</span> <span class="pre">'Polyester',</span> <span class="pre">True).material</span></code> would be <code class="docutils literal notranslate"><span class="pre">'Polyester'</span></code>.</p>

</div>
<textarea data-lang="python" id="inheritance_mixed7_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        h = Hat('small', 30, 'Cotton')
        f = Fedora('large', 50, 'Polyester', True)

        self.assertEqual(h.size, 'small', "Hat('small', 30, 'Cotton').size")
        self.assertEqual(h.cost, 30, "Hat('small', 30, 'Cotton').cost")
        self.assertEqual(h.material, 'Cotton', "Hat('small', 30, 'Cotton').material")
        self.assertEqual(h.__str__(), "size: small, cost: 30, material: Cotton", "Hat('small', 30, 'Cotton').__str__()")

        self.assertEqual(f.size, 'large', "Fedora('large', 50, 'Polyester', True).size")
        self.assertEqual(f.cost, 50, "Fedora('large', 50, 'Polyester', True).cost")
        self.assertEqual(f.material, 'Polyester', "Fedora('large', 50, 'Polyester', True).material")
        self.assertEqual(f.basketweave, True, "Fedora('large', 50, 'Polyester', True).basketweave")
        self.assertEqual(f.__str__(), "size: large, cost: 50, material: Polyester, basketweave: True", "Fedora('large', 50, 'Polyester', True).__str__()")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="inheritance_mixed8" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class called <code class="docutils literal notranslate"><span class="pre">Vehicle</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">mph</span></code>, <code class="docutils literal notranslate"><span class="pre">color</span></code>, and <code class="docutils literal notranslate"><span class="pre">weight</span></code>. Next, create a class called <code class="docutils literal notranslate"><span class="pre">Car</span></code> with a
<code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits from the <code class="docutils literal notranslate"><span class="pre">Vehicle</span></code> constructor and additionally has a parameter <code class="docutils literal notranslate"><span class="pre">electric</span></code>. Lastly, create a class called <code class="docutils literal notranslate"><span class="pre">ElectricCar</span></code> with a
<code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits from the <code class="docutils literal notranslate"><span class="pre">Car</span></code> constructor and additionally has a parameter <code class="docutils literal notranslate"><span class="pre">cost</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">Vehicle(3.5,</span> <span class="pre">'grey',</span> <span class="pre">2000).weight</span></code> would be <code class="docutils literal notranslate"><span class="pre">2000</span></code>,
<code class="docutils literal notranslate"><span class="pre">Car(40,</span> <span class="pre">'blue',</span> <span class="pre">4000,</span> <span class="pre">False).electric</span></code> would be <code class="docutils literal notranslate"><span class="pre">False</span></code>, and <code class="docutils literal notranslate"><span class="pre">ElectricCar(40,</span> <span class="pre">'red',</span> <span class="pre">3000,</span> <span class="pre">True,</span> <span class="pre">30000).cost</span></code> would be <code class="docutils literal notranslate"><span class="pre">30000</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="15"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class Vehicle:
---
    def __init__(self, mph, color, weight):
---
    def init(self, mph, color, weight): #paired
---
        self.mph = mph
        self.color = color
        self.weight = weight
---
class Car(Vehicle):
---
    def __init__(self, mph, color, weight, electric):
---
        super().__init__(mph, color, weight)
---
        self.electric = electric
---
class ElectricCar(Car):
---
class ElectricCar(Vehicle): #paired
---
    def __init__(self, mph, color, weight, electric, cost):
---
        super().__init__(mph, color, weight, electric)
---
        super().__init__(mph, color, weight, electric, cost) #paired
---
        self.cost = cost
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=inheritance_mixed8_ac data-question_label="16">
<div id=inheritance_mixed8_ac_question class="ac_question col-md-12">
<p>Write a class called <code class="docutils literal notranslate"><span class="pre">Vehicle</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">mph</span></code>, <code class="docutils literal notranslate"><span class="pre">color</span></code>, and <code class="docutils literal notranslate"><span class="pre">weight</span></code>. Next, create a class called <code class="docutils literal notranslate"><span class="pre">Car</span></code> with a
<code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits from the <code class="docutils literal notranslate"><span class="pre">Vehicle</span></code> constructor and additionally has a parameter <code class="docutils literal notranslate"><span class="pre">electric</span></code>. Lastly, create a class called <code class="docutils literal notranslate"><span class="pre">ElectricCar</span></code> with a
<code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits from the <code class="docutils literal notranslate"><span class="pre">Car</span></code> constructor and additionally has a parameter <code class="docutils literal notranslate"><span class="pre">cost</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">Vehicle(3.5,</span> <span class="pre">'grey',</span> <span class="pre">2000).weight</span></code> would be <code class="docutils literal notranslate"><span class="pre">2000</span></code>,
<code class="docutils literal notranslate"><span class="pre">Car(40,</span> <span class="pre">'blue',</span> <span class="pre">4000,</span> <span class="pre">False).electric</span></code> would be <code class="docutils literal notranslate"><span class="pre">False</span></code>, and <code class="docutils literal notranslate"><span class="pre">ElectricCar(40,</span> <span class="pre">'red',</span> <span class="pre">3000,</span> <span class="pre">True,</span> <span class="pre">30000).cost</span></code> would be <code class="docutils literal notranslate"><span class="pre">30000</span></code>.</p>

</div>
<textarea data-lang="python" id="inheritance_mixed8_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        v = Vehicle(3.5, 'grey', 2000)
        c = Car(40, 'blue', 4000, False)
        e = ElectricCar(40, 'red', 3000, True, 30000)

        self.assertEqual(v.mph, 3.5, "Vehicle(3.5, 'grey', 2000).mph")
        self.assertEqual(v.color, 'grey', "Vehicle(3.5, 'grey', 2000).color")
        self.assertEqual(v.weight, 2000, "Vehicle(3.5, 'grey', 2000).weight")

        self.assertEqual(c.mph, 40, "Car(40, 'blue', 4000, False).mph")
        self.assertEqual(c.color, 'blue', "Car(40, 'blue', 4000, False).color")
        self.assertEqual(c.weight, 4000, "Car(40, 'blue', 4000, False).weight")
        self.assertEqual(c.electric, False, "Car(40, 'blue', 4000, False).electric")

        self.assertEqual(e.mph, 40, "ElectricCar(40, 'red', 3000, True, 30000).mph")
        self.assertEqual(e.color, 'red', "ElectricCar(40, 'red', 3000, True, 30000).color")
        self.assertEqual(e.weight, 3000, "ElectricCar(40, 'red', 3000, True, 30000).weight")
        self.assertEqual(e.electric, True, "ElectricCar(40, 'red', 3000, True, 30000).electric")
        self.assertEqual(e.cost, 30000, "ElectricCar(40, 'red', 3000, True, 30000).cost")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="inheritance_mixed9" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class called <code class="docutils literal notranslate"><span class="pre">SchoolSupplies</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">size</span></code>, <code class="docutils literal notranslate"><span class="pre">cost</span></code>, and <code class="docutils literal notranslate"><span class="pre">material</span></code>, and a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method
that returns <code class="docutils literal notranslate"><span class="pre">&quot;size:</span> <span class="pre">(size),</span> <span class="pre">cost:</span> <span class="pre">(cost),</span> <span class="pre">material:</span> <span class="pre">(material)&quot;</span></code>. Next, create a class called <code class="docutils literal notranslate"><span class="pre">WritingUtensils</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits
from the <code class="docutils literal notranslate"><span class="pre">SchoolSupplies</span></code> constructor and additionally has a parameter <code class="docutils literal notranslate"><span class="pre">count</span></code>, and a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns <code class="docutils literal notranslate"><span class="pre">&quot;size:</span> <span class="pre">(size),</span> <span class="pre">cost:</span> <span class="pre">(cost),</span>
<span class="pre">material:</span> <span class="pre">(material),</span> <span class="pre">count:</span> <span class="pre">(count)&quot;</span></code>. Lastly, create a class called <code class="docutils literal notranslate"><span class="pre">Pencil</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits from the <code class="docutils literal notranslate"><span class="pre">WritingUtensils</span></code>
constructor and additionally has parameters <code class="docutils literal notranslate"><span class="pre">mechanical</span></code> and <code class="docutils literal notranslate"><span class="pre">lead</span></code>. If a <code class="docutils literal notranslate"><span class="pre">Pencil</span></code> is <code class="docutils literal notranslate"><span class="pre">mechanical</span></code>, it will have a float <code class="docutils literal notranslate"><span class="pre">lead</span></code> size. Otherwise, <code class="docutils literal notranslate"><span class="pre">lead</span></code> will be <code class="docutils literal notranslate"><span class="pre">None</span></code>.
Additionally, the <code class="docutils literal notranslate"><span class="pre">Pencil</span></code> class should have a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns <code class="docutils literal notranslate"><span class="pre">&quot;size:</span> <span class="pre">(size),</span> <span class="pre">cost:</span> <span class="pre">(cost),</span> <span class="pre">material:</span> <span class="pre">(material),</span> <span class="pre">count:</span> <span class="pre">(count),</span>
<span class="pre">mechanical:</span> <span class="pre">(mechanical),</span> <span class="pre">lead:</span> <span class="pre">(lead)&quot;</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">WritingUtensils('medium',</span> <span class="pre">20,</span> <span class="pre">'wood',</span> <span class="pre">12).__str__()</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;size:</span> <span class="pre">medium,</span> <span class="pre">cost:</span> <span class="pre">20,</span> <span class="pre">material:</span>
<span class="pre">wood,</span> <span class="pre">count:</span> <span class="pre">12&quot;</span></code> and <code class="docutils literal notranslate"><span class="pre">Pencil('large',</span> <span class="pre">30,</span> <span class="pre">'wood',</span> <span class="pre">24,</span> <span class="pre">True,</span> <span class="pre">0.7).mechanical</span></code> would be <code class="docutils literal notranslate"><span class="pre">True</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="17"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class SchoolSupplies:
---
    def __init__(self, size, cost, material):
---
        self.size = size
        self.cost = cost
        self.material = material
---
    def __str__(self):
---
        return "size: " + str(self.size) + ", cost: " + str(self.cost) + ", material: " + str(self.material)
---
class WritingUtensils(SchoolSupplies):
---
    def __init__(self, size, cost, material, count):
---
        super().__init__(size, cost, material)
---
        self.count = count
---
    def __str__(self):
---
        return "size: " + str(self.size) + ", cost: " + str(self.cost) + ", material: " + str(self.material) + ", count: " + str(self.count)
---
class Pencil(WritingUtensils):
---
    def __init__(self, size, cost, material, count, mechanical, lead):
---
        super().__init__(size, cost, material, count)
---
        self.mechanical = mechanical
---
        if self.mechanical == True:
---
            self.lead = lead
---
        else:
---
            self.lead = None
---
    def __str__(self):
---
        return "size: " + str(self.size) + ", cost: " + str(self.cost) + ", material: " + str(self.material) + ", count: " + str(self.count) + ", mechanical: " + str(self.mechanical) + ", lead: " + str(self.lead)
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=inheritance_mixed9_ac data-question_label="18">
<div id=inheritance_mixed9_ac_question class="ac_question col-md-12">
<p>Write a class called <code class="docutils literal notranslate"><span class="pre">SchoolSupplies</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">size</span></code>, <code class="docutils literal notranslate"><span class="pre">cost</span></code>, and <code class="docutils literal notranslate"><span class="pre">material</span></code>, and a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method
that returns <code class="docutils literal notranslate"><span class="pre">&quot;size:</span> <span class="pre">(size),</span> <span class="pre">cost:</span> <span class="pre">(cost),</span> <span class="pre">material:</span> <span class="pre">(material)&quot;</span></code>. Next, create a class called <code class="docutils literal notranslate"><span class="pre">WritingUtensils</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits
from the <code class="docutils literal notranslate"><span class="pre">SchoolSupplies</span></code> constructor and additionally has a parameter <code class="docutils literal notranslate"><span class="pre">count</span></code>, and a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns <code class="docutils literal notranslate"><span class="pre">&quot;size:</span> <span class="pre">(size),</span> <span class="pre">cost:</span> <span class="pre">(cost),</span>
<span class="pre">material:</span> <span class="pre">(material),</span> <span class="pre">count:</span> <span class="pre">(count)&quot;</span></code>. Lastly, create a class called <code class="docutils literal notranslate"><span class="pre">Pencil</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits from the <code class="docutils literal notranslate"><span class="pre">WritingUtensils</span></code>
constructor and additionally has parameters <code class="docutils literal notranslate"><span class="pre">mechanical</span></code> and <code class="docutils literal notranslate"><span class="pre">lead</span></code>. If a <code class="docutils literal notranslate"><span class="pre">Pencil</span></code> is <code class="docutils literal notranslate"><span class="pre">mechanical</span></code>, it will have a float <code class="docutils literal notranslate"><span class="pre">lead</span></code> size. Otherwise, <code class="docutils literal notranslate"><span class="pre">lead</span></code> will be <code class="docutils literal notranslate"><span class="pre">None</span></code>.
Additionally, the <code class="docutils literal notranslate"><span class="pre">Pencil</span></code> class should have a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns <code class="docutils literal notranslate"><span class="pre">&quot;size:</span> <span class="pre">(size),</span> <span class="pre">cost:</span> <span class="pre">(cost),</span> <span class="pre">material:</span> <span class="pre">(material),</span> <span class="pre">count:</span> <span class="pre">(count),</span>
<span class="pre">mechanical:</span> <span class="pre">(mechanical),</span> <span class="pre">lead:</span> <span class="pre">(lead)&quot;</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">WritingUtensils('medium',</span> <span class="pre">20,</span> <span class="pre">'wood',</span> <span class="pre">12).__str__()</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;size:</span> <span class="pre">medium,</span> <span class="pre">cost:</span> <span class="pre">20,</span> <span class="pre">material:</span>
<span class="pre">wood,</span> <span class="pre">count:</span> <span class="pre">12&quot;</span></code> and <code class="docutils literal notranslate"><span class="pre">Pencil('large',</span> <span class="pre">30,</span> <span class="pre">'wood',</span> <span class="pre">24,</span> <span class="pre">True,</span> <span class="pre">0.7).mechanical</span></code> would be <code class="docutils literal notranslate"><span class="pre">True</span></code>.</p>

</div>
<textarea data-lang="python" id="inheritance_mixed9_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        s = SchoolSupplies('small', 10, 'plastic')
        w = WritingUtensils('medium', 20, 'wood', 12)
        p1 = Pencil('large', 30, 'wood', 24, True, 0.7)
        p2 = Pencil('large', 30, 'wood', 24, False, 0.5)

        self.assertEqual(s.size, 'small', "SchoolSupplies('small', 10, 'plastic').size")
        self.assertEqual(s.cost, 10, "SchoolSupplies('small', 10, 'plastic').cost")
        self.assertEqual(s.material, 'plastic', "SchoolSupplies('small', 10, 'plastic').material")
        self.assertEqual(s.__str__(), "size: small, cost: 10, material: plastic", "SchoolSupplies('small', 10, 'plastic').__str__()")

        self.assertEqual(w.size, 'medium', "WritingUtensils('medium', 20, 'wood', 12).size")
        self.assertEqual(w.cost, 20, "WritingUtensils('medium', 20, 'wood', 12).cost")
        self.assertEqual(w.material, 'wood', "WritingUtensils('medium', 20, 'wood', 12).material")
        self.assertEqual(w.count, 12, "WritingUtensils('medium', 20, 'wood', 12).count")
        self.assertEqual(w.__str__(), "size: medium, cost: 20, material: wood, count: 12", "WritingUtensils('medium', 20, 'wood', 12).__str__()")

        self.assertEqual(p1.size, 'large', "Pencil('large', 30, 'wood', 24, True, 0.7).size")
        self.assertEqual(p1.cost, 30, "Pencil('large', 30, 'wood', 24, True, 0.7).cost")
        self.assertEqual(p1.material, 'wood', "Pencil('large', 30, 'wood', 24, True, 0.7).material")
        self.assertEqual(p1.count, 24, "Pencil('large', 30, 'wood', 24, True, 0.7).count")
        self.assertEqual(p1.mechanical, True, "Pencil('large', 30, 'wood', 24, True, 0.7).mechanical")
        self.assertEqual(p1.lead, 0.7, "Pencil('large', 30, 'wood', 24, True, 0.7).lead")
        self.assertEqual(p1.__str__(), "size: large, cost: 30, material: wood, count: 24, mechanical: True, lead: 0.7", "Pencil('large', 30, 'wood', 24, True, 0.7).__str__()")

        self.assertEqual(p2.size, 'large', "Pencil('large', 30, 'wood', 24, False, 0.5).size")
        self.assertEqual(p2.cost, 30, "Pencil('large', 30, 'wood', 24, False, 0.5).cost")
        self.assertEqual(p2.material, 'wood', "Pencil('large', 30, 'wood', 24, False, 0.5).material")
        self.assertEqual(p2.count, 24, "Pencil('large', 30, 'wood', 24, False, 0.5).count")
        self.assertEqual(p2.mechanical, False, "Pencil('large', 30, 'wood', 24, False, 0.5).mechancial")
        self.assertEqual(p2.lead, None, "Pencil('large', 30, 'wood', 24, False, 0.5).lead")
        self.assertEqual(p2.__str__(), "size: large, cost: 30, material: wood, count: 24, mechanical: False, lead: None", "Pencil('large', 30, 'wood', 24, False, 0.5).__str__()")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="inheritance_mixed10" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class called <code class="docutils literal notranslate"><span class="pre">FederalCourts</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">size</span></code> and <code class="docutils literal notranslate"><span class="pre">location</span></code>, a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method
that returns <code class="docutils literal notranslate"><span class="pre">&quot;size:</span> <span class="pre">(size),</span> <span class="pre">location:</span> <span class="pre">(location)&quot;</span></code>, and a method called <code class="docutils literal notranslate"><span class="pre">judgment</span></code> that takes in a string parameter <code class="docutils literal notranslate"><span class="pre">decision</span></code>.
If <code class="docutils literal notranslate"><span class="pre">decision</span></code> is <code class="docutils literal notranslate"><span class="pre">'guilty'</span></code>, return <code class="docutils literal notranslate"><span class="pre">'Conviction</span> <span class="pre">needed'</span></code>. If <code class="docutils literal notranslate"><span class="pre">decision</span></code> is <code class="docutils literal notranslate"><span class="pre">'not</span> <span class="pre">guilty'</span></code>, return <code class="docutils literal notranslate"><span class="pre">'Conviction</span> <span class="pre">not</span> <span class="pre">needed'</span></code>.
Otherwise, return <code class="docutils literal notranslate"><span class="pre">'Inconclusive'</span></code>. Next, create a class called <code class="docutils literal notranslate"><span class="pre">SupremeCourt</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits
from the <code class="docutils literal notranslate"><span class="pre">FederalCourts</span></code> constructor and additionally has a parameter <code class="docutils literal notranslate"><span class="pre">leaning_party</span></code>, and a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns <code class="docutils literal notranslate"><span class="pre">&quot;size:</span> <span class="pre">(size),</span> <span class="pre">location:</span> <span class="pre">(location),</span> <span class="pre">leaning_party:</span> <span class="pre">(leaning_party)&quot;</span></code>.
Lastly, create a class called <code class="docutils literal notranslate"><span class="pre">CourtsOfAppeals</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits from the <code class="docutils literal notranslate"><span class="pre">FederalCourts</span></code>
constructor and additionally has a parameter <code class="docutils literal notranslate"><span class="pre">divisions</span></code>. Additionally, the <code class="docutils literal notranslate"><span class="pre">CourtsOfAppeals</span></code> class should have a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns
<code class="docutils literal notranslate"><span class="pre">&quot;size:</span> <span class="pre">(size),</span> <span class="pre">location:</span> <span class="pre">(location),</span> <span class="pre">divisions:</span> <span class="pre">(divisions)&quot;</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">FederalCourts(94,</span> <span class="pre">'New</span> <span class="pre">York').__str__()</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;size:</span> <span class="pre">94,</span> <span class="pre">location:</span> <span class="pre">New</span> <span class="pre">York&quot;</span></code>,
<code class="docutils literal notranslate"><span class="pre">CourtsOfAppeals(179,</span> <span class="pre">'Pennsylvania',</span> <span class="pre">23).divisions</span></code> would be <code class="docutils literal notranslate"><span class="pre">23</span></code>, and <code class="docutils literal notranslate"><span class="pre">FederalCourts(94,</span> <span class="pre">'New</span> <span class="pre">York').judgment('Not</span> <span class="pre">Guilty')</span></code> would be <code class="docutils literal notranslate"><span class="pre">&quot;Conviction</span> <span class="pre">not</span> <span class="pre">needed&quot;</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="19"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class FederalCourts:
---
    def __init__(self, size, location):
---
        self.size = size
        self.location = location
---
    def __str__(self):
---
        return "size: " + str(self.size) + ", location: " + str(self.location)
---
    def judgment(self, decision):
---
        if decision.lower() == 'guilty':
---
        if decision == 'Guilty'.lower(): #paired
---
            return 'Conviction needed'
---
        elif decision.lower() == 'not guilty':
---
        elif decision == 'Not Guilty'.lower(): #paired
---
            return 'Conviction not needed'
---
        else:
---
            return 'Inconclusive'
---
class SupremeCourt(FederalCourts):
---
    def __init__(self, size, location, leaning_party):
---
        super().__init__(size, location)
---
        self.leaning_party = leaning_party
---
    def __str__(self):
---
        return "size: " + str(self.size) + ", location: " + str(self.location) + ", leaning_party: " + str(self.leaning_party)
---
class CourtsOfAppeals(FederalCourts):
---
    def __init__(self, size, location, divisions):
---
        super().__init__(size, location)
---
        self.divisions = divisions
---
    def __str__(self):
---
        return "size: " + str(self.size) + ", location: " + str(self.location) + ", divisions: " + str(self.divisions)
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=inheritance_mixed10_ac data-question_label="20">
<div id=inheritance_mixed10_ac_question class="ac_question col-md-12">
<p>Write a class called <code class="docutils literal notranslate"><span class="pre">FederalCourts</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">size</span></code> and <code class="docutils literal notranslate"><span class="pre">location</span></code>, a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method
that returns <code class="docutils literal notranslate"><span class="pre">&quot;size:</span> <span class="pre">(size),</span> <span class="pre">location:</span> <span class="pre">(location)&quot;</span></code>, and a method called <code class="docutils literal notranslate"><span class="pre">judgment</span></code> that takes in a string parameter <code class="docutils literal notranslate"><span class="pre">decision</span></code>.
If <code class="docutils literal notranslate"><span class="pre">decision</span></code> is <code class="docutils literal notranslate"><span class="pre">'guilty'</span></code>, return <code class="docutils literal notranslate"><span class="pre">'Conviction</span> <span class="pre">needed'</span></code>. If <code class="docutils literal notranslate"><span class="pre">decision</span></code> is <code class="docutils literal notranslate"><span class="pre">'not</span> <span class="pre">guilty'</span></code>, return <code class="docutils literal notranslate"><span class="pre">'Conviction</span> <span class="pre">not</span> <span class="pre">needed'</span></code>.
Otherwise, return <code class="docutils literal notranslate"><span class="pre">'Inconclusive'</span></code>. Next, create a class called <code class="docutils literal notranslate"><span class="pre">SupremeCourt</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits
from the <code class="docutils literal notranslate"><span class="pre">FederalCourts</span></code> constructor and additionally has a parameter <code class="docutils literal notranslate"><span class="pre">leaning_party</span></code>, and a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns <code class="docutils literal notranslate"><span class="pre">&quot;size:</span> <span class="pre">(size),</span> <span class="pre">location:</span> <span class="pre">(location),</span> <span class="pre">leaning_party:</span> <span class="pre">(leaning_party)&quot;</span></code>.
Lastly, create a class called <code class="docutils literal notranslate"><span class="pre">CourtsOfAppeals</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inherits from the <code class="docutils literal notranslate"><span class="pre">FederalCourts</span></code>
constructor and additionally has a parameter <code class="docutils literal notranslate"><span class="pre">divisions</span></code>. Additionally, the <code class="docutils literal notranslate"><span class="pre">CourtsOfAppeals</span></code> class should have a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns
<code class="docutils literal notranslate"><span class="pre">&quot;size:</span> <span class="pre">(size),</span> <span class="pre">location:</span> <span class="pre">(location),</span> <span class="pre">divisions:</span> <span class="pre">(divisions)&quot;</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">FederalCourts(94,</span> <span class="pre">'New</span> <span class="pre">York').__str__()</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;size:</span> <span class="pre">94,</span> <span class="pre">location:</span> <span class="pre">New</span> <span class="pre">York&quot;</span></code>,
<code class="docutils literal notranslate"><span class="pre">CourtsOfAppeals(179,</span> <span class="pre">'Pennsylvania',</span> <span class="pre">23).divisions</span></code> would be <code class="docutils literal notranslate"><span class="pre">23</span></code>, and <code class="docutils literal notranslate"><span class="pre">FederalCourts(94,</span> <span class="pre">'New</span> <span class="pre">York').judgment('Not</span> <span class="pre">Guilty')</span></code> would be <code class="docutils literal notranslate"><span class="pre">&quot;Conviction</span> <span class="pre">not</span> <span class="pre">needed&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="inheritance_mixed10_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        fc = FederalCourts(94, 'New York')
        sc = SupremeCourt(9, 'Washington DC', 'Republican')
        coa = CourtsOfAppeals(179, 'Pennsylvania', 23)

        self.assertEqual(fc.size, 94, "FederalCourts(94, 'New York').size")
        self.assertEqual(fc.location, 'New York', "FederalCourts(94, 'New York').location")
        self.assertEqual(fc.__str__(), "size: 94, location: New York", "FederalCourts(94, 'New York').__str__()")
        self.assertEqual(fc.judgment('Guilty'), 'Conviction needed', "FederalCourts(94, 'New York').judgment('Guilty')")
        self.assertEqual(fc.judgment('Not Guilty'), "Conviction not needed", "FederalCourts(94, 'New York').judgment('Not Guilty')")
        self.assertEqual(fc.judgment('DO NOT KNOW'), "Inconclusive", "FederalCourts(94, 'New York').judgment('DO NOT KNOW')")

        self.assertEqual(sc.size, 9, "SupremeCourt(9, 'Washington DC', 'Republican').size")
        self.assertEqual(sc.location, 'Washington DC', "SupremeCourt(9, 'Washington DC', 'Republican').location")
        self.assertEqual(sc.leaning_party, 'Republican', "SupremeCourt(9, 'Washington DC', 'Republican').leaning_party")
        self.assertEqual(sc.__str__(), "size: 9, location: Washington DC, leaning_party: Republican", "SupremeCourt(9, 'Washington DC', 'Republican').__str__()")

        self.assertEqual(coa.size, 179, "CourtsOfAppeals(179, 'Pennsylvania', 23).size")
        self.assertEqual(coa.location, 'Pennsylvania', "CourtsOfAppeals(179, 'Pennsylvania', 23).location")
        self.assertEqual(coa.divisions, 23, "CourtsOfAppeals(179, 'Pennsylvania', 23).divisions")
        self.assertEqual(coa.__str__(), "size: 179, location: Pennsylvania, divisions: 23", "CourtsOfAppeals(179, 'Pennsylvania', 23).__str__()")

myTests().main()
</textarea>
</div>
</div>
</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>