<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Mixed-Up Code Questions &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/637.css?v=bddd2011b20db76a20ea" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.css?v=7c5388ad4389eae1a318" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/runtime.bundle.js?v=aa7c717c38c5405af9e6"></script>
    <script src="../_static/637.bundle.js?v=1123ad75f9b821cdd636"></script>
    <script src="../_static/runestone.bundle.js?v=9c8e466e8cbed4a3af5a"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'py4e-int';
    eBookConfig.basecourse = 'py4e-int';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          py4e-int
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=mixedupcode_orig'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
  
<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="mixed-up-code-questions">
<h1>Mixed-Up Code Questions<a class="headerlink" href="#mixed-up-code-questions" title="Permalink to this headline">¶</a></h1>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="inheritance_mixed1_v3" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class <code class="docutils literal notranslate"><span class="pre">Person</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that takes <code class="docutils literal notranslate"><span class="pre">first_name</span></code> and <code class="docutils literal notranslate"><span class="pre">last_name</span></code> and inititalize those attributes in the current object. Next, create a class called <code class="docutils literal notranslate"><span class="pre">Customer</span></code> that inherits from <code class="docutils literal notranslate"><span class="pre">Person</span></code> with a constructor that takes <code class="docutils literal notranslate"><span class="pre">first_name</span></code>, <code class="docutils literal notranslate"><span class="pre">last_name</span></code>, and <code class="docutils literal notranslate"><span class="pre">id</span></code>.  Call the constructor in <code class="docutils literal notranslate"><span class="pre">Person</span></code> to initialize <code class="docutils literal notranslate"><span class="pre">first_name</span></code> and <code class="docutils literal notranslate"><span class="pre">last_name</span></code> and then also set the <code class="docutils literal notranslate"><span class="pre">id</span></code> attribute in the <code class="docutils literal notranslate"><span class="pre">Customer</span></code> object.  For example, <code class="docutils literal notranslate"><span class="pre">Person(&quot;Barb&quot;,</span> <span class="pre">&quot;Ericson&quot;).last_name</span> <span class="pre">==</span> <span class="pre">&quot;Ericson&quot;</span></code> and <code class="docutils literal notranslate"><span class="pre">Customer(&quot;Barb&quot;,</span> <span class="pre">&quot;Ericson&quot;,</span> <span class="pre">&quot;a1&quot;).id</span> <span class="pre">==</span> <span class="pre">``&quot;a1&quot;</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="1"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class Person:
---
    def __init__(self, first_name, last_name):
---
    def init(first_name, last_name): #paired
---
        self.first_name = first_name
        self.last_name = last_name
---
class Customer(Person):
---
class Customer extends Person: #paired
---
    def __init__(self, first_name, last_name, id):
---
    def __init__(first_name, last_name, id): #paired
---
        super().__init__(first_name, last_name)
---
        super.__init__(first_name, last_name) #paired
---
        self.id = id
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=inheritance_mixed1_ac_v3 data-question_label="2">
<div id=inheritance_mixed1_ac_v3_question class="ac_question col-md-12">
<p>Write a class <code class="docutils literal notranslate"><span class="pre">Person</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that takes <code class="docutils literal notranslate"><span class="pre">first_name</span></code> and <code class="docutils literal notranslate"><span class="pre">last_name</span></code> and inititalize those attributes in the current object. Next, create a class called <code class="docutils literal notranslate"><span class="pre">Customer</span></code> that inherits from <code class="docutils literal notranslate"><span class="pre">Person</span></code> with a constructor that takes <code class="docutils literal notranslate"><span class="pre">first_name</span></code>, <code class="docutils literal notranslate"><span class="pre">last_name</span></code>, and <code class="docutils literal notranslate"><span class="pre">id</span></code>.  Call the constructor in <code class="docutils literal notranslate"><span class="pre">Person</span></code> to initialize <code class="docutils literal notranslate"><span class="pre">first_name</span></code> and <code class="docutils literal notranslate"><span class="pre">last_name</span></code> and then also set the <code class="docutils literal notranslate"><span class="pre">id</span></code> attribute in the <code class="docutils literal notranslate"><span class="pre">Customer</span></code> object.  For example, <code class="docutils literal notranslate"><span class="pre">Person(&quot;Barb&quot;,</span> <span class="pre">&quot;Ericson&quot;).last_name</span> <span class="pre">==</span> <span class="pre">&quot;Ericson&quot;</span></code> and <code class="docutils literal notranslate"><span class="pre">Customer(&quot;Barb&quot;,</span> <span class="pre">&quot;Ericson&quot;,</span> <span class="pre">&quot;a1&quot;).id</span> <span class="pre">==</span> <span class="pre">``&quot;a1&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="inheritance_mixed1_ac_v3_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        p = Person("Barb", "Ericson")
        c = Customer("Barb", "Ericson", "a1")
        self.assertEqual(p.first_name, "Barb", 'Person("Barb", "Ericson").first_name')
        self.assertEqual(p.last_name, "Ericson", 'Person("Barb", "Ericson").last_name')
        self.assertEqual(c.first_name, "Barb", 'Customer("Barb", "Ericson", "a1").first_name')
        self.assertEqual(c.last_name, "Ericson", 'Customer("Barb", "Ericson", "a1").last_name')
        self.assertEqual(c.id, "a1", 'Customer("Barb", "Ericson", "a1").id')


myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="inheritance_mixed2_v3" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class called <code class="docutils literal notranslate"><span class="pre">Person</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that takes <code class="docutils literal notranslate"><span class="pre">name</span></code> and <code class="docutils literal notranslate"><span class="pre">age</span></code>. Next, create a class called <code class="docutils literal notranslate"><span class="pre">Student</span></code> that inherits from <code class="docutils literal notranslate"><span class="pre">Person</span></code>.
The constructor for <code class="docutils literal notranslate"><span class="pre">Student</span></code> takes <code class="docutils literal notranslate"><span class="pre">name</span></code>, <code class="docutils literal notranslate"><span class="pre">age</span></code>, and <code class="docutils literal notranslate"><span class="pre">gpa</span></code>. Have the constructor for the <code class="docutils literal notranslate"><span class="pre">Student</span></code> class call the constructor for the <code class="docutils literal notranslate"><span class="pre">Person</span></code> class to initialize the <code class="docutils literal notranslate"><span class="pre">name</span></code> and <code class="docutils literal notranslate"><span class="pre">age</span></code> and then set the value for the <code class="docutils literal notranslate"><span class="pre">gpa</span></code>.  For example, <code class="docutils literal notranslate"><span class="pre">Person('Susan',</span> <span class="pre">13).name</span></code> would be <code class="docutils literal notranslate"><span class="pre">&quot;Susan&quot;</span></code>
and <code class="docutils literal notranslate"><span class="pre">Student('Sally',</span> <span class="pre">16,</span> <span class="pre">3.7).gpa</span></code> would be <code class="docutils literal notranslate"><span class="pre">3.7</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="3"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class Person:
---
    def __init__(self, name, age):
---
    def init(self, name, age): #paired
---
        self.name = name
        self.age = age
---
class Student(Person):
---
class Student(): #paired
---
    def __init__(self, name, age, gpa):
---
    def __init__(name, age, gpa): #paired
---
        super().__init__(name, age)
---
        super.__init__(name,age) #paired
---
        self.gpa = gpa
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=inheritance_mixed2_ac_v3 data-question_label="4">
<div id=inheritance_mixed2_ac_v3_question class="ac_question col-md-12">
<p>Write a class called <code class="docutils literal notranslate"><span class="pre">Person</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that takes <code class="docutils literal notranslate"><span class="pre">name</span></code> and <code class="docutils literal notranslate"><span class="pre">age</span></code>. Next, create a class called <code class="docutils literal notranslate"><span class="pre">Student</span></code> that inherits from <code class="docutils literal notranslate"><span class="pre">Person</span></code>.
The constructor for <code class="docutils literal notranslate"><span class="pre">Student</span></code> takes <code class="docutils literal notranslate"><span class="pre">name</span></code>, <code class="docutils literal notranslate"><span class="pre">age</span></code>, and <code class="docutils literal notranslate"><span class="pre">gpa</span></code>. Have the constructor for the <code class="docutils literal notranslate"><span class="pre">Student</span></code> class call the constructor for the <code class="docutils literal notranslate"><span class="pre">Person</span></code> class to initialize the <code class="docutils literal notranslate"><span class="pre">name</span></code> and <code class="docutils literal notranslate"><span class="pre">age</span></code> and then set the value for the <code class="docutils literal notranslate"><span class="pre">gpa</span></code>.  For example, <code class="docutils literal notranslate"><span class="pre">Person('Susan',</span> <span class="pre">13).name</span></code> would be <code class="docutils literal notranslate"><span class="pre">&quot;Susan&quot;</span></code>
and <code class="docutils literal notranslate"><span class="pre">Student('Sally',</span> <span class="pre">16,</span> <span class="pre">3.7).gpa</span></code> would be <code class="docutils literal notranslate"><span class="pre">3.7</span></code>.</p>

</div>
<textarea data-lang="python" id="inheritance_mixed2_ac_v3_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        p = Person('Susan', 13)
        s = Student('Sally', 16, 3.7)
        self.assertEqual(p.name, 'Susan', "Person('Susan', 13).name")
        self.assertEqual(p.age, 13, "Person('Susan', 13).age")
        self.assertEqual(s.name, 'Sally', "Student('Sally', 16, 3.7).name")
        self.assertEqual(s.age, 16, "Student('Sally', 16, 3.7).age")
        self.assertEqual(s.gpa, 3.7, "Student('Sally', 16, 3.7).gpa")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="inheritance_mixed3_v3" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class called <code class="docutils literal notranslate"><span class="pre">Building</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inititalizes the attributes <code class="docutils literal notranslate"><span class="pre">floors</span></code> and <code class="docutils literal notranslate"><span class="pre">sqft</span></code>. Next, create a class called <code class="docutils literal notranslate"><span class="pre">House</span></code> that inherits from <code class="docutils literal notranslate"><span class="pre">Building</span></code>. The constructor for <code class="docutils literal notranslate"><span class="pre">House</span></code> should take <code class="docutils literal notranslate"><span class="pre">floors</span></code>, <code class="docutils literal notranslate"><span class="pre">sqft</span></code>, and <code class="docutils literal notranslate"><span class="pre">bedrooms</span></code> and call the construtor in <code class="docutils literal notranslate"><span class="pre">Building</span></code> to initialze the <code class="docutils literal notranslate"><span class="pre">floors</span></code> and <code class="docutils literal notranslate"><span class="pre">sqft</span></code> before setting the attribute <code class="docutils literal notranslate"><span class="pre">bedrooms</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">Building(15,</span> <span class="pre">10000).floors</span></code> would be <code class="docutils literal notranslate"><span class="pre">15</span></code>
and <code class="docutils literal notranslate"><span class="pre">House(3,</span> <span class="pre">2000,</span> <span class="pre">3).bedrooms</span></code> would be <code class="docutils literal notranslate"><span class="pre">3</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="5"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class Building:
---
    def __init__(self, floors, sqft):
---
    def __init__(floors, sqft): #paired
---
        self.floors = floors
        self.sqft = sqft
---
class House(Building):
---
class House: #paired
---
    def __init__(self, floors, sqft, bedrooms):
---
    def __init__(floors, sqft, bedrooms): #paired
---
        super().__init__(floors, sqft)
---
        super.__init__(floors, sqft) #paired
---
        self.bedrooms = bedrooms
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=inheritance_mixed3_ac_v3 data-question_label="6">
<div id=inheritance_mixed3_ac_v3_question class="ac_question col-md-12">
<p>Write a class called <code class="docutils literal notranslate"><span class="pre">Building</span></code> with a <code class="docutils literal notranslate"><span class="pre">constructor</span></code> that inititalizes the attributes <code class="docutils literal notranslate"><span class="pre">floors</span></code> and <code class="docutils literal notranslate"><span class="pre">sqft</span></code>. Next, create a class called <code class="docutils literal notranslate"><span class="pre">House</span></code> that inherits from <code class="docutils literal notranslate"><span class="pre">Building</span></code>. The constructor for <code class="docutils literal notranslate"><span class="pre">House</span></code> should take <code class="docutils literal notranslate"><span class="pre">floors</span></code>, <code class="docutils literal notranslate"><span class="pre">sqft</span></code>, and <code class="docutils literal notranslate"><span class="pre">bedrooms</span></code> and call the construtor in <code class="docutils literal notranslate"><span class="pre">Building</span></code> to initialze the <code class="docutils literal notranslate"><span class="pre">floors</span></code> and <code class="docutils literal notranslate"><span class="pre">sqft</span></code> before setting the attribute <code class="docutils literal notranslate"><span class="pre">bedrooms</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">Building(15,</span> <span class="pre">10000).floors</span></code> would be <code class="docutils literal notranslate"><span class="pre">15</span></code>
and <code class="docutils literal notranslate"><span class="pre">House(3,</span> <span class="pre">2000,</span> <span class="pre">3).bedrooms</span></code> would be <code class="docutils literal notranslate"><span class="pre">3</span></code>.</p>

</div>
<textarea data-lang="python" id="inheritance_mixed3_ac_v3_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        b = Building(15, 10000)
        h = House(3, 2000, 4)

        self.assertEqual(b.floors, 15, "Building(15, 10000).floors")
        self.assertEqual(b.sqft, 10000, "Building(15, 10000).sqft")
        self.assertEqual(h.floors, 3, "House(3, 2000, 4).floors")
        self.assertEqual(h.sqft, 2000, "House(3, 2000, 4).sqft")
        self.assertEqual(h.bedrooms, 4, "House(3, 2000, 4).bedrooms")


myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="inheritance_mixed4_v3" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class <code class="docutils literal notranslate"><span class="pre">Account</span></code> with an attribute of <code class="docutils literal notranslate"><span class="pre">balance</span></code> and methods of <code class="docutils literal notranslate"><span class="pre">deposit</span></code> and <code class="docutils literal notranslate"><span class="pre">withdraw</span></code> that both take an <code class="docutils literal notranslate"><span class="pre">amount</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">Account(200).withdraw(50)</span></code> should set the balance to <code class="docutils literal notranslate"><span class="pre">150</span></code> and <code class="docutils literal notranslate"><span class="pre">Account(200).deposit(100)</span></code> should set the balance to <code class="docutils literal notranslate"><span class="pre">300</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="7"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class Account:
---
    def __init__(self, balance):
---
        self.balance = balance
---
    def deposit(self, amount):
---
    def deposit(amount): #paired
---
        self.balance += amount
---
        balance += amount #paired
---
    def withdraw(self, amount):
---
    def withdraw(amount): #paired
---
        self.balance -= amount
---
        balance -= amount #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=inheritance_mixed4_ac_v3 data-question_label="8">
<div id=inheritance_mixed4_ac_v3_question class="ac_question col-md-12">
<p>Write a class <code class="docutils literal notranslate"><span class="pre">Account</span></code> with an attribute of <code class="docutils literal notranslate"><span class="pre">balance</span></code> and methods of <code class="docutils literal notranslate"><span class="pre">deposit</span></code> and <code class="docutils literal notranslate"><span class="pre">withdraw</span></code> that both take an <code class="docutils literal notranslate"><span class="pre">amount</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">Account(200).withdraw(50)</span></code> should set the balance to <code class="docutils literal notranslate"><span class="pre">150</span></code> and <code class="docutils literal notranslate"><span class="pre">Account(200).deposit(100)</span></code> should set the balance to <code class="docutils literal notranslate"><span class="pre">300</span></code>.</p>

</div>
<textarea data-lang="python" id="inheritance_mixed4_ac_v3_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        a = Account(200)
        b = Account(200)
        c = Account(50)
        d = Account(50)
        e = Account(-50)
        a.withdraw(50)
        self.assertEqual(a.balance, 150, 'Account(200).withdraw(50)')
        b.deposit(100)
        self.assertEqual(b.balance, 300, 'Account(200).deposit(100)')
        c.withdraw(75)
        self.assertEqual(c.balance, -25, 'Account(50).withdraw(75)')
        d.deposit(50)
        self.assertEqual(d.balance, 100, 'Account(50).deposit(50)')
        e.deposit(50)
        self.assertEqual(e.balance, 0, 'Account(-50).deposit(50)')

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="inheritance_mixed5_v3" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class <code class="docutils literal notranslate"><span class="pre">Animal</span></code> with an attribute of <code class="docutils literal notranslate"><span class="pre">name</span></code> and a method <code class="docutils literal notranslate"><span class="pre">make_noise</span></code> that returns <code class="docutils literal notranslate"><span class="pre">&quot;Noise&quot;</span></code>.  Then create a <code class="docutils literal notranslate"><span class="pre">Dog</span></code> class that inherits from <code class="docutils literal notranslate"><span class="pre">Animal</span></code>.  Have the constructor in <code class="docutils literal notranslate"><span class="pre">Dog</span></code> call the constructor in <code class="docutils literal notranslate"><span class="pre">Animal</span></code> to initialze the <code class="docutils literal notranslate"><span class="pre">name</span></code>. In <code class="docutils literal notranslate"><span class="pre">Dog</span></code> also override the <code class="docutils literal notranslate"><span class="pre">make_noise</span></code> method inherited from <code class="docutils literal notranslate"><span class="pre">Animal</span></code> to return <code class="docutils literal notranslate"><span class="pre">&quot;Bark&quot;</span></code>.   For example, <code class="docutils literal notranslate"><span class="pre">Animal(&quot;Diana&quot;).make_noise()</span></code> returns <code class="docutils literal notranslate"><span class="pre">&quot;Noise&quot;</span></code> and <code class="docutils literal notranslate"><span class="pre">Dog(&quot;Percy&quot;).make_noise()</span></code> returns <code class="docutils literal notranslate"><span class="pre">&quot;Bark&quot;</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="9"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class Animal:
---
def class Animal: #paired
---
    def __init__(self, name):
---
        self.name = name
---
    def make_noise(self):
---
    def make_noise(): #paired
---
        return "Noise"
---
class Dog(Animal):
---
class Dog: #paired
---
    def __init__(self, name):
---
        super().__init__(name)
---
    def make_noise(self):
---
    def make_noise(): #paired
---
        return "Bark"
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=inheritance_mixed5_ac_v3 data-question_label="10">
<div id=inheritance_mixed5_ac_v3_question class="ac_question col-md-12">
<p>Write a class <code class="docutils literal notranslate"><span class="pre">Animal</span></code> with an attribute of <code class="docutils literal notranslate"><span class="pre">name</span></code> and a method <code class="docutils literal notranslate"><span class="pre">make_noise</span></code> that returns <code class="docutils literal notranslate"><span class="pre">&quot;Noise&quot;</span></code>.  Then create a <code class="docutils literal notranslate"><span class="pre">Dog</span></code> class that inherits from <code class="docutils literal notranslate"><span class="pre">Animal</span></code>.  Have the constructor in <code class="docutils literal notranslate"><span class="pre">Dog</span></code> call the constructor in <code class="docutils literal notranslate"><span class="pre">Animal</span></code> to initialze the <code class="docutils literal notranslate"><span class="pre">name</span></code>. In <code class="docutils literal notranslate"><span class="pre">Dog</span></code> also override the <code class="docutils literal notranslate"><span class="pre">make_noise</span></code> method inherited from <code class="docutils literal notranslate"><span class="pre">Animal</span></code> to return <code class="docutils literal notranslate"><span class="pre">&quot;Bark&quot;</span></code>.   For example, <code class="docutils literal notranslate"><span class="pre">Animal(&quot;Diana&quot;).make_noise()</span></code> returns <code class="docutils literal notranslate"><span class="pre">&quot;Noise&quot;</span></code> and <code class="docutils literal notranslate"><span class="pre">Dog(&quot;Percy&quot;).make_noise()</span></code> returns <code class="docutils literal notranslate"><span class="pre">&quot;Bark&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="inheritance_mixed5_ac_v3_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        a = Animal('Diana')
        d = Dog('Percy')
        self.assertEqual(a.make_noise(),"Noise", 'Animal("Diana").make_noise()')
        self.assertEqual(a.name,"Diana", 'Animal("Diana").name')
        self.assertEqual(d.make_noise(),"Bark", 'Dog("Percy").make_noise()')
        self.assertEqual(d.name,"Percy", 'Dog("Percy").name')



myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="inheritance_mixed6_v3" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class <code class="docutils literal notranslate"><span class="pre">Animal</span></code> with an attribute of <code class="docutils literal notranslate"><span class="pre">name</span></code> and a method <code class="docutils literal notranslate"><span class="pre">make_noise</span></code> that returns <code class="docutils literal notranslate"><span class="pre">&quot;Noise&quot;</span></code>.  Then create a <code class="docutils literal notranslate"><span class="pre">Cat</span></code> class that inherits from <code class="docutils literal notranslate"><span class="pre">Animal</span></code>.  Have the constructor in <code class="docutils literal notranslate"><span class="pre">Cat</span></code> call the constructor in <code class="docutils literal notranslate"><span class="pre">Animal</span></code> to initialze the <code class="docutils literal notranslate"><span class="pre">name</span></code>. In <code class="docutils literal notranslate"><span class="pre">Cat</span></code> also override the <code class="docutils literal notranslate"><span class="pre">make_noise</span></code> method inherited from <code class="docutils literal notranslate"><span class="pre">Animal</span></code> to return <code class="docutils literal notranslate"><span class="pre">&quot;Meow&quot;</span></code>.   For example, <code class="docutils literal notranslate"><span class="pre">Animal(&quot;Spike&quot;).make_noise()</span></code> returns <code class="docutils literal notranslate"><span class="pre">&quot;Noise&quot;</span></code> and <code class="docutils literal notranslate"><span class="pre">Cat(&quot;Diana&quot;).make_noise()</span></code> returns <code class="docutils literal notranslate"><span class="pre">&quot;Meow&quot;</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="11"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class Animal:
---
def class Animal: #paired
---
    def __init__(self, name):
---
        self.name = name
---
    def make_noise(self):
---
    def make_noise(): #paired
---
        return "Noise"
---
class Cat(Animal):
---
class Cat: #paired
---
    def __init__(self, name):
---
        super().__init__(name)
---
    def make_noise(self):
---
    def make_noise(): #paired
---
        return "Meow"
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=inheritance_mixed6_ac_v3 data-question_label="12">
<div id=inheritance_mixed6_ac_v3_question class="ac_question col-md-12">
<p>Write a class <code class="docutils literal notranslate"><span class="pre">Animal</span></code> with an attribute of <code class="docutils literal notranslate"><span class="pre">name</span></code> and a method <code class="docutils literal notranslate"><span class="pre">make_noise</span></code> that returns <code class="docutils literal notranslate"><span class="pre">&quot;Noise&quot;</span></code>.  Then create a <code class="docutils literal notranslate"><span class="pre">Cat</span></code> class that inherits from <code class="docutils literal notranslate"><span class="pre">Animal</span></code>.  Have the constructor in <code class="docutils literal notranslate"><span class="pre">Cat</span></code> call the constructor in <code class="docutils literal notranslate"><span class="pre">Animal</span></code> to initialze the <code class="docutils literal notranslate"><span class="pre">name</span></code>. In <code class="docutils literal notranslate"><span class="pre">Cat</span></code> also override the <code class="docutils literal notranslate"><span class="pre">make_noise</span></code> method inherited from <code class="docutils literal notranslate"><span class="pre">Animal</span></code> to return <code class="docutils literal notranslate"><span class="pre">&quot;Meow&quot;</span></code>.   For example, <code class="docutils literal notranslate"><span class="pre">Animal(&quot;Spike&quot;).make_noise()</span></code> returns <code class="docutils literal notranslate"><span class="pre">&quot;Noise&quot;</span></code> and <code class="docutils literal notranslate"><span class="pre">Cat(&quot;Diana&quot;).make_noise()</span></code> returns <code class="docutils literal notranslate"><span class="pre">&quot;Meow&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="inheritance_mixed6_ac_v3_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        a = Animal('Spike')
        c = Cat('Diana')
        self.assertEqual(a.make_noise(),"Noise", 'Animal("Spike").make_noise()')
        self.assertEqual(a.name,"Spike", 'Animal("Spike").name')
        self.assertEqual(c.make_noise(),"Meow", 'Cat("Diana").make_noise()')
        self.assertEqual(c.name,"Diana", 'Cat("Diana").name')

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="inheritance_mixed7_v3" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Given a class <code class="docutils literal notranslate"><span class="pre">Point</span></code> that has attributes of <code class="docutils literal notranslate"><span class="pre">x</span></code> and <code class="docutils literal notranslate"><span class="pre">y</span></code> and a method <code class="docutils literal notranslate"><span class="pre">distance_to(self,</span> <span class="pre">other)</span></code> that returns the distance between the current point (self) and the other point, create a <code class="docutils literal notranslate"><span class="pre">MyCircle</span></code> class that has a constructor that takes two attibutes <code class="docutils literal notranslate"><span class="pre">p1</span></code> and <code class="docutils literal notranslate"><span class="pre">p2</span></code> both objects of the class <code class="docutils literal notranslate"><span class="pre">Point</span></code>.  Also create a <code class="docutils literal notranslate"><span class="pre">radius</span></code> method that returns the radius of the circle (half the distance between the two points that define circle).</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="13"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class MyCircle:
---
class MyCircle(Point): #paired
---
    def __init__(self, p1, p2):
---
    def __init__(p1, p2): #paired
---
        self.p1 = p1
        self.p2 = p2
---
    def radius(self):
---
    def radius(): #paired
---
        return self.p1.distance_to(self.p2) / 2
---
        return p1.distance_to(p2) / 2 #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=inheritance_mixed7_ac_v3 data-question_label="14">
<div id=inheritance_mixed7_ac_v3_question class="ac_question col-md-12">
<p>Given a class <code class="docutils literal notranslate"><span class="pre">Point</span></code> defined below that has attributes of <code class="docutils literal notranslate"><span class="pre">x</span></code> and <code class="docutils literal notranslate"><span class="pre">y</span></code> and a method <code class="docutils literal notranslate"><span class="pre">distance_to(self,</span> <span class="pre">other)</span></code> that returns the distance between the current point (self) and the other point, write a <code class="docutils literal notranslate"><span class="pre">MyCircle</span></code> class that that has a constructor that takes two attibutes <code class="docutils literal notranslate"><span class="pre">p1</span></code> and <code class="docutils literal notranslate"><span class="pre">p2</span></code> both objects of the class <code class="docutils literal notranslate"><span class="pre">Point</span></code>.  Also create a <code class="docutils literal notranslate"><span class="pre">radius</span></code> method that returns the radius of the circle (half the distance between the two points that define circle).</p>

</div>
<textarea data-lang="python" id="inheritance_mixed7_ac_v3_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Point:
    def __init__(self, x, y):
        self.x = x
        self.y = y
    def distance_to(self, other):
        return (((self.x - other.x) ** 2) + ((self.y - other.y) ** 2)) ** 0.5

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        p1 = Point(5,5)
        p2 = Point(5, 10)
        p3 = Point(5, 15)
        c = MyCircle(p1, p2)
        c2 = MyCircle(p1, p3)

        self.assertEqual(c.p1, p1, "test of MyCircle p1 for first circle")
        self.assertEqual(c.p2, p2, "test of MyCircle p2 for first circle")
        self.assertEqual(c.radius(), p1.distance_to(p2) /  2, "test of MyCircle radius for first circle")
        self.assertEqual(c2.p2, p3, "test of MyCircle p2 for second circle")
        self.assertEqual(c2.radius(), p1.distance_to(p3) /  2, "test of MyCircle radius for second circle")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="inheritance_mixed8_v3" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Given a class <code class="docutils literal notranslate"><span class="pre">Item</span></code> with an attributes of <code class="docutils literal notranslate"><span class="pre">name</span></code> and <code class="docutils literal notranslate"><span class="pre">price</span></code>, create an <code class="docutils literal notranslate"><span class="pre">Order</span></code> class that has an attribute <code class="docutils literal notranslate"><span class="pre">item_list</span></code> that is intitialized to the empty list in the constructor.  Then add an <code class="docutils literal notranslate"><span class="pre">add_item</span></code> method that takes an <code class="docutils literal notranslate"><span class="pre">item</span></code> and appends it to the <code class="docutils literal notranslate"><span class="pre">item_list</span></code> attribute.   Then create a <code class="docutils literal notranslate"><span class="pre">get_total</span></code> method that returns the total price for all the items in <code class="docutils literal notranslate"><span class="pre">item_list</span></code> attribute.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="15"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class Order:
---
class Order(Item): #paired
---
    def __init__(self):
---
        self.item_list = []
---
        self.item_list = () #paired
---
    def add_item(self, item):
---
        self.item_list.appand(item)
---
        item_list.appand(item) #paired
---
    def get_total(self):
---
        total = 0
---
        for item in self.item_list:
---
        for item in item_list: #paired
---
            total += item.price
---
        return total
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=inheritance_mixed8_ac_v3 data-question_label="16">
<div id=inheritance_mixed8_ac_v3_question class="ac_question col-md-12">
<p>Given a class <code class="docutils literal notranslate"><span class="pre">Item</span></code> define below with attributes of <code class="docutils literal notranslate"><span class="pre">name</span></code> and <code class="docutils literal notranslate"><span class="pre">price</span></code>, write an <code class="docutils literal notranslate"><span class="pre">Order</span></code> class that has an attribute <code class="docutils literal notranslate"><span class="pre">item_list</span></code> that is intitialized to the empty list in the constructor.  Then add an <code class="docutils literal notranslate"><span class="pre">add_item</span></code> method that takes an <code class="docutils literal notranslate"><span class="pre">item</span></code> and appends it to the <code class="docutils literal notranslate"><span class="pre">item_list</span></code> attribute.   Then create a <code class="docutils literal notranslate"><span class="pre">get_total</span></code> method that returns the total price for all the items in <code class="docutils literal notranslate"><span class="pre">item_list</span></code> attribute.</p>

</div>
<textarea data-lang="python" id="inheritance_mixed8_ac_v3_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Item:
    def __init__(self, name, price):
        self.name = name
        self.price = price


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        i1 = Item('hamburger', 6.99)
        i2 = Item('fries', 2.99)
        i3 = Item('drink', 1.99)
        i4 = Item("dessert", 1.99)
        o1 = Order()
        o1.add_item(i1)
        o1.add_item(i2)
        o1.add_item(i3)
        self.assertEqual(o1.get_total(), 11.97, 'get_total on order 1')
        o1.add_item(i4)
        self.assertEqual(o1.get_total(), 13.96, 'get total on order 2')
        o2 = Order()
        o2.add_item(i1)
        o2.add_item(i4)
        self.assertEqual(o2.get_total(), 8.98, 'get_total on order 3')
        o2.add_item(i2)
        self.assertEqual(o2.get_total(), 11.97, 'get_total on order 4')
        o2.add_item(i3)
        self.assertEqual(o2.get_total(), 13.96, 'get_total on order 5')



myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="inheritance_mixed9_v3" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Given a class <code class="docutils literal notranslate"><span class="pre">Account</span></code> with an attribute <code class="docutils literal notranslate"><span class="pre">balance</span></code> and methods of <code class="docutils literal notranslate"><span class="pre">deposit</span></code> and <code class="docutils literal notranslate"><span class="pre">withdraw</span></code>, create a <code class="docutils literal notranslate"><span class="pre">SavingsAccount</span></code> class that inherits from <code class="docutils literal notranslate"><span class="pre">Account</span></code>.  Create a constructor that takes <code class="docutils literal notranslate"><span class="pre">balance</span></code> and <code class="docutils literal notranslate"><span class="pre">interest</span></code>.  In the <code class="docutils literal notranslate"><span class="pre">SavingsAccount</span></code> constructor call the <code class="docutils literal notranslate"><span class="pre">Account</span></code> constructor to intitialze the <code class="docutils literal notranslate"><span class="pre">balance</span></code> and then initialize the <code class="docutils literal notranslate"><span class="pre">interest</span></code>.   Then create a <code class="docutils literal notranslate"><span class="pre">pay_interest</span></code> method that adds the interest to the balance (divide the interest rate by 100 and multiply it by the balance to determine the amount to add).</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="17"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class SavingsAccount(Account):
---
class SavingsAccount: #paired
---
    def __init__(self, balance, interest):
---
    def __init__(balance, interest): #paired
---
        super().__init__(balance)
---
        super.__init__(balance) #paired
---
        self.interest = interest
---
    def pay_interest(self):
---
    def pay_interest(): #paired
---
        amount = self.interest / 100
---
        self.balance += amount * self.balance
---
        self.balance += self.amount * self.balance #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=inheritance_mixed9_ac_v3 data-question_label="18">
<div id=inheritance_mixed9_ac_v3_question class="ac_question col-md-12">
<p>Given a class <code class="docutils literal notranslate"><span class="pre">Account</span></code> with an attribute <code class="docutils literal notranslate"><span class="pre">balance</span></code> and methods of <code class="docutils literal notranslate"><span class="pre">deposit</span></code> and <code class="docutils literal notranslate"><span class="pre">withdraw</span></code>, create a <code class="docutils literal notranslate"><span class="pre">SavingsAccount</span></code> class hat inherits from <code class="docutils literal notranslate"><span class="pre">Account</span></code>.  Write a constructor that takes <code class="docutils literal notranslate"><span class="pre">balance</span></code> and <code class="docutils literal notranslate"><span class="pre">interest</span></code>.  In the <code class="docutils literal notranslate"><span class="pre">SavingsAccount</span></code> constructor call the <code class="docutils literal notranslate"><span class="pre">Account</span></code> constructor to intitialze the <code class="docutils literal notranslate"><span class="pre">balance</span></code> and then initialize the <code class="docutils literal notranslate"><span class="pre">interest</span></code>.   Then create a <code class="docutils literal notranslate"><span class="pre">pay_interest</span></code> method that adds the interest to the balance (divide the interest rate by 100 and multiply it by the balance to determine the amount to add).</p>

</div>
<textarea data-lang="python" id="inheritance_mixed9_ac_v3_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Account:
    def __init__(self, balance):
        self.balance = balance
    def deposit(self, amount):
        self.balance += amount
    def withdraw(self, amount):
        self.balance -= amount



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        s = SavingsAccount(200, 1.5)
        s.withdraw(50)
        self.assertEqual(s.balance, 150, "test of withdraw")
        s.deposit(75)
        self.assertEqual(s.balance, 225, "test of deposit")
        s.pay_interest()
        self.assertAlmostEqual(s.balance, 228.375, 3, "test of pay_interest")




myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="inheritance_mixed10_v3" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Given a class <code class="docutils literal notranslate"><span class="pre">Treasure</span></code> that has attributes of <code class="docutils literal notranslate"><span class="pre">name</span></code> and <code class="docutils literal notranslate"><span class="pre">points</span></code>, create a class <code class="docutils literal notranslate"><span class="pre">Room</span></code> that has attributes of <code class="docutils literal notranslate"><span class="pre">name</span></code> and <code class="docutils literal notranslate"><span class="pre">treasures</span></code>.  Initialize <code class="docutils literal notranslate"><span class="pre">treasures</span></code> to the empty list in the constructor.  Create a <code class="docutils literal notranslate"><span class="pre">add_treasure</span></code> method in <code class="docutils literal notranslate"><span class="pre">Room</span></code> that adds a passed <code class="docutils literal notranslate"><span class="pre">Treasure</span></code> object to <code class="docutils literal notranslate"><span class="pre">treasures</span></code>.  Next create a <code class="docutils literal notranslate"><span class="pre">get_points</span></code> method in <code class="docutils literal notranslate"><span class="pre">Room</span></code> that returns the total of all of the points for the <code class="docutils literal notranslate"><span class="pre">Treasure</span></code> in <code class="docutils literal notranslate"><span class="pre">treasures</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="19"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class Room:
---
class Room (Treasure): #paired
---
    def __init__(self, name):
---
        self.name = name
        self.treasures = []
---
    def add_treasure(self, treasure):
---
        self.treasures.append(treasure)
---
        treasures.append(treasure) #paired
---
    def get_points(self):
---
        total = 0
---
        for treasure in self.treasures:
---
        for treasure in treasures: #paired
---
            total += treasure.points
---
        return total
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=inheritance_mixed10_ac_v3 data-question_label="20">
<div id=inheritance_mixed10_ac_v3_question class="ac_question col-md-12">
<p>Given a class <code class="docutils literal notranslate"><span class="pre">Treasure</span></code> below that has attributes of <code class="docutils literal notranslate"><span class="pre">name</span></code> and <code class="docutils literal notranslate"><span class="pre">points</span></code>, create a class <code class="docutils literal notranslate"><span class="pre">Room</span></code> that has attributes of <code class="docutils literal notranslate"><span class="pre">name</span></code> and <code class="docutils literal notranslate"><span class="pre">treasures</span></code>.  Initialize <code class="docutils literal notranslate"><span class="pre">treasures</span></code> to the empty list in the constructor.  Create a <code class="docutils literal notranslate"><span class="pre">add_treasure</span></code> method in <code class="docutils literal notranslate"><span class="pre">Room</span></code> that adds a passed <code class="docutils literal notranslate"><span class="pre">Treasure</span></code> object to <code class="docutils literal notranslate"><span class="pre">treasures</span></code>.  Next create a <code class="docutils literal notranslate"><span class="pre">get_points</span></code> method in <code class="docutils literal notranslate"><span class="pre">Room</span></code> that returns the total of all of the points for the <code class="docutils literal notranslate"><span class="pre">Treasure</span></code> in <code class="docutils literal notranslate"><span class="pre">treasures</span></code>.</p>

</div>
<textarea data-lang="python" id="inheritance_mixed10_ac_v3_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Treasure:
    def __init__(self, name, points):
        self.name = name
        self.points = points



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        t1 = Treasure("Ring", 100)
        t2 = Treasure("Lamp", 300)
        t3 = Treasure("Necklace", 200)
        t4 = Treasure("Shoes", 150)
        r1 = Room("Kitchen")
        r1.add_treasure(t1)
        r1.add_treasure(t2)
        self.assertEqual(r1.get_points(), 400, "test of get_points on room 1")
        r2 = Room("Billards")
        r2.add_treasure(t3)
        r2.add_treasure(t4)
        self.assertEqual(r2.get_points(), 350, "test of get_points on room 2")
        r3 = Room("Living Room")
        r3.add_treasure(t1)
        self.assertEqual(r3.get_points(), 100, "test of get_points on room 3")
        r4 = Room("Pantry")
        r4.add_treasure(t2)
        self.assertEqual(r4.get_points(), 300, "test of get_points on room 4")




myTests().main()
</textarea>
</div>
</div>
</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>