<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <title>Writing code problems &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=DB047029" />
    <link rel="stylesheet" type="text/css" href="https://cdn.jsdelivr.net/npm/handsontable@7.2.2/dist/handsontable.full.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-3.4.1/css/bootstrap.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/presenter_mode.css?v=9FF1320F" />
    <link rel="stylesheet" type="text/css" href="../_static/jquery-ui-1.10.3.custom.min.css?v=97E0178A" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/user-highlights.css?v=A2326E52" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone-custom-sphinx-bootstrap.css?v=5.0.8" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/language_data.js"></script>
    <script src="../_static/jquery.highlight.js?v=F1C496C8"></script>
    <script src="../_static/sql-wasm.js?v=44FE0709"></script>
    <script src="https://cdn.jsdelivr.net/npm/handsontable@7.2.2/dist/handsontable.full.js"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/pytutor-embed.bundle.js?v=268B5C04"></script>
    <script src="../_static/html4css1.css?v=4D8F3FB4"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/Blob.js?v=AEC2447"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.emitter.bidi.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.emitter.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.fallbacks.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.messagestore.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.parser.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.language.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega@4.0.0-rc.2/build/vega.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega-lite@2.5.0/build/vega-lite.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega-embed@3.14.0/build/vega-embed.js"></script>
    <script src="../_static/runestone.js?v=3EAD674C"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js?v=2B07FE4F"></script>
    <script src="../_static/bootstrap-3.4.1/js/bootstrap.min.js"></script>
    <script src="../_static/jquery-fix.js?v=1226591B"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/jquery.idle-timer.js?v=E942DAF1"></script>
    <script src="../_static/presenter_mode.js?v=2DF16FB4"></script>
    <script src="../_static/theme.js?v=D899FB97"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="Functions" href="../04-functions/toctree.html" />
    <link rel="prev" title="Mixed-up code exercises" href="03-MixedupCode.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'py4e-int';
    eBookConfig.basecourse = 'py4e-int';
    eBookConfig.isLoggedIn = false;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;


</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar" aria-label="navbar toggle" class="visuallyhidden">toggle navbar</span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
          <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          py4e-int
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span aria-label="Social" class="visuallyhidden">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span aria-label="Search" class="visuallyhidden">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span aria-label="User" class="visuallyhidden">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:ACFactory.toggleScratchActivecode()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span aria-label="Scratch Activecode" class="visuallyhidden">Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span aria-label="Help" class="visuallyhidden">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=03-WriteCode'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - Mixed-up code exercises" data-toggle="tooltip">
    <a href="03-MixedupCode.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Functions' data-toggle="tooltip" >              
    <a href="../04-functions/toctree.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

      </div> <!-- navbar -->







    </div> <!-- container -->
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
  
  <div class="section" id="writing-code-problems">
<h1>Writing code problems<a class="headerlink" href="#writing-code-problems" title="Permalink to this headline">¶</a></h1>
<ol class="arabic">
<li><div id='cndtnl-exercise-five' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<p>Fix the errors in the code, and change it to use only <strong>one</strong> if statement. The code should print
“The number is 5” when the number is 5, and should print “The number is NOT 5” when it is not.</p>

<div data-childcomponent="cndtnl-wc-fiveq" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=cndtnl-wc-fiveq data-lang="python" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static>
x = 5
if x == 5:
print("The number is 5")
if x != 5:
print("The number is NOT 5")
</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="cndtnl-wc-fivea" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=cndtnl-wc-fivea data-lang="python" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static>
x = 5
if x == 5:
    print("The number is 5")
else:
    print("The number is NOT 5")
</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function cndtnlFiveDisc_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'cndtnlFiveDisc_' + eBookConfig.course;
                    this.page.url = 'https://cndtnlFiveDisc_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'cndtnlFiveDisc_' + eBookConfig.course;
            disqus_url = 'http://cndtnlFiveDisc_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="cndtnlFiveDisc" onclick="cndtnlFiveDisc_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='cndtnlFiveDisc']").attr('data-disqus-identifier', 'cndtnlFiveDisc_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li><div id='cndtnl-exercise-pay' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<p>Complete the pay computation to give the employee 1.5 times the hourly rate for hours worked
above 40 hours, if the regular pay rate is $10 an hour. Then set <code class="docutils literal notranslate"><span class="pre">grossPay</span></code> equal to the
amount an employee would be paid for working 45 hours.</p>

<div data-childcomponent="cndtnl-wc-payq" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=cndtnl-wc-payq data-lang="python" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static>
hours = 45
rate = 10
overtimeRate =  #your code here
grossPay = 0
if hours  #your code here
    grossPay = hours * rate
else:
    overTime =  #your code here
    grossPay =  #your code here
print(grossPay)

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(grossPay, 475.0)

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="cndtnl-wc-paya" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=cndtnl-wc-paya data-lang="python" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static>
#initializing variables
hours = 45
rate = 10
overtimeRate = rate * 1.5
grossPay = 0

#starting if statement to see if hours are within regular pay
if hours <= 40:
    #if within 40 hours, pay will be hours * rate
    grossPay = hours * rate
#else statement if hours are greater than 40
else:
    #creating a variable for overtime hours
    overTime = hours % 40
    #pay will equal the regular rate for 40 hours, plus the overtime rate for the extra hours
    grossPay = (rate * 40) + (overTime * overtimeRate)
print(grossPay)
</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function cndtnlPayDisc_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'cndtnlPayDisc_' + eBookConfig.course;
                    this.page.url = 'https://cndtnlPayDisc_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'cndtnlPayDisc_' + eBookConfig.course;
            disqus_url = 'http://cndtnlPayDisc_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="cndtnlPayDisc" onclick="cndtnlPayDisc_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='cndtnlPayDisc']").attr('data-disqus-identifier', 'cndtnlPayDisc_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li><div id='cndtnl-exercise-pay2' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<p>Rewrite your pay program using <code class="docutils literal notranslate"><span class="pre">try</span></code> and <code class="docutils literal notranslate"><span class="pre">except</span></code> so that your program handles non-numeric
input gracefully by printing a message and exiting the program. The following shows two
executions of the program:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">Enter</span> <span class="n">Hours</span><span class="p">:</span> <span class="mi">20</span>
<span class="n">Enter</span> <span class="n">Rate</span><span class="p">:</span> <span class="n">nine</span>
<span class="n">Error</span><span class="p">,</span> <span class="n">please</span> <span class="n">enter</span> <span class="n">numeric</span> <span class="nb">input</span>
</pre></div>
</div>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">Enter</span> <span class="n">Hours</span><span class="p">:</span> <span class="n">forty</span>
<span class="n">Error</span><span class="p">,</span> <span class="n">please</span> <span class="n">enter</span> <span class="n">numeric</span> <span class="nb">input</span>
</pre></div>
</div>

<div data-childcomponent="cndtnl-wc-pay2q" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=cndtnl-wc-pay2q data-lang="python" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static>


</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function cndtnlPay2Disc_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'cndtnlPay2Disc_' + eBookConfig.course;
                    this.page.url = 'https://cndtnlPay2Disc_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'cndtnlPay2Disc_' + eBookConfig.course;
            disqus_url = 'http://cndtnlPay2Disc_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="cndtnlPay2Disc" onclick="cndtnlPay2Disc_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='cndtnlPay2Disc']").attr('data-disqus-identifier', 'cndtnlPay2Disc_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li><div id='cndtnl-exercise-distance' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<p>Write the code to calculate and print the cost of a 14 mile cab ride. If the distance traveled
is less than or equal to 12 miles the cost is $2.00 a mile, and if the distance traveled is more
than 12 miles the cost is $1.50 a mile.</p>

<div data-childcomponent="cndtnl-wc-distanceq" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=cndtnl-wc-distanceq data-lang="python" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static>


</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="cndtnl-wc-distancea" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=cndtnl-wc-distancea data-lang="python" 
      data-timelimit=25000     data-optional=true
    data-audio=''      
           data-wasm=/_static>
distance = 14
# SET CONDITIONS
if distance <= 12:
    rate = 2.00
if distance > 12:
    rate = 1.50
# CALCULATE TRIP COST
total = distance * rate
print("Total cost of trip: " + str(total))
</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function cndtnlDistanceDisc_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'cndtnlDistanceDisc_' + eBookConfig.course;
                    this.page.url = 'https://cndtnlDistanceDisc_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'cndtnlDistanceDisc_' + eBookConfig.course;
            disqus_url = 'http://cndtnlDistanceDisc_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="cndtnlDistanceDisc" onclick="cndtnlDistanceDisc_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='cndtnlDistanceDisc']").attr('data-disqus-identifier', 'cndtnlDistanceDisc_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li><div id='cndtnl-exercise-score' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<p>Write a program to prompt for a score between 0.0 and 1.0. If the score is out of range,
print an error message. If the score is between 0.0 and 1.0, print a grade using the
following table:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span> <span class="n">Score</span>   <span class="n">Grade</span>
<span class="o">&gt;=</span> <span class="mf">0.9</span>     <span class="n">A</span>
<span class="o">&gt;=</span> <span class="mf">0.8</span>     <span class="n">B</span>
<span class="o">&gt;=</span> <span class="mf">0.7</span>     <span class="n">C</span>
<span class="o">&gt;=</span> <span class="mf">0.6</span>     <span class="n">D</span>
 <span class="o">&lt;</span> <span class="mf">0.6</span>     <span class="n">F</span>
</pre></div>
</div>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">Enter</span> <span class="n">score</span><span class="p">:</span> <span class="mf">0.95</span>
<span class="n">A</span>
</pre></div>
</div>

<div data-childcomponent="cndtnl-wc-scoreq" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=cndtnl-wc-scoreq data-lang="python" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static>


</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function cndtnlScoreDisc_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'cndtnlScoreDisc_' + eBookConfig.course;
                    this.page.url = 'https://cndtnlScoreDisc_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'cndtnlScoreDisc_' + eBookConfig.course;
            disqus_url = 'http://cndtnlScoreDisc_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="cndtnlScoreDisc" onclick="cndtnlScoreDisc_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='cndtnlScoreDisc']").attr('data-disqus-identifier', 'cndtnlScoreDisc_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li><div id='cndtnl-exercise-weight' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<p>Fix the example such that the cost of frozen yogurt is 0 if you pour exactly 1 lb. in your cup.</p>

<div data-childcomponent="cndtnl-wc-weightq" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=cndtnl-wc-weightq data-lang="python" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static>
weight = 0.5
if weight < 1:
    price = 1.45
if weight > 1:
    price = 1.15
total = weight * price
print(weight)
print(price)
print(total)
</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="cndtnl-wc-weighta" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=cndtnl-wc-weighta data-lang="python" 
      data-timelimit=25000     data-optional=true
    data-audio=''      
           data-wasm=/_static>
weight = 0.5
if weight < 1:
    price = 1.45
if weight == 1:
    price = 0
if weight > 1:
    price = 1.15
total = weight * price
print(weight)
print(price)
print(total)
</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function cndtnlWeightDisc_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'cndtnlWeightDisc_' + eBookConfig.course;
                    this.page.url = 'https://cndtnlWeightDisc_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'cndtnlWeightDisc_' + eBookConfig.course;
            disqus_url = 'http://cndtnlWeightDisc_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="cndtnlWeightDisc" onclick="cndtnlWeightDisc_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='cndtnlWeightDisc']").attr('data-disqus-identifier', 'cndtnlWeightDisc_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li><div id='cndtnl-exercise-wallet' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<p>Write a procedure that takes 2 ints, total price, and amount in wallet. Print “You have enough
money” if the difference between the wallet and price is 0 or greater; otherwise,
print “Get more money”</p>

<div data-childcomponent="cndtnl-wc-walletq" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=cndtnl-wc-walletq data-lang="python" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static>


</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function cndtnlWalletDisc_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'cndtnlWalletDisc_' + eBookConfig.course;
                    this.page.url = 'https://cndtnlWalletDisc_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'cndtnlWalletDisc_' + eBookConfig.course;
            disqus_url = 'http://cndtnlWalletDisc_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="cndtnlWalletDisc" onclick="cndtnlWalletDisc_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='cndtnlWalletDisc']").attr('data-disqus-identifier', 'cndtnlWalletDisc_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li><div id='cndtnl-exercise-even' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<p>Write a program that takes an integer and sets <code class="docutils literal notranslate"><span class="pre">isEven</span></code> to <code class="docutils literal notranslate"><span class="pre">True</span></code> if the argument
is an <strong>even number</strong> and <code class="docutils literal notranslate"><span class="pre">False</span></code> if it is <strong>odd</strong>.</p>

<div data-childcomponent="cndtnl-wc-evenq" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=cndtnl-wc-evenq data-lang="python" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static>
def is_even(n):
    # your code here
</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="cndtnl-wc-evena" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=cndtnl-wc-evena data-lang="python" 
      data-timelimit=25000     data-optional=true
    data-audio=''      
           data-wasm=/_static>
isEven = True

if n % 2 == 0:
    isEven = True
else:
    isEven = False
print(isEven)
</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function cndtnlEvenDisc_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'cndtnlEvenDisc_' + eBookConfig.course;
                    this.page.url = 'https://cndtnlEvenDisc_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'cndtnlEvenDisc_' + eBookConfig.course;
            disqus_url = 'http://cndtnlEvenDisc_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="cndtnlEvenDisc" onclick="cndtnlEvenDisc_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='cndtnlEvenDisc']").attr('data-disqus-identifier', 'cndtnlEvenDisc_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li><div id='cndtnl-exercise-hi' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<p>Complete and finish the code so that the code prints “Hi” if x satisfies the condition on line 2 and prints “Hello” otherwise.</p>

<div data-childcomponent="cndtnl-wc-hiq" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=cndtnl-wc-hiq data-lang="python" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static>
x =
if not x != 3:
</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function cndtnlHiDisc_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'cndtnlHiDisc_' + eBookConfig.course;
                    this.page.url = 'https://cndtnlHiDisc_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'cndtnlHiDisc_' + eBookConfig.course;
            disqus_url = 'http://cndtnlHiDisc_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="cndtnlHiDisc" onclick="cndtnlHiDisc_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='cndtnlHiDisc']").attr('data-disqus-identifier', 'cndtnlHiDisc_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li><div id='cndtnl-exercise-leapYear' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<p>3 criteria must be taken into account to identify leap years:</p>
<ul class="simple">
<li><p>The year is evenly divisible by 4;</p></li>
<li><p>If the year can be evenly divided by 100, it is NOT a leap year, unless;</p></li>
<li><p>The year is also evenly divisible by 400. Then it is a leap year.</p></li>
</ul>
<p>Write a program that takes a year as a parameter and sets <code class="docutils literal notranslate"><span class="pre">leapYear</span></code> equal to <code class="docutils literal notranslate"><span class="pre">True</span></code> if the year is a
leap year, <code class="docutils literal notranslate"><span class="pre">False</span></code> otherwise. (use a few different years to test your work)</p>

<div data-childcomponent="cndtnl-wc-leapYearq" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=cndtnl-wc-leapYearq data-lang="python" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static>


</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="cndtnl-wc-leapYeara" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=cndtnl-wc-leapYeara data-lang="python" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static>
year = 1900 #TEST WITH ANY YEAR
# Initializing leapYear to False since most years are not leap years
leapYear = False

# Use modulus to see if the year is divisible by 4
if year % 4 == 0:
    # If divisible by 4, use modulus to see if divisible by 100
    if year % 100 == 0:
        # if divisible by 4 and 100, use modulus to see if divisible by 400
        if year % 400 == 0:
            # if divisible by 4, 100, and 400 it is a leap year
            leapYear = True
    else: # if divisible by 4 and not 100 it is a leap year
        leapYear = True
print(leapYear)
''' Since leapYear started as False, we do not need to reset it to False
    for the instances that the year is not a leap year '''

# ANOTHER OPTION
year = 1900
leapYear = False
if year % 400 == 0:
    leapYear = True
elif year % 4 == 0 and year % 100 != 0:
    leapYear = True
print(leapYear)
</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function cndtnlLeapYearDisc_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'cndtnlLeapYearDisc_' + eBookConfig.course;
                    this.page.url = 'https://cndtnlLeapYearDisc_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'cndtnlLeapYearDisc_' + eBookConfig.course;
            disqus_url = 'http://cndtnlLeapYearDisc_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="cndtnlLeapYearDisc" onclick="cndtnlLeapYearDisc_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='cndtnlLeapYearDisc']").attr('data-disqus-identifier', 'cndtnlLeapYearDisc_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
</ol>
</div>


  <div id=scprogresscontainer>
    You have attempted <span id=scprogresstotal></span> of <span id=scprogressposs></span> activities on this page <div id="subchapterprogress"></div>
  </div>
  
      
<ul>
  <li id="relations-prev" style="opacity:0.4" onmouseover="mouseOverPrev()" onmouseout="mouseOutPrev()"class="navLink" title='Previous Section - Mixed-up code exercises' data-toggle="tooltip">
    <a href="03-MixedupCode.html" aria-label="Mixed-up code exercises">
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">
            <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
        </div>
    </a>
  </li>
</ul>

  
<ul>
  <li id="relations-next" style="opacity:0.4" onmouseover="mouseOverNext()" onmouseout="mouseOutNext()" class="navLink" title='Next Section - Functions' data-toggle="tooltip" >
    <a href="../04-functions/toctree.html" aria-label="Functions">
      <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">
          <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
      </div>
    </a>
  </li>
</ul>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});

</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 CS Learning 4 U group.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>




<script> getSwitch(); </script>


  </body>
</html>