<?xml version="1.0"?>
<section xml:id="strings_write-code-questions">
  <title>Write Code Questions</title>

  <activity xml:id="str-ex-hiA" label="str-ex-hiA">
    <statement>
      <p>Fix line 2 so that it prints <q>Hi</q> instead of <q>hi</q>.</p>
    </statement>
    <program interactive="activecode" language="python">
      <code>
s1 = "hi"
s1.capitalize()
print(s1)

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(s1, "Hi")
myTests().main()
      </code>
    </program>
    <solution>
      <p>Fix line 2 so that it prints <q>Hi</q> instead of <q>hi</q>.</p>
      <program xml:id="str-ex-hiA-sol" interactive="activecode" language="python">
        <code>
s1 = "hi"
s1 = s1.capitalize()
print(s1)
        </code>
      </program>
    </solution>
  </activity>

  <activity label="str-ex-meow_sq"><select grade="first" questions="str-ex-meowacq str-ex-meowansw"/></activity>

  <activity label="str-ex-eggs" xml:id="str-ex-eggs">
    <statement>
      <p>Write code to evaluate the length of the string <q>I like green eggs</q> and print it. It should print <q>The length is 17</q>.</p>
    </statement>
    <program interactive="activecode" language="python">
      <code>

      </code>
    </program>
    <solution>
       <p>Write code to evaluate the length of the string <q>I like green eggs</q> and print it. It should print <q>The length is 17</q>.</p>
       <program xml:id="str-ex-eggsacq-sol" interactive="activecode" language="python">
        <code>
sentence = "I like green eggs"
print("The length is " + str(len(sentence)))
        </code>
       </program>
    </solution>
  </activity>


  <activity label="str-ex-count-sq"><select grade="first" questions="str-ex-countacq str-ex-countansw"/></activity>

  <activity label="str-ex-floatq" xml:id="str-ex-floatq">
    <statement>
      <p>Take the following Python code that stores this string: <c>string = "X-DSPAM-Confidence: 0.8475"</c>. Use <c>find</c> and string slicing to extract the portion of the string after the colon character and then use the <c>float</c> function to convert the extracted string into a floating point number called <c>num</c>.</p>
    </statement>
    <program interactive="activecode" language="python">
      <code>
string = "X-DSPAM-Confidence: 0.8475"

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(num, 0.8475)
myTests().main()
      </code>
    </program>
    <solution>
      <p>Take the following Python code that stores this string: <c>string = "X-DSPAM-Confidence: 0.8475"</c>. Use <c>find</c> and string slicing to extract the portion of the string after the colon character and then use the <c>float</c> function to convert the extracted string into a floating point number called <c>num</c>.</p>
      <program xml:id="str-ex-floatq-sol" interactive="activecode" language="python">
      <code>
string = "X-DSPAM-Confidence: 0.8475"
colonpos = string.find(':')
num = float(string[colonpos+1:])
print(num)
      </code>
    </program>
    </solution>
  </activity>

  <activity label="tr-ex-ducks-sq"><select grade="first" questions="str-ex-ducksacq str-ex-ducksansw"/></activity>

  <activity label="str-ex-digitsq" xml:id="str-ex-digitsq">
    <statement>
      <p>Write a function <c>numDigits</c> that will return the number of digits in an integer <c>n</c>.</p>
    </statement>
    <program interactive="activecode" language="python">
      <code>
def numDigits(n):
    # your code here

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(numDigits(2), 1, "Tested numDigits on input of 2")
        self.assertEqual(numDigits(55), 2, "Tested numDigits on input of 55")
        self.assertEqual(numDigits(1352), 4, "Tested numDigits on input of 1352")
        self.assertEqual(numDigits(444), 3, "Tested numDigits on input of 444")
myTests().main()
      </code>
    </program>
    <solution>
      <p>Write a function <c>numDigits</c> that will return the number of digits in an integer <c>n</c>.</p>
      <program xml:id="str-ex-digitsacq-sol" interactive="activecode" language="python">
      <code>
def numDigits(n):
  return len(str(n))
print(numDigits(420328))
print(numDigits(5))
      </code>
    </program>
    </solution>
  </activity>

  <activity label="tr-ex-name-sq"><select grade="first" questions="str-ex-nameacq str-ex-nameansw"/></activity>


    <activity xml:id="str-ex-nameInputq"  label="str-ex-nameInputq">
    <statement>
      <p>Write a program that asks a user for their name and from the input prints the first letter of their name in lowercase.</p>
    </statement>
    <program  interactive="activecode" language="python">
      <code>

      </code>
    </program>
    <solution>
      <p>Write a program that asks a user for their name and from the input prints the first letter of their name in lowercase.</p>
      <program xml:id="str-ex-nameacq-sol" interactive="activecode" language="python">
      <code>
prompt = "What is your name?"
name = input(prompt)
first = name[0]
lowerFirst = first.lower()
print(lowerFirst)    
      </code>
    </program>
    </solution>
  </activity>

  <activity label="str-ex-lower-sq"><select grade="first" questions="str-ex-loweracq str-ex-loweransw"/></activity>


</section>
