<?xml version="1.0"?>
<section xml:id="strings_write-code-questions">
  <title>Write Code Questions</title>

  <exercise label="str-ex-hi-toggle">
    <statement>
      <p>Fix line 2 so that it prints <q>Hi</q> instead of <q>hi</q>.</p>
    </statement>
    <program xml:id="str-ex-hiacq" interactive="activecode" language="python">
      <code>
s1 = "hi"
s1.capitalize()
print(s1)

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(s1, "Hi")
myTests().main()
      </code>
    </program>
    <solution>
      <p>Strings are immutable, meaning method alterations like <c>.capitalize()</c> do not modify the string in place. You must assign the result back to <c>s1</c>.</p>
       <code>
s1 = "hi"
s1 = s1.capitalize()
print(s1)

      </code>
    </solution>
  </exercise>

  <exercise label="str-ex-meow-toggle">
    <statement>
      <p>Fix the code so that only <q>meow</q> is printed.</p>
    </statement>
    <program xml:id="str-ex-meowacq" interactive="activecode" language="python">
      <code>
sentence = "The cat goes meow."
s2 = sentence[16:13]
print(s2)

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(s2, "meow")
myTests().main()
      </code>
    </program>
    <blocks xml:id="str-ex-meowansw" submission="program">
      <block order="1">
        <cline>sentence = "The cat goes meow."</cline>
      </block>
      <block order="2">
        <cline>s2 = sentence[16:13]</cline>
      </block>
      <block order="3">
        <cline>print(s2)</cline>
      </block>
    </blocks>
  </exercise>

  <exercise label="str-ex-eggs-toggle">
    <statement>
      <p>Write code to evaluate the length of the string <q>I like green eggs</q> and print it. Save the string as <c>s1</c>. It should print <q>The length is 17</q>.</p>
    </statement>
    <program xml:id="str-ex-eggsacq" interactive="activecode" language="python">
      <code>

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(len(s1), 17, "Verifying string length registration.")
myTests().main()
      </code>
    </program>
    <solution>
      <p>Wrap <c>len()</c> inside <c>str()</c> to safely concatenate it with the leading literal descriptor.</p>
      <code>
s1 = "I like green eggs"
print("The length is " + str(len(s1)))
      </code>
    </solution>
  </exercise>

  <exercise label="str-ex-count-toggle">
    <statement>
      <p>Create a function named <c>count</c> that accepts a string and a letter as arguments, then returns the count of that letter in the string. Hint: use the count method.</p>
    </statement>
    <program xml:id="str-ex-countacq" interactive="activecode" language="python">
      <code>


====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(count('banana','a'), 3)
        self.assertEqual(count('pineapple','s'), 0)
myTests().main()
      </code>
    </program>
    <blocks xml:id="str-ex-countansw" submission="program">
      <block order="1">
        <cline>def count(str, letter):</cline>
      </block>
      <block order="2">
        <cline>    return str.count(letter)</cline>
      </block>
    </blocks>
  </exercise>

  <exercise label="str-ex-float-toggle">
    <statement>
      <p>Extract the portion of the string after the colon character and use the <c>float</c> function to convert it into a float number called <c>num</c>.</p>
    </statement>
    <program xml:id="str-ex-floatacq" interactive="activecode" language="python">
      <code>
string = "X-DSPAM-Confidence: 0.8475"

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(num, 0.8475)
myTests().main()
      </code>
    </program>
    <solution>
      <p>Isolate the colon position index dynamically, slice safely forward from that marker, and run type casting via <c>float()</c>.</p>
      <code>
string = "X-DSPAM-Confidence: 0.8475"
colonpos = string.find(':')
num = float(string[colonpos+1:])
      </code>
    </solution>
  </exercise>

  <exercise label="str-ex-ducks-toggle">
    <statement>
      <p>Fix the loop sequence so that Ouack and Quack are spelled correctly inside execution.</p>
    </statement>
    <program xml:id="str-ex-ducksacq" interactive="activecode" language="python">
      <code>
prefixes = "JKLMNOPQ"
suffix = "ack"

for letter in prefixes:
    if letter == 'O' or letter == 'Q':
        print(letter + "u" + suffix)
    else:
        print(letter + suffix)

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(prefixes, "JKLMNOPQ")
myTests().main()
      </code>
    </program>
    <blocks xml:id="str-ex-ducksansw" submission="program">
      <block order="1">
        <cline>prefixes = "JKLMNOPQ"</cline>
        <cline>suffix = "ack"</cline>
      </block>
      <block order="2">
        <cline>for letter in prefixes:</cline>
      </block>
      <block order="3">
        <cline>    if letter == 'O' or letter == 'Q':</cline>
        <cline>        print(letter + "u" + suffix)</cline>
      </block>
      <block order="4">
        <cline>    else:</cline>
        <cline>        print(letter + suffix)</cline>
      </block>
    </blocks>
    <solution>
      <p>Add a conditional check for edge case prefix strings to inject the missing vowel wrapper matching structural patterns.</p>
    </solution>
  </exercise>

  <exercise label="str-ex-digits-toggle">
    <statement>
      <p>Write a function <c>numDigits</c> that will return the number of digits in an integer <c>n</c>.</p>
    </statement>
    <program xml:id="str-ex-digitsacq" interactive="activecode" language="python">
      <code>
def numDigits(n):
    return len(str(n))

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(numDigits(1352), 4)
myTests().main()
      </code>
    </program>
    <blocks xml:id="str-ex-digitsansw" submission="program">
      <block order="1">
        <cline>def numDigits(n):</cline>
      </block>
      <block order="2">
        <cline>    return len(str(n))</cline>
      </block>
    </blocks>
    <solution>
      <p>Convert integers into type strings to allow measuring aggregate character lengths directly via <c>len()</c>.</p>
    </solution>
  </exercise>

  <exercise label="str-ex-name-toggle">
    <statement>
      <p>Construct a target statement matching specified variable rules without raw string injections.</p>
    </statement>
    <program xml:id="str-ex-nameacq" interactive="activecode" language="python">
      <code>
s1 = "hi"
s2 = "My name is Bob, and he and I love to eat ham."
statement = s1.capitalize() + s2[2:3] + s2.lower()[0:3] + s2[3:14] + s2[22:29] + s2[4:6] + s2[2:3] + str(len(s1))

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(statement, "Hi my name is Bob and I am 2")
myTests().main()
      </code>
    </program>
    <blocks xml:id="str-ex-nameansw" submission="program">
      <block order="1">
        <cline>s1 = "hi"</cline>
        <cline>s2 = "My name is Bob, and he and I love to eat ham."</cline>
      </block>
      <block order="2">
        <cline>statement = s1.capitalize() + s2[2:3] + s2.lower()[0:3] + s2[3:14] + s2[22:29] + s2[4:6] + s2[2:3] + str(len(s1))</cline>
      </block>
    </blocks>
    <solution>
      <p>Isolate and map targeted segment spaces via index slicing arrays cleanly to dynamically form words.</p>
    </solution>
  </exercise>

  <exercise label="str-ex-nameInput-standalone">
    <statement>
      <p>Write a program that asks a user for their name and from the input prints the first letter of their name in lowercase.</p>
    </statement>
    <program xml:id="str-ex-nameInputacq" interactive="activecode" language="python">
      <code>
name = input("What is your name? ")
print(name[0].lower())

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertTrue(True)
myTests().main()
      </code>
    </program>
    <blocks xml:id="str-ex-nameInputansw" submission="program">
      <block order="1">
        <cline>name = input("What is your name? ")</cline>
      </block>
      <block order="2">
        <cline>print(name[0].lower())</cline>
      </block>
    </blocks>
    <solution>
      <p>Capture runtime input data arrays, isolate your structural starting index via position 0, and run <c>.lower()</c> adjustments.</p>
    </solution>
  </exercise>

  <exercise label="str-ex-lower-standalone">
    <statement>
      <p>Write a program that asks for user input and prints their input in all lowercase, as well as the length of their string.</p>
    </statement>
    <program xml:id="str-ex-loweracq" interactive="activecode" language="python">
      <code>
user_input = input("Enter a string: ")
print(user_input.lower() + " " + str(len(user_input)))

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertTrue(True)
myTests().main()
      </code>
    </program>
    <blocks xml:id="str-ex-loweransw" submission="program">
      <block order="1">
        <cline>user_input = input("Enter a string: ")</cline>
      </block>
      <block order="2">
        <cline>print(user_input.lower() + " " + str(len(user_input)))</cline>
      </block>
    </blocks>
    <solution>
      <p>Convert inputs cleanly via string utility actions before transforming character lengths safely via casting structures.</p>
    </solution>
  </exercise>
</section>