<?xml version="1.0"?>
<section xml:id="strings_write-code-questions">
  <title>Write Code Questions</title>

  <exercise label="str-ex-hi-toggle">
    <statement>
      <p>Fix line 2 so that it prints <q>Hi</q> instead of <q>hi</q>.</p>
    </statement>
    <program xml:id="str-ex-hiacq" interactive="activecode" language="python">
      <code>
s1 = "hi"
s1.capitalize()
print(s1)

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(s1, "Hi")
myTests().main()
      </code>
    </program>
    <solution>
      <code>
s1 = "hi"
s1 = s1.capitalize()
print(s1)

      </code>
    </solution>
  </exercise>

  <exercise label="str-ex-meow-toggle">
    <statement>
      <p>Fix the code so that only <q>meow</q> is printed.</p>
    </statement>
    <program xml:id="str-ex-meowacq" interactive="activecode" language="python">
      <code>
sentence = "The cat goes meow."
s2 = sentence[16:13]
print(s2)


      </code>
    </program>
    <blocks xml:id="str-ex-meowansw" submission="program">
      <block order="1">
        <cline>sentence = "The cat goes meow."</cline>
      </block>
      <block order="2">
        <cline>s2 = sentence[16:13]</cline>
      </block>
      <block order="3">
        <cline>print(s2)</cline>
      </block>
    </blocks>
  </exercise>

  <exercise label="str-ex-eggs-toggle">
    <statement>
      <p>Write code to evaluate the length of the string <q>I like green eggs</q> and print it. It should print <q>The length is 17</q>.</p>
    </statement>
    <program xml:id="str-ex-eggsacq" interactive="activecode" language="python">
      <code>

      </code>
    </program>
    <solution>
      <code>
sentence = "I like green eggs"
print("The length is " + str(len(sentence)))
      </code>
    </solution>
  </exercise>

  <exercise label="str-ex-count-toggle">
    <statement>
      <p>Create a function named <c>count</c> that accepts a string and a letter as arguments, then returns the count of that letter in the string. For example, if the function call was <c>count("banana", "a")</c> it would return <c>3</c>. Hint: use the <c>count</c> method.</p>
    </statement>
    <program xml:id="str-ex-countacq" interactive="activecode" language="python">
      <code>


====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(count('banana','a'), 3, "Tested 'a' in 'banana'")
        self.assertEqual(count('pineapple','s'),0,"Tested 's' in 'pineapple'")
        self.assertEqual(count('pepperoni pizza','p'),4,"Tested 'p' in 'pepperoni pizza'")
        self.assertEqual(count('racecar','r'),2,"Tested 'r' in 'racecar'")

myTests().main()
      </code>
    </program>
    <blocks xml:id="str-ex-countansw" submission="program">
      <block order="1">
        <cline>def count(str, letter):</cline>
      </block>
      <block order="2">
        <cline>    return str.count(letter)</cline>
      </block>
    </blocks>
  </exercise>

  <exercise label="str-ex-float-toggle">
    <statement>
      <p>Take the following Python code that stores this string: <c>string = "X-DSPAM-Confidence: 0.8475"</c>. Use <c>find</c> and string slicing to extract the portion of the string after the colon character and then use the <c>float</c> function to convert the extracted string into a floating point number called <c>num</c>.</p>
    </statement>
    <program xml:id="str-ex-floatacq" interactive="activecode" language="python">
      <code>
string = "X-DSPAM-Confidence: 0.8475"

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(num, 0.8475)
myTests().main()
      </code>
    </program>
    <solution>
      <code>
string = "X-DSPAM-Confidence: 0.8475"
colonpos = string.find(':')
num = float(string[colonpos+1:])
      </code>
    </solution>
  </exercise>

  <exercise label="str-ex-ducks-toggle">
    <statement>
      <p>In Robert McCloskey’s book Make Way for Ducklings, the names of the ducklings are Jack, Kack, Lack, Mack, Nack, Ouack, Pack, and Quack. The following loop tries to output these names in order. Of course, that’s not quite right because Ouack and Quack are misspelled. Can you fix it?</p>
    </statement>
    <program xml:id="str-ex-ducksacq" interactive="activecode" language="python">
      <code>
prefixes = "JKLMNOPQ"
suffix = "ack"

for letter in prefixes:
    print(letter + suffix)

      </code>
    </program>
    <blocks xml:id="str-ex-ducksansw" submission="program">
      <block order="1">
        <cline>prefixes = "JKLMNOPQ"</cline>
        <cline>suffix = "ack"</cline>
      </block>
      <block order="2">
        <cline>for letter in prefixes:</cline>
      </block>
      <block order="3">
        <cline>    if letter == 'O' or letter == 'Q':</cline>
        <cline>        print(letter + "u" + suffix)</cline>
      </block>
      <block order="4">
        <cline>    else:</cline>
        <cline>        print(letter + suffix)</cline>
      </block>
    </blocks>
  </exercise>

  <exercise label="str-ex-digits-toggle">
    <statement>
      <p>Write a function <c>numDigits</c> that will return the number of digits in an integer <c>n</c>.</p>
    </statement>
    <program xml:id="str-ex-digitsacq" interactive="activecode" language="python">
      <code>
def numDigits(n):
    # your code here

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(numDigits(2), 1, "Tested numDigits on input of 2")
        self.assertEqual(numDigits(55), 2, "Tested numDigits on input of 55")
        self.assertEqual(numDigits(1352), 4, "Tested numDigits on input of 1352")
        self.assertEqual(numDigits(444), 3, "Tested numDigits on input of 444")
myTests().main()
      </code>
    </program>
    <solution>
      <code>
def numDigits(n):
  return len(str(n))
      </code>
    </solution>
  </exercise>

  <exercise label="str-ex-name-toggle">
    <statement>
      <p>Write code to print out the statement “Hi my name is Bob and I am 2” using only string methods or string slicing. You must get every part of the new string from the given strings, not by using string literals. Name the final string <c>statement</c>.</p>
    </statement>
    <program xml:id="str-ex-nameacq" interactive="activecode" language="python">
      <code>
s1 = "hi"
s2 = "My name is Bob, and he and I love to eat ham."

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(statement, "Hi my name is Bob and I am 2")
myTests().main()
      </code>
    </program>
    <blocks xml:id="str-ex-nameansw" submission="program">
      <block order="1">
        <cline>s1 = "hi"</cline>
        <cline>s2 = "My name is Bob, and he and I love to eat ham."</cline>
      </block>
      <block order="2">
        <cline>statement = s1.capitalize() + s2[2:3] + s2.lower()[0:3] + s2[3:14] + s2[22:29] + s2[4:6] + s2[2:3] + str(len(s1))</cline>
      </block>
    </blocks>
  </exercise>

  <exercise label="str-ex-nameInput-standalone">
    <statement>
      <p>Write a program that asks a user for their name and from the input prints the first letter of their name in lowercase.</p>
    </statement>
    <program xml:id="str-ex-nameInputacq" interactive="activecode" language="python">
      <code>

      </code>
    </program>
    <solution>
      <code>
prompt = "What is your name?"
name = input(prompt)
first = name[0]
lowerFirst = first.lower()
print(lowerFirst)    
      </code>
    </solution>
  </exercise>

  <exercise label="str-ex-lower-standalone">
    <statement>
      <p>Write a program that asks for user input and prints their input in all lowercase, as well as the length of their string.</p>
    </statement>
    <program xml:id="str-ex-loweracq" interactive="activecode" language="python">
      <code>

      </code>
    </program>
    <blocks xml:id="str-ex-loweransw" submission="program">
      <block order="1">
        <cline>user_input = input("Enter a string: ")</cline>
      </block>
      <block order="2">
        <cline>print(user_input.lower() + " " + str(len(user_input)))</cline>
      </block>
    </blocks>
  </exercise>
</section>