<?xml version="1.0"?>
<section xml:id="strings_write-code-questions-orig" component="source-for-select-questions">
  <title>Write Code Questions Orig</title>

  <activity label="str-ex-meowacq" xml:id="str-ex-meowacq">
    <statement>
      <p>Fix the code so that only <q>meow</q> is printed.</p>
    </statement>
    <program interactive="activecode" language="python">
      <code>
sentence = "The cat goes meow."
s2 = sentence[0:17]
print(s2)
      </code>
    </program>
  </activity>

  <activity xml:id="str-ex-meowansw" label="str-ex-meowansw" adaptive="yes" indentation="hide" language="python">
    <statement>
      <p>Arrange the blocks on the right side so that only "meow" is printed.</p>
    </statement>
    <blocks numbered="left">
      <block order="1">
        <cline>sentence = "The cat goes meow."</cline>
      </block>
      <block order="2">
        <choice correct="yes">
          <cline>s2 = sentence[13:17]</cline>
        </choice>
        <choice>
          <cline>s2 = sentence[0:17]</cline>
        </choice>
      </block>
      <block order="3">
        <cline>print(s2)</cline>
      </block>
    </blocks>
  </activity>

  <activity xml:id="str-ex-countacq" label="str-ex-countacq">
    <statement>
      <p>Create a function named <c>count</c> that accepts a string and a letter as arguments, then returns the count of that letter in the string. For example, if the function call was <c>count("banana", "a")</c> it would return <c>3</c>. Hint: use the <c>count</c> method.</p>
    </statement>
    <program interactive="activecode" language="python">
      <code>


====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(count('banana','a'), 3, "Tested 'a' in 'banana'")
        self.assertEqual(count('pineapple','s'),0,"Tested 's' in 'pineapple'")
        self.assertEqual(count('pepperoni pizza','p'),4,"Tested 'p' in 'pepperoni pizza'")
        self.assertEqual(count('racecar','r'),2,"Tested 'r' in 'racecar'")

myTests().main()
      </code>
    </program>
  </activity>

  <activity xml:id="str-ex-countansw" label="str-ex-countansw" adaptive="yes" indentation="hide" language="python">
    <statement>
      <p>Create a function named <c>count</c> that accepts a string and a letter as arguments, then returns the count of that letter in the string. For example, if the function call was <c>count("banana", "a")</c> it would return <c>3</c>. Hint: use the <c>count</c> method.</p>
    </statement>
    <blocks numbered="left">
      <block order="1">
        <choice correct="yes">
           <cline>def count(str, letter):</cline>
        </choice>
        <choice>
          <cline>def count(str, letter)</cline>
        </choice>
      </block>
      <block order="2">
        <choice correct="yes">
        <cline>    return str.count(letter)</cline>
        </choice>
        <choice>
          <cline>    return count(letter)</cline>
        </choice>
      </block>
    </blocks>
  </activity>

  <activity xml:id="str-ex-ducksacq" label="str-ex-ducksacq">
    <statement>
      <p>In Robert McCloskey’s book Make Way for Ducklings, the names of the ducklings are Jack, Kack, Lack, Mack, Nack, Ouack, Pack, and Quack. The following loop tries to output these names in order. Of course, that’s not quite right because Ouack and Quack are misspelled. Can you fix it?</p>
    </statement>
    <program interactive="activecode" language="python">
      <code>
prefixes = "JKLMNOPQ"
suffix = "ack"

for letter in prefixes:
    print(letter + suffix)
      </code>
    </program>
  </activity>

  <activity xml:id="str-ex-ducksansw" label="str-ex-ducksansw" adaptive="yes" indentation="hide" language="python">
    <statement>
      <p>In Robert McCloskey’s book Make Way for Ducklings, the names of the ducklings are Jack, Kack, Lack, Mack, Nack, Ouack, Pack, and Quack. The following loop tries to output these names in order. Of course, that’s not quite right because Ouack and Quack are misspelled. Can you put the blocks in order to fix it?</p>
    </statement>
    <blocks numbered="left">
      <block order="1">
        <cline>prefixes = "JKLMNOPQ"</cline>
        <cline>suffix = "ack"</cline>
      </block>
      <block order="2">
        <cline>for letter in prefixes:</cline>
      </block>
      <block order="3">
        <choice correct="yes">
          <cline>    if letter == 'O' or letter == 'Q':</cline>
          <cline>        print(letter + "u" + suffix)</cline>
        </choice>
        <choice>
          <cline>    if letter == 'o' or letter == 'Q':</cline>
          <cline>        print(letter + "u" + suffix)</cline>
        </choice>
      </block>
      <block order="4">
        <cline>    else:</cline>
        <cline>        print(letter + suffix)</cline>
      </block>
    </blocks>
  </activity>

  <activity xml:id="str-ex-nameacq" label="str-ex-nameacq">
    <statement>
      <p>Write code to print out the statement “Hi my name is Bob2” using only string methods or string slicing. You must get every part of the new string from the given strings, not by using string literals. Name the final string <c>statement</c>.</p>
    </statement>
    <program  interactive="activecode" language="python">
      <code>
s1 = "hi"
s2 = "My name is Bob, and he and I love to eat ham."

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(statement, "Hi my name is Bob2")
myTests().main()
      </code>
    </program>
  </activity>

  <activity xml:id="str-ex-nameansw" label="str-ex-nameansw" adaptive="yes" indentation="hide" language="python">
    <statement>
      <p>Create code to print out the statement “Hi my name is Bob2” using only string methods or string slicing. You must get every part of the new string from the given strings, not by using string literals. Name the final string <c>statement</c>.</p>
    </statement>
    <blocks numbered="left">
      <block order="1">
        <cline>s1 = "hi"</cline>
        <cline>s2 = "My name is Bob, and he and I love to eat ham."</cline>
      </block>
      <block order="2">
        <choice correct="yes">
          <cline>statement = s1.capitalize() + s2[2:3] + s2[0:1].lower() + s2[1:14] + str(len(s1))</cline>
        </choice>
        <choice>
          <cline>statement = s1.capitalize() + s2.[0:13] + len(s1)</cline>
        </choice>
      </block>
    </blocks>
  </activity>

  <activity xml:id="str-ex-loweracq" label="str-ex-loweracq">
    <statement>
      <p>Write a program that asks for user input and prints their input in all lowercase, as well as the length of their string.</p>
    </statement>
    <program interactive="activecode" language="python">
      <code>

      </code>
    </program>
  </activity>

  <activity xml:id="str-ex-loweransw" label="str-ex-loweransw" adaptive="yes" indentation="hide" language="python">
    <statement>
      <p>Create a program that asks for user input and prints their input in all lowercase, as well as the length of their string.</p>
    </statement>
    <blocks numbered="left">
      <block order="1">
        <cline>user_input = input("Enter a string: ")</cline>
      </block>
      <block order="2">
        <choice correct="yes">
          <cline>print(user_input.lower() + " " + str(len(user_input)))</cline>
        </choice>
        <choice>
          <cline>print(user_input.lower() + " " + len(user_input))</cline>
        </choice>
      </block>
    </blocks>
  </activity>

  <activity xml:id="str-ex-hiA-pp" label="str-ex-hiA-pp" adaptive="yes" indentation="hide" language="python">
    <statement>
      <p>Arrange the blocks to fix the code so that it prints <q>Hi</q> instead of <q>hi</q>.</p>
    </statement>
    <blocks numbered="left">
      <block order="1">
        <cline>s1 = "hi"</cline>
      </block>
      <block order="2">
        <choice correct="yes">
          <cline>s1 = s1.capitalize()</cline>
        </choice>
        <choice>
          <cline>s1.capitalize()</cline>
        </choice>
      </block>
      <block order="3">
        <cline>print(s1)</cline>
      </block>
    </blocks>
  </activity>

  <activity xml:id="str-ex-eggs-pp" label="str-ex-eggs-pp" adaptive="yes" indentation="hide" language="python">
    <statement>
      <p>Arrange the blocks to evaluate the length of the string <q>I like green eggs</q>, store it in a variable named <c>length</c>, and print it, so that it prints <q>The length is 17</q>.</p>
    </statement>
    <blocks numbered="left">
      <block order="1">
        <cline>sentence = "I like green eggs"</cline>
      </block>
      <block order="2">
        <choice correct="yes">
          <cline>length = len(sentence)</cline>
        </choice>
        <choice>
          <cline>length = str(sentence)</cline>
        </choice>
      </block>
      <block order="3">
        <cline>print("The length is " + str(length))</cline>
      </block>
    </blocks>
  </activity>

  <activity xml:id="str-ex-floatq-pp" label="str-ex-floatq-pp" adaptive="yes" indentation="hide" language="python">
    <statement>
      <p>Arrange the blocks to use <c>find</c> and string slicing to extract the portion of <c>string</c> after the colon character and convert it into a floating point number called <c>num</c>.</p>
    </statement>
    <blocks numbered="left">
      <block order="1">
        <cline>string = "X-DSPAM-Confidence: 0.8475"</cline>
      </block>
      <block order="2">
        <choice correct="yes">
          <cline>colonpos = string.find(':')</cline>
        </choice>
        <choice>
          <cline>colonpos = string.find(".")</cline>
        </choice>
      </block>
      <block order="3">
        <choice correct="yes">
          <cline>num = float(string[colonpos+1:])</cline>
        </choice>
        <choice>
          <cline>num = float(string[colonpos:])</cline>
        </choice>
      </block>
      <block order="4">
        <cline>print(num)</cline>
      </block>
    </blocks>
  </activity>

  <activity xml:id="str-ex-digitsq-pp" label="str-ex-digitsq-pp" adaptive="yes" indentation="hide" language="python">
    <statement>
      <p>Arrange the blocks to define a function <c>numDigits</c> that returns the number of digits in an integer <c>n</c>.</p>
    </statement>
    <blocks numbered="left">
      <block order="1">
        <choice correct="yes">
          <cline>def numDigits(n):</cline>
        </choice>
        <choice>
          <cline>def numDigits(n)</cline>
        </choice>
      </block>
      <block order="2">
        <choice correct="yes">
          <cline>    return len(str(n))</cline>
        </choice>
        <choice>
          <cline>    return len(n)</cline>
        </choice>
      </block>
    </blocks>
  </activity>

  <activity xml:id="str-ex-nameInputq-pp" label="str-ex-nameInputq-pp" adaptive="yes" indentation="hide" language="python">
    <statement>
      <p>Arrange the blocks to ask the user for their name and print the first letter of their name in lowercase.</p>
    </statement>
    <blocks numbered="left">
      <block order="1">
        <cline>prompt = "What is your name?"</cline>
        <cline>name = input(prompt)</cline>
      </block>
      <block order="2">
        <cline>first = name[0]</cline>
      </block>
      <block order="3">
        <choice correct="yes">
          <cline>lowerFirst = first.lower()</cline>
        </choice>
        <choice>
          <cline>lowerFirst = first.upper()</cline>
        </choice>
      </block>
      <block order="4">
        <cline>print(lowerFirst)</cline>
      </block>
    </blocks>
  </activity>
</section>