<?xml version="1.0"?>
<worksheet xml:id="fsc_strsconds_groupsub" groupwork="yes" groupsize="4" label="funcWithStrsAndConds">
  <title>Group Work: Functions with Strings and Conditionals</title>
   <p>It is best to use a POGIL approach with the following. In POGIL students work
            in groups on activities and each member has an assigned role.  For more information see <url href="https://cspogil.org/Home" visual="https://cspogil.org/Home">https://cspogil.org/Home</url>.</p>
  <note>
    <p>If you work in a group, have only one member of the group fill in the answers on this page.  You will be able to share your answers with the group at the bottom of the page.</p>
  </note>
<section xml:id="fsc_group-work-functions-with-strings-and-conditionals">
  <title>Learning Objectives</title>
  <introduction>
  <p>Students will know and be able to do the following.</p>
  <p>
    <em>Content Objectives:</em>
  </p>
  <p>
    <ul>
      <li>
        <p>Learn how to display a value in Python using print.</p>
      </li>
      <li>
        <p>Match terms for functions to their definitiions.</p>
      </li>
      <li>
        <p>Recognize common symbols (:, #), operators (+, -, /), and keywords (def, return) in Python.</p>
      </li>
      <li>
        <p>Recognize strings.</p>
      </li>
      <li>
        <p>Use string indices.</p>
      </li>
      <li>
        <p>Understand how slice works for both positive and negative indices.</p>
      </li>
      <li>
        <p>Use input and convert between strings and numbers.</p>
      </li>
      <li>
        <p>Recognize common string methods.</p>
      </li>
      <li>
        <p>Recognize the keywords used in conditionals in Python (if, elif, else).</p>
      </li>
      <li>
        <p>Predict the output from simple and complex conditionals.</p>
      </li>
      <li>
        <p>Recognize the keywords used in complex conditionals (and, or, not).</p>
      </li>
    </ul>
  </p>
  <p>
    <em>Process Objectives:</em>
  </p>
  <p>
    <ul>
      <li>
        <p>Learn how to properly indent the body of a function</p>
      </li>
      <li>
        <p>Predict the output from code.</p>
      </li>
      <li>
        <p>Modify code to produce the correct output.</p>
      </li>
      <li>
        <p>Create good test data for functions with conditionals</p>
      </li>
      <li>
        <p>Modify code to use complex conditionals</p>
      </li>
    </ul>
  </p>
</introduction>
  <subsection xml:id="fsc_str_cond_funct_intro">
    <title>Function Example</title>
    <p>A function is a name for one or more lines of code. You first define a function using the <c>def</c> keyword and then execute 
       it using <c>function_name(arguments)</c>.  
    </p>
    <exercise label="fsc_fitb_print_test_first_line_v3">
      <statement>
        <p>What is the first thing that will be printed when the code below runs? <var/>  </p>
      </statement>
      <setup>
        <var>
          <condition string="Hello!">
            <feedback>
              <p>The main function prints this before it calls the test function.</p>
            </feedback>
          </condition>
          <condition string=".*">
            <feedback>
              <p>The first function that is executed is <c>main</c>. What does that print first?</p>
            </feedback>
          </condition>
        </var>
      </setup>
    </exercise>
    <exercise label="fsc_ac_pogil_print_message_v2">
      <statement>
        <p>Run the code below to see what it prints and then use it to answer the following questions.</p>
      </statement>
      <program xml:id="fsc_ac_pogil_print_message_v2_ac_editor" interactive="activecode" language="python">
        <input>
# function definition
def test(a, b = 2):
    print("Welcome")
    print("Learn the power of functions!")
    print(a + b)
    print(a - b)
    print(a * b)
    print(a / b)
    print(a // b)
    a = 5 % 2
    return a

# function definition
def main():
    print("Hello!")

    # function call
    print(test(3))

# function call
main()
        </input>
      </program>
    </exercise>

  <note>
    <p>You do not declare the type of a variable in Python. Python requires that all statements that are part of the body of a function must be indented. Use four spaces to indent.</p>
  </note>

    <exercise label="fsc_fitb_print_test_value_b">
      <statement>
        <p>What is the value of <c>b</c> when function <c>test</c> starts executing?</p>
      </statement>
      <setup>
        <var>
          <condition number="2">
            <feedback>
              <p>Since the value of <c>b</c> was not specified it defaults to 2 (it was set to 2 as a default in the function definition)</p>
            </feedback>
          </condition>
          <condition string=".*">
            <feedback>
              <p>What is the default value specified for <c>b</c> in the function definition?</p>
            </feedback>
          </condition>
        </var>
      </setup>
    </exercise>
    <p>Delete the last line of the test function above and run the code again.</p>
    <exercise label="fsc_fitb_return_None">
      <statement>
        <p>What value is returned from a function that doesn't have a <c>return</c> keyword? <var/>  </p>
      </statement>
      <setup>
        <var>
          <condition string="None">
            <feedback>
              <p>A function without a <c>return</c> still returns the keyword None</p>
            </feedback>
          </condition>
          <condition string=".*">
            <feedback>
              <p>What is the last thing that is printed when you run the code above after deleting the return from the test function?</p>
            </feedback>
          </condition>
        </var>
      </setup>
    </exercise>

    <exercise label="fsc_pp_print_name_v2" numbered="yes" adaptive="yes" indentation="hide" language="python">
      <statement>
        <p>Drag the blocks from the left and put them in the correct order on the right to define a function <c>print_greeting</c> that asks for your name and prints <q>Hello Name</q>. Then define a <c>main</c> function that calls <c>print_greeting</c>.  Be sure to also call the <c>main</c> function. Note that you will have to indent the lines that are in the body of each function.  Click the <em>Check</em> button to check your solution.</p>
      </statement>
      <blocks>
        <block order="5">
          <choice correct="yes">
            <cline>def print_greeting():</cline>
          </choice>
          <choice>
            <cline>def print_greeting() </cline>
          </choice>
        </block>
        <block order="8">
          <cline>    name = input("What is your name?")</cline>
        </block>
        <block order="9">
          <cline>    print("Hello " + name)</cline>
        </block>
        <block order="7">
          <choice correct="yes">
            <cline>def main():</cline>
          </choice>
          <choice>
            <cline>Def main(): </cline>
          </choice>
        </block>
        <block order="4">
          <choice correct="yes">
            <cline>    print_greeting()</cline>
          </choice>
          <choice>
            <cline>    print_greeting </cline>
          </choice>
        </block>
        <block order="3">
          <cline>main()</cline>
        </block>
      </blocks>
    </exercise>
  </subsection>
<subsection xml:id="fsc_str_cond_string_def">
    <title>String Definition</title>
    <p>A <term>string</term> is a sequence of characters enclosed in quotes.  In Python you can use pairs of single
or double quotes to enclose a string like in <c>"hello"</c> or <c>'hello'</c>. This is especially useful when a string includes a single quote <c>"they're"</c>.  You can even use tripe quotes when
a string covers more than one line. </p>
</subsection>
<subsection xml:id="fsc_str_cond_string_index">
    <title>String Indices</title>
    <p>You can get a character from a string at an index (position).</p>
    <exercise label="fsc_fitb_pogil_initials">
      <statement>
        <p>What is the last thing that will be printed when the code below runs? <var/>  </p>
      </statement>
      <setup>
        <var>
          <condition string="Jk">
            <feedback>
              <p>It prints the first letter of the first name and last letter of the last name.</p>
            </feedback>
          </condition>
          <condition string=".*">
            <feedback>
              <p>String indices start with 0 and -1 is the index of the last letter in a string.</p>
            </feedback>
          </condition>
        </var>
      </setup>
    </exercise>
    <exercise label="fsc_ac_pogil_initials_v2">
      <statement>
        <p>Run the code below to see what it prints.  Then fix it to pass the given test.  It should return a string with the
                first character of the first name and first character of the last name.</p>
      </statement>
      <program xml:id="fsc_ac_pogil_initials_v2_editor" interactive="activecode" language="python">
        <input>
# function definition
def get_initials(first, last):
    return first[0] + last[-1]

# function definition
def main():
    print(type("Hello"))
    print(type('Class'))
    print(type(42))

    print(get_initials("J'Quan",'Alik'))


# function call
main()

from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(get_initials("J'Quan",'Alik'),"JA",'''get_initials("J'Quan",'Alik')''')

myTests().main()
        </input>
      </program>
    </exercise>
</subsection>
<subsection xml:id="fsc_string-slices">
    <title>String Slices</title>
    <exercise label="fsc_fitb_pogil_short_name">
      <statement>
        <p>What is the last thing that will be printed when the code below runs? <var/>  </p>
      </statement>
      <setup>
        <var>
          <condition string="Sibs">
            <feedback>
              <p>The last thing it prints are the first two letters of the first name and the last two letters of the last name.</p>
            </feedback>
          </condition>
          <condition string=".*">
            <feedback>
              <p>A slice starts with the first number and ends before the second.  If the second is left off it goes to the end of the string. The last character in a string is at index -1.</p>
            </feedback>
          </condition>
        </var>
      </setup>
    </exercise>
    <exercise label="fsc_ac_pogil_short_name">
      <statement>
        <p>Run the code below to see what it prints.</p>
      </statement>
      <program xml:id="fsc_ac_pogil_short_name_editor" interactive="activecode" language="python">
        <input>
# function definition
def get_short_name(first, last):
    print(len(first))
    print(len(last))
    return first[:2] + last[-2:]

# function definition
def main():
    print(get_short_name('Simona',"Jacobs"))

# function call
main()
        </input>
      </program>
    </exercise>
    <note>
      <p>Use the slice [start:end] operator to get a slice (substring) from a string. It will return a new string starting at the start and including all the characters up to just before the end (end - 1).</p>
    </note>
    <exercise label="fsc_fitb_three_char_slice">
      <statement>
        <p>Use the slice operator to return the first three characters from the variable <c>dna</c>? <var/>  </p>
      </statement>
      <setup>
        <var>
          <condition string="dna\[0\:3\]|\[:3\]">
            <feedback>
              <p>This will return a new string with the characters from index 0 to 2.</p>
            </feedback>
          </condition>
          <condition string=".*">
            <feedback>
              <p>Look at the note above and try again.</p>
            </feedback>
          </condition>
        </var>
      </setup>
    </exercise>
  </subsection>
</section>
</worksheet>
