<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Mixed-up Code Questions &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/637.css?v=bddd2011b20db76a20ea" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.css?v=7c5388ad4389eae1a318" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/runtime.bundle.js?v=aa7c717c38c5405af9e6"></script>
    <script src="../_static/637.bundle.js?v=1123ad75f9b821cdd636"></script>
    <script src="../_static/runestone.bundle.js?v=9c8e466e8cbed4a3af5a"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'py4e-int';
    eBookConfig.basecourse = 'py4e-int';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          py4e-int
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=MixedupCode_orig'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
  
<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="mixed-up-code-questions">
<h1>Mixed-up Code Questions<a class="headerlink" href="#mixed-up-code-questions" title="Permalink to this headline">¶</a></h1>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="str-mixed-emotion" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create the function, <code class="docutils literal notranslate"><span class="pre">mood(emotion)</span></code>, to return a string with the <code class="docutils literal notranslate"><span class="pre">emotion</span></code>
followed by <code class="docutils literal notranslate"><span class="pre">&quot;</span> <span class="pre">4</span> <span class="pre">you&quot;</span></code>. For example, if <code class="docutils literal notranslate"><span class="pre">emotion</span></code> is “So happy” it should
return <code class="docutils literal notranslate"><span class="pre">&quot;So</span> <span class="pre">happy</span> <span class="pre">4</span> <span class="pre">you!&quot;</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="1"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def mood(emotion):
---
    out = emotion + " 4 you!"
---
    out = emotion + "4 you!" #paired
---
    return out
---
    return emotion #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=str-mixed-emotion-ac data-question_label="2">
<div id=str-mixed-emotion-ac_question class="ac_question col-md-12">
<p>Write the function, <code class="docutils literal notranslate"><span class="pre">mood(emotion)</span></code>, to return a string with the <code class="docutils literal notranslate"><span class="pre">emotion</span></code>
followed by <code class="docutils literal notranslate"><span class="pre">&quot;</span> <span class="pre">4</span> <span class="pre">you&quot;</span></code>. For example, if <code class="docutils literal notranslate"><span class="pre">emotion</span></code> is “So happy” it should
return <code class="docutils literal notranslate"><span class="pre">&quot;So</span> <span class="pre">happy</span> <span class="pre">4</span> <span class="pre">you!&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="str-mixed-emotion-ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(mood("So happy"), "So happy 4 you!", 'mood("So happy")')
        self.assertEqual(mood("So HaPpY"), "So HaPpY 4 you!",  'mood("So HaPpY")')
        self.assertEqual(mood("HAPPY"), "HAPPY 4 you!", 'mood("HAPPY")')
        self.assertEqual(mood("happy"), "happy 4 you!", 'mood("happy")')
        self.assertEqual(mood("Happy"), "Happy 4 you!", 'mood("Happy")')
        self.assertEqual(mood("So glad"), "So glad 4 you!", 'mood("So glad")')

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="str-mixed-book" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create the function, <code class="docutils literal notranslate"><span class="pre">item_purchase(item,</span> <span class="pre">price)</span></code>, to return a string with the <code class="docutils literal notranslate"><span class="pre">item</span></code> purchased followed by the <code class="docutils literal notranslate"><span class="pre">price</span></code>.
For example if <code class="docutils literal notranslate"><span class="pre">item</span></code> is “new book” and price is <code class="docutils literal notranslate"><span class="pre">12</span></code> the function should return <code class="docutils literal notranslate"><span class="pre">&quot;My</span> <span class="pre">new</span> <span class="pre">book</span> <span class="pre">cost</span> <span class="pre">$12&quot;</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="3"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def item_purchase(item, price):
---
def item_purchase(price, item): #paired
---
    out = "My " + item + " cost $" + str(price)
---
    out = "My " + item + " cost $" + price #paired
---
    return out
---
    Return out #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=str-mixed-book-ac data-question_label="4">
<div id=str-mixed-book-ac_question class="ac_question col-md-12">
<p>Write the function, <code class="docutils literal notranslate"><span class="pre">item_purchase(item,</span> <span class="pre">price)</span></code>, to return a string with the <code class="docutils literal notranslate"><span class="pre">item</span></code> purchased followed by the <code class="docutils literal notranslate"><span class="pre">price</span></code>.
For example if <code class="docutils literal notranslate"><span class="pre">item</span></code> is “new book” and price is <code class="docutils literal notranslate"><span class="pre">12</span></code> the function should return <code class="docutils literal notranslate"><span class="pre">&quot;My</span> <span class="pre">new</span> <span class="pre">book</span> <span class="pre">cost</span> <span class="pre">$12&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="str-mixed-book-ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(item_purchase("new book", "12"), "My new book cost $12", 'item_purchase("new book", "12")')
        self.assertEqual(item_purchase("makeup", "14"), "My makeup cost $14", 'item_purchase("makeup", "14")')
        self.assertEqual(item_purchase("new bag", "12.31"), "My new bag cost $12.31", 'item_purchase("new bag", "12.31")')
        self.assertEqual(item_purchase("backpack", "0"), "My backpack cost $0", 'item_purchase("backpack", "0")')

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="str-mixed-code" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create the function <code class="docutils literal notranslate"><span class="pre">phrase(person,</span> <span class="pre">thing)</span></code>, to return a string with a <code class="docutils literal notranslate"><span class="pre">person</span></code> of your choosing followed by a <code class="docutils literal notranslate"><span class="pre">thing</span></code>
of what that person likes to do. For example, if the <code class="docutils literal notranslate"><span class="pre">person</span></code> is <code class="docutils literal notranslate"><span class="pre">sam</span></code> and <code class="docutils literal notranslate"><span class="pre">thing</span></code> is <code class="docutils literal notranslate"><span class="pre">Likes</span> <span class="pre">to</span> <span class="pre">code</span></code>, the
returned string should be <code class="docutils literal notranslate"><span class="pre">&quot;Sam</span> <span class="pre">likes</span> <span class="pre">to</span> <span class="pre">code&quot;</span></code>. Make sure that <code class="docutils literal notranslate"><span class="pre">person</span></code> is capitalized and <code class="docutils literal notranslate"><span class="pre">thing</span></code> is lowercase.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="5"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def phrase(person, thing):
---
def phrase(person, thing) #paired
---
    person = person.capitalized()
    thing = thing.lower()
---
    person.capitalized() #paired
    thing.lower()
---
    return person + " " + thing
---
    return person + thing #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=str-mixed-code-ac data-question_label="6">
<div id=str-mixed-code-ac_question class="ac_question col-md-12">
<p>Write the function, <code class="docutils literal notranslate"><span class="pre">phrase(person,</span> <span class="pre">thing)</span></code>, to return a string with a <code class="docutils literal notranslate"><span class="pre">person</span></code> of your choosing followed by a <code class="docutils literal notranslate"><span class="pre">thing</span></code>
of what that person likes to do. For example, if the <code class="docutils literal notranslate"><span class="pre">person</span></code> is <code class="docutils literal notranslate"><span class="pre">Sam</span></code> and <code class="docutils literal notranslate"><span class="pre">thing</span></code> is <code class="docutils literal notranslate"><span class="pre">Likes</span> <span class="pre">to</span> <span class="pre">code</span></code>, the
returned string should be <code class="docutils literal notranslate"><span class="pre">&quot;Sam</span> <span class="pre">likes</span> <span class="pre">to</span> <span class="pre">code&quot;</span></code>.  Make sure that <code class="docutils literal notranslate"><span class="pre">person</span></code> is capitalized and <code class="docutils literal notranslate"><span class="pre">thing</span></code> is lowercase.</p>

</div>
<textarea data-lang="python" id="str-mixed-code-ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(phrase("sam", "Likes to code"), "Sam likes to code", 'phrase("sam", " Likes to code")')
        self.assertEqual(phrase("mary-anne", "likes to sing"), "Mary-anne likes to sing", 'phrase("mary-anne", " likes to sing")')
        self.assertEqual(phrase("ANNA", "likes to dance"), "Anna likes to dance", 'phrase("ANNA", " likes to dance")')


myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="str-mixed-time" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Let’s imagine that last weekend you went on a camping trip with your family. Create a function
<code class="docutils literal notranslate"><span class="pre">time(hours,</span> <span class="pre">minutes)</span></code>. Return a string that says how long it took you to get home in minutes. For example,
if <code class="docutils literal notranslate"><span class="pre">hours</span></code> was <code class="docutils literal notranslate"><span class="pre">3</span></code> and <code class="docutils literal notranslate"><span class="pre">minutes</span></code> was <code class="docutils literal notranslate"><span class="pre">13</span></code>, the returned string should be, <code class="docutils literal notranslate"><span class="pre">&quot;It</span> <span class="pre">took</span> <span class="pre">us</span> <span class="pre">193</span> <span class="pre">minutes</span> <span class="pre">to</span> <span class="pre">get</span> <span class="pre">home&quot;</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="7"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def time(hours, minutes):
---
def time(hours, minutes) #paired
---
    total = hours * 60 + minutes
---
    total = hours + 60 * minutes #paired
---
    return "It took us " + str(total) + " minutes to get home"
---
    return "It took us " + str(total) " minutes to get home"  #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=str-mixed-time-ac data-question_label="8">
<div id=str-mixed-time-ac_question class="ac_question col-md-12">
<p>Let’s imagine that last weekend you went on a camping trip with your family. Write a function
<code class="docutils literal notranslate"><span class="pre">time(hours,</span> <span class="pre">minutes)</span></code>.  Return a string that says how long it took you to get home in minutes. For example,
if <code class="docutils literal notranslate"><span class="pre">hours</span></code> was <code class="docutils literal notranslate"><span class="pre">3</span></code> and <code class="docutils literal notranslate"><span class="pre">minutes</span></code> was <code class="docutils literal notranslate"><span class="pre">13</span></code>, the returned string should be, <code class="docutils literal notranslate"><span class="pre">&quot;It</span> <span class="pre">took</span> <span class="pre">us</span> <span class="pre">193</span> <span class="pre">minutes</span> <span class="pre">to</span> <span class="pre">get</span> <span class="pre">home&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="str-mixed-time-ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(time(2, 45), "It took us 165 minutes to get home", 'time(2, 45)')
        self.assertEqual(time(2, 15), "It took us 135 minutes to get home", 'time(2, 15)')
        self.assertEqual(time(1, 23), "It took us 83 minutes to get home", 'time(1, 23)')

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="str-mixed-feeling" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create the function, <code class="docutils literal notranslate"><span class="pre">food(fave,</span> <span class="pre">least_fave)</span></code>, to return a string with your <code class="docutils literal notranslate"><span class="pre">fave</span></code> food followed by your <code class="docutils literal notranslate"><span class="pre">least_fave</span></code> food both in lowercase.
For example, if your <code class="docutils literal notranslate"><span class="pre">fave</span></code> food is “PIZZA” and your <code class="docutils literal notranslate"><span class="pre">least_fave</span></code> is “broccoli”, return
<code class="docutils literal notranslate"><span class="pre">&quot;My</span> <span class="pre">favorite</span> <span class="pre">food</span> <span class="pre">is</span> <span class="pre">pizza,</span> <span class="pre">but</span> <span class="pre">I</span> <span class="pre">don't</span> <span class="pre">like</span> <span class="pre">broccoli&quot;</span></code>. Be sure to change <code class="docutils literal notranslate"><span class="pre">fave</span></code> and <code class="docutils literal notranslate"><span class="pre">least_fave</span></code> to lowercase.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="9"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def food(fave, least_fave):
---
    fave_l = fave.lower()
    least_l = least_fave.lower()
    start = "My favorite food is "
    end = ", but I don't like "
---
    fave.lower() #paired
    least_fave.lower()
    start = "My favorite food is "
    end = ", but I don't like "
---
    return start + fave_l + end + least_l
---
    return start + fave_l + least_l + end #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=str-mixed-feeling-ac data-question_label="10">
<div id=str-mixed-feeling-ac_question class="ac_question col-md-12">
<p>Write the function, <code class="docutils literal notranslate"><span class="pre">food(fave,</span> <span class="pre">least_fave)</span></code>, to return a string with your <code class="docutils literal notranslate"><span class="pre">fave</span></code> food followed by your <code class="docutils literal notranslate"><span class="pre">least_fave</span></code> food both in lowercase.
For example, if your <code class="docutils literal notranslate"><span class="pre">fave</span></code> food is “PIZZA” and your <code class="docutils literal notranslate"><span class="pre">least_fave</span></code> is “broccoli”, return
<code class="docutils literal notranslate"><span class="pre">&quot;My</span> <span class="pre">favorite</span> <span class="pre">food</span> <span class="pre">is</span> <span class="pre">pizza,</span> <span class="pre">but</span> <span class="pre">I</span> <span class="pre">don't</span> <span class="pre">like</span> <span class="pre">broccoli&quot;</span></code>. Be sure to change <code class="docutils literal notranslate"><span class="pre">fave</span></code> and <code class="docutils literal notranslate"><span class="pre">least_fave</span></code> to lowercase.</p>

</div>
<textarea data-lang="python" id="str-mixed-feeling-ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(food("pizza", "broccoli"), "My favorite food is pizza, but I don't like broccoli", 'food("pizza", "broccoli")')
        self.assertEqual(food("pizza", "pizza"), "My favorite food is pizza, but I don't like pizza", 'food("pizza", "pizza")')
        self.assertEqual(food("HAMBURGER", "caRrOts"), "My favorite food is hamburger, but I don't like carrots", 'food("HAMBURGER", "caRrOts")')

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="str-mixed-animal" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create the function, <code class="docutils literal notranslate"><span class="pre">animal(type)</span></code>, to return the string, “My <code class="docutils literal notranslate"><span class="pre">type_animal</span></code> looks like a Muppet” with <code class="docutils literal notranslate"><span class="pre">type</span></code> being
any animal of your choosing. For example, if your <code class="docutils literal notranslate"><span class="pre">type</span></code> is “sheepdog”, then the returned string should
be <code class="docutils literal notranslate"><span class="pre">&quot;My</span> <span class="pre">sheepdog</span> <span class="pre">looks</span> <span class="pre">like</span> <span class="pre">a</span> <span class="pre">Muppet&quot;</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="11"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def animal(type):
---
    end = " looks like a Muppet"
---
    end = "looks like a Muppet" #paired
---
    return "My " + type + end
---
    return "My" + type + end #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=str-mixed-animal-ac data-question_label="12">
<div id=str-mixed-animal-ac_question class="ac_question col-md-12">
<p>Finish the function, <code class="docutils literal notranslate"><span class="pre">animal(type)</span></code>, to return the string, “My <code class="docutils literal notranslate"><span class="pre">type_animal</span></code> looks like a Muppet” with <code class="docutils literal notranslate"><span class="pre">type</span></code> being
any animal of your choosing. For example, if your <code class="docutils literal notranslate"><span class="pre">type</span></code> is “sheepdog”, then the returned string should
be <code class="docutils literal notranslate"><span class="pre">&quot;My</span> <span class="pre">sheepdog</span> <span class="pre">looks</span> <span class="pre">like</span> <span class="pre">a</span> <span class="pre">Muppet&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="str-mixed-animal-ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(animal("sheepdog"), "My sheepdog looks like a Muppet", 'animal("sheepdog")')
        self.assertEqual(animal("FROG"), "My FROG looks like a Muppet", 'animal("FROG")')
        self.assertEqual(animal("c a t"), "My c a t looks like a Muppet", 'animal("c a t")')

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="str-mixed-add" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Let’s imagine you went on a huge shopping spree. Create the function, <code class="docutils literal notranslate"><span class="pre">total(store1,</span> <span class="pre">store2,</span> <span class="pre">store3)</span></code>,
to return the total of how much you spent at the three stores. For example, if you spent <code class="docutils literal notranslate"><span class="pre">30</span></code> in <code class="docutils literal notranslate"><span class="pre">store1</span></code>, <code class="docutils literal notranslate"><span class="pre">13</span></code> in <code class="docutils literal notranslate"><span class="pre">store2</span></code>, and <code class="docutils literal notranslate"><span class="pre">56</span></code> in <code class="docutils literal notranslate"><span class="pre">store3</span></code>,
<code class="docutils literal notranslate"><span class="pre">ans</span></code> return <code class="docutils literal notranslate"><span class="pre">&quot;$30</span> <span class="pre">+</span> <span class="pre">$13</span> <span class="pre">+</span> <span class="pre">$56</span> <span class="pre">=</span> <span class="pre">$99&quot;</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="13"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def total(store1, store2, store3):
---
    total = store1 + store2 + store3
---
    s1 = "$" + str(store1)
    s2 = "$" + str(store2)
    s3 = "$" + str(store3)
    t = "$" + str(total)
---
    return s1 + " + " + s2 + " + " + s3 + " = " + t
---
    return s1 + s2 + s3 + t #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=str-mixed-add-ac data-question_label="14">
<div id=str-mixed-add-ac_question class="ac_question col-md-12">
<p>Let’s imagine you went on a huge shopping spree. Write the function, <code class="docutils literal notranslate"><span class="pre">total(store1,</span> <span class="pre">store2,</span> <span class="pre">store3)</span></code>,
to return the total of how much you spent at the three stores. For example, if you spent <code class="docutils literal notranslate"><span class="pre">30</span></code> in <code class="docutils literal notranslate"><span class="pre">store1</span></code>, <code class="docutils literal notranslate"><span class="pre">13</span></code> in <code class="docutils literal notranslate"><span class="pre">store2</span></code>, and <code class="docutils literal notranslate"><span class="pre">56</span></code> in <code class="docutils literal notranslate"><span class="pre">store3</span></code>,
<code class="docutils literal notranslate"><span class="pre">ans</span></code> return <code class="docutils literal notranslate"><span class="pre">&quot;$30</span> <span class="pre">+</span> <span class="pre">$13</span> <span class="pre">+</span> <span class="pre">$56</span> <span class="pre">=</span> <span class="pre">$99&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="str-mixed-add-ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(total(30,13,56), "$30 + $13 + $56 = $99", 'total(30,13,56,99)')
        self.assertEqual(total(1.5,3.5,13), "$1.5 + $3.5 + $13 = $18.0", 'total(1.5,3.5,13,17)')
        self.assertEqual(total(1,1,1), "$1 + $1 + $1 = $3", 'total(1,1,1,3)')


myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="str-mixed-wizard" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create the function <code class="docutils literal notranslate"><span class="pre">combine(var1,</span> <span class="pre">var2)</span></code>, to return a string by combining the two variables <code class="docutils literal notranslate"><span class="pre">var1</span></code> and <code class="docutils literal notranslate"><span class="pre">var2</span></code>.
For example, if <code class="docutils literal notranslate"><span class="pre">var1</span></code> is “i am a” and <code class="docutils literal notranslate"><span class="pre">var2</span></code> is “CS wizard”, the function should return the string, <code class="docutils literal notranslate"><span class="pre">&quot;I</span> <span class="pre">am</span> <span class="pre">a</span> <span class="pre">cs</span> <span class="pre">wizard&quot;</span></code>.
Make sure <code class="docutils literal notranslate"><span class="pre">var1</span></code> is capitalized and <code class="docutils literal notranslate"><span class="pre">var2</span></code> is lowercase.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="15"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def combine(var1, var2):
---
    var1 = var1.capitalize()
    var2 = var2.lower()
---
    var1.capitalize() #paired
    var2.lower()
---
    return var1 + " " + var2
---
    return var1 + var2 #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=str-mixed-wizard-ac data-question_label="16">
<div id=str-mixed-wizard-ac_question class="ac_question col-md-12">
<p>Write the function, <code class="docutils literal notranslate"><span class="pre">combine(var1,</span> <span class="pre">var2)</span></code>, to return a string by combining the two variables <code class="docutils literal notranslate"><span class="pre">var1</span></code> and <code class="docutils literal notranslate"><span class="pre">var2</span></code>.
For example, if <code class="docutils literal notranslate"><span class="pre">var1</span></code> is “i am a ” and <code class="docutils literal notranslate"><span class="pre">var2</span></code> is “CS wizard”, the function should return the string, <code class="docutils literal notranslate"><span class="pre">&quot;I</span> <span class="pre">am</span> <span class="pre">a</span> <span class="pre">cs</span> <span class="pre">wizard&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="str-mixed-wizard-ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(combine("i am a", "CS wizard"), "I am a cs wizard", 'combine("i am a", "CS wizard")')
        self.assertEqual(combine("Potatos", "are cool"), "Potatos are cool", 'combine("Potatos", "are cool")')
        self.assertEqual(combine("i am a", "Meatloaf"), "I am a meatloaf", 'combine("i am a", "Meatloaf")')
        self.assertEqual(combine("I like", "bubbles"), "I like bubbles", 'combine("I like", "bubbles")')

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="str-mixed-colors" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create the function, <code class="docutils literal notranslate"><span class="pre">colors(col1,</span> <span class="pre">col2)</span></code>, should return a string combining the variables <code class="docutils literal notranslate"><span class="pre">col1</span></code> and <code class="docutils literal notranslate"><span class="pre">col2</span></code>. For example,
if <code class="docutils literal notranslate"><span class="pre">col1</span></code> and <code class="docutils literal notranslate"><span class="pre">col2</span></code> is “‘red’” and “‘blue’”, then the function should return <code class="docutils literal notranslate"><span class="pre">Red</span> <span class="pre">is</span> <span class="pre">a</span> <span class="pre">primary</span> <span class="pre">color</span> <span class="pre">and</span>
<span class="pre">so</span> <span class="pre">is</span> <span class="pre">blue</span></code>.  Be sure to capitalize the first color and lowercase at second.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="17"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def colors(col1, col2):
---
    mid = " is a primary color, and so is "
---
    mid = "is a primary color, and so is" #paired
---
    col1 = col1.capitalize() + mid
    col2 = col2.lower()
---
    col1 = col1.lower() + mid #paired
    col2 = col2.capitalize()
---
    return col1 + col2
---
    return col1 + " " + col2 #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=str-mixed-colors-ac data-question_label="18">
<div id=str-mixed-colors-ac_question class="ac_question col-md-12">
<p>Write the function, <code class="docutils literal notranslate"><span class="pre">colors(col1,</span> <span class="pre">col2)</span></code>, to return a string combining the variables <code class="docutils literal notranslate"><span class="pre">col1</span></code> and <code class="docutils literal notranslate"><span class="pre">col2</span></code>.
For example, if <code class="docutils literal notranslate"><span class="pre">col1</span></code> and <code class="docutils literal notranslate"><span class="pre">col2</span></code> is “‘Red’” and “‘blue’”, then the function should return <code class="docutils literal notranslate"><span class="pre">Red</span> <span class="pre">is</span> <span class="pre">a</span> <span class="pre">primary</span> <span class="pre">color</span> <span class="pre">and</span>
<span class="pre">so</span> <span class="pre">is</span> <span class="pre">blue</span></code>.  Be sure to capitalize the first color and lowercase the second.</p>

</div>
<textarea data-lang="python" id="str-mixed-colors-ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(colors("red", "blue"), "Red is a primary color, and so is blue", "colors('red', 'blue')")
        self.assertEqual(colors("yellow", "Blue"), "Yellow is a primary color, and so is blue", "colors('yellow', 'Blue')")
        self.assertEqual(colors("Yellow", "red"), "Yellow is a primary color, and so is red", "colors('Yellow', 'red')")
        self.assertEqual(colors("Red", "Yellow"), "Red is a primary color, and so is yellow", "colors('Red', 'Yellow')")
        self.assertEqual(colors("blue", "RED"), "Blue is a primary color, and so is red", "colors('blue', 'RED')")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="str-mixed-dogs" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create the function, <code class="docutils literal notranslate"><span class="pre">fave_animal(animal)</span></code>, to return the string, “Petting <code class="docutils literal notranslate"><span class="pre">animal</span></code> makes me happier than anything else” where <code class="docutils literal notranslate"><span class="pre">animal</span></code>
is any animal of your choice. For example, if <code class="docutils literal notranslate"><span class="pre">animal</span></code> is “dogs”, then the function should return, <code class="docutils literal notranslate"><span class="pre">&quot;Petting</span> <span class="pre">dogs</span> <span class="pre">makes</span>
<span class="pre">me</span> <span class="pre">happier</span> <span class="pre">than</span> <span class="pre">anything</span> <span class="pre">else&quot;</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="19"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def fave_animal(animal):
---
    end = " makes me happier than anything else"
---
    end = "makes me happier than anything else" #paired
---
    return "Petting " + animal + end
---
    return "Petting " + "animal" + end #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=str-mixed-dogs-ac data-question_label="20">
<div id=str-mixed-dogs-ac_question class="ac_question col-md-12">
<p>Finish the function, <code class="docutils literal notranslate"><span class="pre">fave_animal(animal)</span></code>, to return the string, “Petting <code class="docutils literal notranslate"><span class="pre">animal</span></code> makes me happier than anything else” where <code class="docutils literal notranslate"><span class="pre">animal</span></code>
is any animal of your choice. For example, if <code class="docutils literal notranslate"><span class="pre">animal</span></code> is “dogs”, then the function should return, <code class="docutils literal notranslate"><span class="pre">&quot;Petting</span> <span class="pre">dogs</span> <span class="pre">makes</span>
<span class="pre">me</span> <span class="pre">happier</span> <span class="pre">than</span> <span class="pre">anything</span> <span class="pre">else&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="str-mixed-dogs-ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(fave_animal("dogs"), "Petting dogs makes me happier than anything else", "fave_animal('dogs')")
        self.assertEqual(fave_animal("ca ts"), "Petting ca ts makes me happier than anything else", 'fave_animal("ca ts")')
        self.assertEqual(fave_animal("Horses"), "Petting Horses makes me happier than anything else", 'fave_animal("Horses")')

myTests().main()
</textarea>
</div>
</div>
</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>