<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Mixed-Up Code Exercises &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/637.css?v=bddd2011b20db76a20ea" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.css?v=7c5388ad4389eae1a318" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/runtime.bundle.js?v=aa7c717c38c5405af9e6"></script>
    <script src="../_static/637.bundle.js?v=1123ad75f9b821cdd636"></script>
    <script src="../_static/runestone.bundle.js?v=9c8e466e8cbed4a3af5a"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'py4e-int';
    eBookConfig.basecourse = 'py4e-int';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          py4e-int
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=mixedupcode_orig'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
  
<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="mixed-up-code-exercises">
<h1>Mixed-Up Code Exercises<a class="headerlink" href="#mixed-up-code-exercises" title="Permalink to this headline">¶</a></h1>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="obj_mixedupcode_1" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a <code class="docutils literal notranslate"><span class="pre">Dog</span></code> class and define the <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method
that has one parameter, <code class="docutils literal notranslate"><span class="pre">name</span></code>. Then define the <code class="docutils literal notranslate"><span class="pre">getName</span></code> method, which returns the
<code class="docutils literal notranslate"><span class="pre">name</span></code> when the method is called.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="1"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class Dog:
---
class Dog #paired
---
    def __init__(self, name):
---
    def __init__(name): #paired
---
        self.name = name
---
        name = self.name #paired
---
    def getName(self):
---
        return self.name
---
        return name #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=obj_mixedupcode_1_ac data-question_label="2">
<div id=obj_mixedupcode_1_ac_question class="ac_question col-md-12">
<p>Write a <code class="docutils literal notranslate"><span class="pre">Dog</span></code> class and define the <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method that has one parameter, <code class="docutils literal notranslate"><span class="pre">name</span></code>. Then define the
<code class="docutils literal notranslate"><span class="pre">getName</span></code> method, which returns the <code class="docutils literal notranslate"><span class="pre">name</span></code> when the method is called.</p>

</div>
<textarea data-lang="python" id="obj_mixedupcode_1_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def setUp(self):
        self.Bill = Dog("Bill")
        self.Dot = Dog("Dot")
        self.Spot = Dog("Spot")
        self.Lovey = Dog("Lovey")
        self.Brownie = Dog("Brownie")

    def testOne(self):
        self.assertEqual(self.Bill.getName(), "Bill" ,"self.Bill.getName()")
        self.assertEqual(self.Dot.getName(), "Dot" ,"self.Dot.getName()")
        self.assertEqual(self.Spot.getName(), "Spot" ,"self.Spot.getName()")
        self.assertEqual(self.Lovey.getName(), "Lovey" ,"self.Lovey.getName()")
        self.assertEqual(self.Brownie.getName(), "Brownie" ,"self.Brownie.getName()")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="obj_mixedupcode_2_v2" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class named <code class="docutils literal notranslate"><span class="pre">Dog</span></code>. Define the <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method with two parameters, <code class="docutils literal notranslate"><span class="pre">name</span></code> and <code class="docutils literal notranslate"><span class="pre">age</span></code>.
Also create the method <code class="docutils literal notranslate"><span class="pre">updateAge</span></code>, that increases the <code class="docutils literal notranslate"><span class="pre">age</span></code> by 1 and returns <code class="docutils literal notranslate"><span class="pre">age</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="3"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class Dog:
---
class Dog #paired
---
    def __init__(self, name, age):
---
    def __init__(name, age): #paired
---
        self.name = name
        self.age = age
---
        name = self.name #paired
        age = self.age
---
    def updateAge(self):
---
    def updateAge(age): #paired
---
        self.age += 1
---
        return self.age
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=obj_mixedupcode_2_v2_ac data-question_label="4">
<div id=obj_mixedupcode_2_v2_ac_question class="ac_question col-md-12">
<p>Write a class named <code class="docutils literal notranslate"><span class="pre">Dog</span></code>. Define the <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method with two parameters, <code class="docutils literal notranslate"><span class="pre">name</span></code> and <code class="docutils literal notranslate"><span class="pre">age</span></code>.
Also create the method <code class="docutils literal notranslate"><span class="pre">updateAge</span></code>, that increases the <code class="docutils literal notranslate"><span class="pre">age</span></code> by 1 and returns <code class="docutils literal notranslate"><span class="pre">age</span></code>.</p>

</div>
<textarea data-lang="python" id="obj_mixedupcode_2_v2_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def setUp(self):
        self.Lark = Dog("Lark", 10)
        self.Bill = Dog("Bill", 10)
        self.Tom = Dog("Tom", 22)
        self.James = Dog("James", 18)
        self.Terry = Dog("Terry", 29)
        self.Tina = Dog("Tina", 17)

    def testOne(self):
        self.assertEqual(self.Lark.updateAge(), 11 ,"self.Lark.updateAge()")
        self.assertEqual(self.Bill.updateAge(), 11 ,"self.Bill.updateAge()")
        self.assertEqual(self.Tom.updateAge(), 23 ,"self.Tom.updateAge()")
        self.assertEqual(self.James.updateAge(), 19 ,"self.James.updateAge()")
        self.assertEqual(self.Terry.updateAge(), 30 ,"self.Terry.updateAge()")
        self.assertEqual(self.Tina.updateAge(), 18 ,"self.Tina.updateAge()")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="obj_mixedupcode_3" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create the class <code class="docutils literal notranslate"><span class="pre">Cat</span></code>. Define the <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method, it has two parameters,
<code class="docutils literal notranslate"><span class="pre">name</span></code> and <code class="docutils literal notranslate"><span class="pre">age</span></code>. Next define the the <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method to return a string with the cat’s information:
<code class="docutils literal notranslate"><span class="pre">&quot;Name:</span> <span class="pre">name,</span> <span class="pre">Age:</span> <span class="pre">age&quot;</span></code>. Then define the <code class="docutils literal notranslate"><span class="pre">make_sound</span></code> method, which should return
the string <code class="docutils literal notranslate"><span class="pre">&quot;Meow&quot;</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="5"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class Cat:
---
    def __init__(self, name, age):
---
        self.name = name
        self.age = age
---
    def __str__(self):
---
        return f"Name: {self.name}, Age: {self.age}"
---
    def make_sound(self):
---
    def make_sound(sound): #paired
---
        return "Meow"
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=obj_mixedupcode_3_ac data-question_label="6">
<div id=obj_mixedupcode_3_ac_question class="ac_question col-md-12">
<p>Write the class <code class="docutils literal notranslate"><span class="pre">Cat</span></code>. Define the <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method, it has two parameters,
<code class="docutils literal notranslate"><span class="pre">name</span></code> and <code class="docutils literal notranslate"><span class="pre">age</span></code>. Next define the the <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method to return a string with the cat’s information:
<code class="docutils literal notranslate"><span class="pre">&quot;Name:</span> <span class="pre">name,</span> <span class="pre">Age:</span> <span class="pre">age&quot;</span></code>. Then define the <code class="docutils literal notranslate"><span class="pre">make_sound</span></code> method, which should return
the string <code class="docutils literal notranslate"><span class="pre">&quot;Meow&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="obj_mixedupcode_3_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def setUp(self):
        self.Lucky = Cat("Lucky", 10)
        self.Jack = Cat("Jack", 10)
        self.Chip = Cat("Chip", 22)
        self.Tink = Cat("Tink", 18)
        self.Pinky = Cat("Pinky", 29)


    def testOne(self):
        self.assertEqual(self.Lucky.__str__(), "Name: Lucky, Age: 10" ,"self.Lucky.__str__()")
        self.assertEqual(self.Lucky.make_sound(), "Meow", 'self.Lucky.make_sound()')
        self.assertEqual(self.Jack.__str__(), "Name: Jack, Age: 10" ,"self.Jack.__str__()")
        self.assertEqual(self.Jack.make_sound(), "Meow", 'self.Jack.make_sound()')
        self.assertEqual(self.Chip.__str__(), "Name: Chip, Age: 22" ,"self.Chip.__str__()")
        self.assertEqual(self.Tink.__str__(), "Name: Tink, Age: 18" ,"self.Tink.__str__()")
        self.assertEqual(self.Pinky.__str__(), "Name: Pinky, Age: 29" ,"self.Pinky.__str__()")
myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="obj_mixedupcode_4" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class named <code class="docutils literal notranslate"><span class="pre">Book</span></code> that has an <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method with the parameters <code class="docutils literal notranslate"><span class="pre">title</span></code>
and <code class="docutils literal notranslate"><span class="pre">author</span></code>.  Then create a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns <code class="docutils literal notranslate"><span class="pre">&quot;Title:</span> <span class="pre">title,</span> <span class="pre">Author:</span> <span class="pre">author&quot;</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="7"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class Book:
---
Class Book: #paired
---
    def __init__(self, title, author):
---
        self.title = title
        self.author = author
---
        title = self.title #paired
        author = self.author
---
    def __str__(self):
---
        a = self.author
        t = self.title
        return f"Title: {t}, Author: {a}"
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=obj_mixedupcode_4_ac data-question_label="8">
<div id=obj_mixedupcode_4_ac_question class="ac_question col-md-12">
<p>Write a class named <code class="docutils literal notranslate"><span class="pre">Book</span></code> that has an <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method with two parameters, title
and author.  Then create a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns <code class="docutils literal notranslate"><span class="pre">&quot;Title:</span> <span class="pre">title,</span> <span class="pre">Author:</span> <span class="pre">author&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="obj_mixedupcode_4_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def setUp(self):
        self.newbook = Book("The Odyssey", "Homer")
        self.newbook2 = Book("The Fault in Our Stars", "John Green")


    def testOne(self):
        self.assertEqual(self.newbook.title, "The Odyssey", "Testing title")
        self.assertEqual(self.newbook.author, "Homer", "Testing author")
        self.assertEqual(self.newbook2.title, "The Fault in Our Stars", "Testing title")
        self.assertEqual(self.newbook2.author, "John Green", "Testing author")
        self.assertEqual(self.newbook.__str__(), "Title: The Odyssey, Author: Homer", "Testing the __str__ method")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="obj_mixedupcode_5_v3" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class named <code class="docutils literal notranslate"><span class="pre">Dog</span></code> with the <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method taking <code class="docutils literal notranslate"><span class="pre">name</span></code> as its parameter
and setting <code class="docutils literal notranslate"><span class="pre">self.tricks</span></code> to an empty list.  Then create a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns a string
with <code class="docutils literal notranslate"><span class="pre">&quot;Name:</span> <span class="pre">name&quot;</span></code>.  Then create a method, <code class="docutils literal notranslate"><span class="pre">updateTricks</span></code>, that adds
a new trick to the list.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="9"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class Dog:
---
    def __init__(self, name):
---
    def __init__(name): #paired
---
        self.name = name
        self.tricks = []
---
        this.name = name #paired
        this.tricks = []
---
    def __str__(self):
---
        return f'Name: {self.name}'
---
    def updateTricks(self, trick):
---
        self.tricks.append(trick)
---
        self.tricks.append(self.trick) #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=obj_mixedupcode_5_v3_ac data-question_label="10">
<div id=obj_mixedupcode_5_v3_ac_question class="ac_question col-md-12">
<p>Write a class named <code class="docutils literal notranslate"><span class="pre">Dog</span></code> with the <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method taking <code class="docutils literal notranslate"><span class="pre">name</span></code> as its parameter
and setting <code class="docutils literal notranslate"><span class="pre">self.tricks</span></code> to an empty list.   Then create a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns a string
with <code class="docutils literal notranslate"><span class="pre">&quot;Name:</span> <span class="pre">name&quot;</span></code>. Then create a method, <code class="docutils literal notranslate"><span class="pre">updateTricks</span></code>, that adds
a new trick to the list.</p>

</div>
<textarea data-lang="python" id="obj_mixedupcode_5_v3_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def setUp(self):
        self.Spot = Dog("Spot")
        self.Jack = Dog("Jack")
        self.Chip = Dog("Chip")
        self.Tink = Dog("Tink")
        self.Pinky = Dog("Pinky")


    def testOne(self):
        self.assertEqual(self.Spot.__str__(), "Name: Spot" ,"self.Spot.__str__()")
        self.Spot.updateTricks("jump")
        self.assertEqual(self.Spot.tricks, ['jump'], "Testing self.Spot.tricks after add 'jump'")
        self.Spot.updateTricks("sit")
        self.assertEqual(self.Spot.tricks, ['jump', 'sit'], "Testing self.Spot.trick after add 'sit'")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="obj_mixedupcode_6" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class named <code class="docutils literal notranslate"><span class="pre">Square</span></code> with the <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method taking the length, <code class="docutils literal notranslate"><span class="pre">len</span></code>, with a default of <code class="docutils literal notranslate"><span class="pre">10</span></code> if it is not specified.  Then create a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns a string
with <code class="docutils literal notranslate"><span class="pre">&quot;Length:</span> <span class="pre">length&quot;</span></code>.  Then create an <code class="docutils literal notranslate"><span class="pre">area</span></code> method that returns the area (<code class="docutils literal notranslate"><span class="pre">length</span></code> times <code class="docutils literal notranslate"><span class="pre">length</span></code>).</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="11"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class Square:
---
def Square: #paired
---
    def __init__(self, len = 10):
---
    def __init__(len = 10): #paired
---
        self.len = len
---
    def __str__(self):
---
        return f'Length: {self.len}'
---
        return f'Length: {len}' #paired
---
    def area(self):
---
    def __area__(self): #paired
---
        return self.len * self.len
---
        return len * len #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=obj_mixedupcode_6_ac data-question_label="12">
<div id=obj_mixedupcode_6_ac_question class="ac_question col-md-12">
<p>Write a class named <code class="docutils literal notranslate"><span class="pre">Square</span></code> with the <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method taking the length, <code class="docutils literal notranslate"><span class="pre">len</span></code>, with a default of <code class="docutils literal notranslate"><span class="pre">10</span></code> if it is not specified.  Then create a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns a string
with <code class="docutils literal notranslate"><span class="pre">&quot;Length:</span> <span class="pre">length&quot;</span></code>.  Then create an <code class="docutils literal notranslate"><span class="pre">area</span></code> method that returns the area (<code class="docutils literal notranslate"><span class="pre">length</span></code> times <code class="docutils literal notranslate"><span class="pre">length</span></code>).</p>

</div>
<textarea data-lang="python" id="obj_mixedupcode_6_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):


    def testOne(self):
        s1 = Square()
        self.assertEqual(s1.len, 10, "Testing the default length")
        self.assertEqual(s1.area(), 100, "Testing s1.area() with length 10")
        s2 = Square(5)
        self.assertEqual(s2.area(), 25, "Testing s2.area() with length 5")
        s3 = Square(8)
        self.assertEqual(s3.area(), 64, "Testing s3.area() with length 8")
        self.assertEqual(s1.__str__(), "Length: 10", "Testing s1 __str__ method")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="obj_mixedupcode_7_v2" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class named <code class="docutils literal notranslate"><span class="pre">Rectangle</span></code> with the <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method taking a <code class="docutils literal notranslate"><span class="pre">length</span></code>, and <code class="docutils literal notranslate"><span class="pre">width</span></code>.  Then create a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns a string
with <code class="docutils literal notranslate"><span class="pre">&quot;l:</span> <span class="pre">length,</span> <span class="pre">w:</span> <span class="pre">width&quot;</span></code>.  Then create an <code class="docutils literal notranslate"><span class="pre">area</span></code> method that returns the area (<code class="docutils literal notranslate"><span class="pre">length</span></code> times <code class="docutils literal notranslate"><span class="pre">width</span></code>).</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="13"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class Rectangle:
---
Class Rectangle: #paired
---
    def __init__(self, length, width):
---
    def __init__(Length, width): #paired
---
        self.length = length
        self.width = width
---
    def __str__(self):
---
        return f'l: {self.length}, w: {self.width}'
---
        return f'l: {len}, w: {width}' #paired
---
    def area(self):
---
    def __area__(self): #paired
---
        return self.length * self.width
---
        return length * width #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=obj_mixedupcode_7_v2_ac data-question_label="14">
<div id=obj_mixedupcode_7_v2_ac_question class="ac_question col-md-12">
<p>Create a class named <code class="docutils literal notranslate"><span class="pre">Rectangle</span></code> with the <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method taking a <code class="docutils literal notranslate"><span class="pre">length</span></code>, and <code class="docutils literal notranslate"><span class="pre">width</span></code>.  Then create a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns a string
with <code class="docutils literal notranslate"><span class="pre">&quot;l:</span> <span class="pre">length,</span> <span class="pre">w:</span> <span class="pre">width&quot;</span></code>.  Then create an <code class="docutils literal notranslate"><span class="pre">area</span></code> method that returns the area (<code class="docutils literal notranslate"><span class="pre">length</span></code> times <code class="docutils literal notranslate"><span class="pre">width</span></code>).</p>

</div>
<textarea data-lang="python" id="obj_mixedupcode_7_v2_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):


    def testOne(self):
        r1 = Rectangle(10,5)
        self.assertEqual(r1.length, 10, "Testing that length is 10")
        self.assertEqual(r1.width, 5, "Testing that width is 5")
        self.assertEqual(r1.area(), 50, "Testing area with a length of 10 and width of 5")
        self.assertEqual(r1.__str__(), "l: 10, w: 5", "Testing __str__")
        r1 = Rectangle(2,3)
        self.assertEqual(r1.length, 2, "Testing that length is 2")
        self.assertEqual(r1.width, 3, "Testing that width is 3")
        self.assertEqual(r1.area(), 6, "Testing area with a length of 2 and width of 3")
        self.assertEqual(r1.__str__(), "l: 2, w: 3", "Testing __str__")


myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="obj_mixedupcode_8" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Given a class <code class="docutils literal notranslate"><span class="pre">Point</span></code> with a method <code class="docutils literal notranslate"><span class="pre">distance_from_point(point)</span></code>, create a class named <code class="docutils literal notranslate"><span class="pre">Triangle</span></code> with the <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method taking three <code class="docutils literal notranslate"><span class="pre">Point</span></code> objects and creating an attribute <code class="docutils literal notranslate"><span class="pre">points</span></code>.
Then create an <code class="docutils literal notranslate"><span class="pre">base</span></code> method that returns length of the longest side.
The length of a side is the distance between two points.  For example, <code class="docutils literal notranslate"><span class="pre">t</span> <span class="pre">=</span> <span class="pre">Triangle(Point(0,0),</span> <span class="pre">Point(1,0),</span> <span class="pre">Point(0,2))</span></code> would return <code class="docutils literal notranslate"><span class="pre">2.23</span></code> from <code class="docutils literal notranslate"><span class="pre">base</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="15"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        class Triangle:
---
    def __init__(self, p1, p2, p3):
---
        self.points = [p1, p2, p3]
---
    def base(self):
---
        dist = []
---
        for i in range(len(self.points)):
---
        for i in range(len(points)): #paired
---
            curr = self.points[i]
---
            prev = self.points[i-1]
---
            d = curr.distance_from_point(prev)
---
            dist.append(d)
---
            self.dist.append(d) #paired
---
        return max(dist)
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=obj_mixedupcode_8_ac data-question_label="16">
<div id=obj_mixedupcode_8_ac_question class="ac_question col-md-12">
<p>Given a class <code class="docutils literal notranslate"><span class="pre">Point</span></code> with a method <code class="docutils literal notranslate"><span class="pre">distance_from_point(point)</span></code>, write a class named <code class="docutils literal notranslate"><span class="pre">Triangle</span></code> with the <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method taking three <code class="docutils literal notranslate"><span class="pre">Point</span></code> objects and creating an attribute <code class="docutils literal notranslate"><span class="pre">points</span></code>.
Then create an <code class="docutils literal notranslate"><span class="pre">base</span></code> method that returns length of the longest side.
The length of a side is the distance between two points.  For example, <code class="docutils literal notranslate"><span class="pre">t</span> <span class="pre">=</span> <span class="pre">Triangle(Point(0,0),</span> <span class="pre">Point(1,0),</span> <span class="pre">Point(0,2))</span></code> would return <code class="docutils literal notranslate"><span class="pre">2.23</span></code> from <code class="docutils literal notranslate"><span class="pre">base</span></code>.</p>

</div>
<textarea data-lang="python" id="obj_mixedupcode_8_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====

import math
from unittest.gui import TestCaseGui

class Point:
    def __init__(self, x, y):
        self.x = x
        self.y = y

    def __str__(self):
        return f"({self.y}, {self.y})"

    def distance_from_point(self, p):
        return math.sqrt((p.x - self.x) ** 2 + (p.y - self.y) ** 2)

class myTests(TestCaseGui):

    def testOne(self):
        t1 = Triangle(Point(0,0), Point(1,0), Point(0,2))
        self.assertAlmostEqual(t1.base(), 2.23, 1, "Testing base of 2.2")
        t1 = Triangle(Point(0,0), Point(1,0), Point(0,3))
        self.assertAlmostEqual(t1.base(), 3.16, 1, "Testing base of 3.1")
        t1 = Triangle(Point(-3,-3), Point(0,0), Point(0,3))
        self.assertAlmostEqual(t1.base(), 6.70, 1, "Testing base of 6.7")
myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="obj_mixedupcode_9" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class named <code class="docutils literal notranslate"><span class="pre">PickItem</span></code> with the <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method taking a list of items.
Then create an <code class="docutils literal notranslate"><span class="pre">pick</span></code> method that returns one of the items at random.  You can use the <code class="docutils literal notranslate"><span class="pre">random</span></code> module’s <code class="docutils literal notranslate"><span class="pre">choice</span></code> method for this.  Next create <code class="docutils literal notranslate"><span class="pre">add(item)</span></code> that adds the item to the list of items.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="17"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        import random
class PickItem:
---
    def __init__(self, item_list):
---
    def __init__(item_list): #paired
---
        self.item_list = item_list
---
    def pick(self):
---
        return random.choice(self.item_list)
---
        return random.choice(item_list) #paired
---
    def add(self,item):
---
    def add(item): #paired
---
        self.item_list.append(item)
---
        self.item_list.append(self.item) #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=obj_mixedupcode_9_ac data-question_label="18">
<div id=obj_mixedupcode_9_ac_question class="ac_question col-md-12">
<p>Create a class named <code class="docutils literal notranslate"><span class="pre">PickItem</span></code> with the <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method taking a list of items.
Then create an <code class="docutils literal notranslate"><span class="pre">pick</span></code> method that returns one of the items at random.  You can use the <code class="docutils literal notranslate"><span class="pre">random</span></code> module’s <code class="docutils literal notranslate"><span class="pre">choice</span></code> method for this.  Next create <code class="docutils literal notranslate"><span class="pre">add(item)</span></code> that adds the item to the list of items.</p>

</div>
<textarea data-lang="python" id="obj_mixedupcode_9_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        l = ['red', 'green', 'blue']
        p = PickItem(l)
        self.assertEqual(p.item_list, l, "Testing item list")
        self.assertIn(p.pick(), l, "Testing pick")
        p.add('yellow')
        self.assertIn('yellow', p.item_list, "Testing add yellow")
        l = [1, 2, 3]
        p = PickItem(l)
        self.assertEqual(p.item_list, l, "Testing item list")
        self.assertIn(p.pick(), l, "Testing pick")
        p.add(4)
        self.assertIn(4, p.item_list, "Testing add 4")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="obj_mixedupcode_10" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a class <code class="docutils literal notranslate"><span class="pre">Dice</span></code> with an <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method that takes the number of sides, <code class="docutils literal notranslate"><span class="pre">num_sides</span></code>.  Use a default of 6 if <code class="docutils literal notranslate"><span class="pre">num_sides</span></code> isn’t passed in. Also set an attribute <code class="docutils literal notranslate"><span class="pre">history</span></code> to an empty list in the <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method.
Then create a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns <code class="docutils literal notranslate"><span class="pre">&quot;Dice</span> <span class="pre">with</span> <span class="pre">num_sides</span> <span class="pre">sides&quot;</span></code> when <code class="docutils literal notranslate"><span class="pre">num_sides</span></code> is the number of sides.
Next create an <code class="docutils literal notranslate"><span class="pre">roll</span></code> method that picks a random value from one to the number of sides, adds the value to the end of the <code class="docutils literal notranslate"><span class="pre">history</span></code> list, and returns it.
You can use the <code class="docutils literal notranslate"><span class="pre">random</span></code> module’s <code class="docutils literal notranslate"><span class="pre">randint(start,end)</span></code> method to return a random number from start to end inclusvie.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="19"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        import random
---
class Dice:
---
    def __init__(self, num_sides = 6):
---
    def __init__(self, num_sides): #paired
---
        self.num_sides = num_sides
        self.history = []
---
    def __str__(self):
---
        return f"Dice with {self.num_sides} sides"
---
    def roll(self):
---
        value = random.randint(1, self.num_sides)
        self.history.append(value)
---
        return value
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=obj_mixedupcode_10_ac data-question_label="20">
<div id=obj_mixedupcode_10_ac_question class="ac_question col-md-12">
<p>Write a class <code class="docutils literal notranslate"><span class="pre">Dice</span></code> with an <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method that takes the number of sides, <code class="docutils literal notranslate"><span class="pre">num_sides</span></code>.  Use a default of 6 if <code class="docutils literal notranslate"><span class="pre">num_sides</span></code> isn’t passed in. Also set an attribute <code class="docutils literal notranslate"><span class="pre">history</span></code> to an empty list in the <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method.
Then create a <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method that returns <code class="docutils literal notranslate"><span class="pre">&quot;Dice</span> <span class="pre">with</span> <span class="pre">num_sides</span> <span class="pre">sides&quot;</span></code> when <code class="docutils literal notranslate"><span class="pre">num_sides</span></code> is the number of sides.
Next create an <code class="docutils literal notranslate"><span class="pre">roll</span></code> method that picks a random value from one to the number of sides, adds the value to the end of the <code class="docutils literal notranslate"><span class="pre">history</span></code> list, and returns it.
You can use the <code class="docutils literal notranslate"><span class="pre">random</span></code> module’s <code class="docutils literal notranslate"><span class="pre">randint(start,end)</span></code> method to return a random number from start to end inclusvie.</p>

</div>
<textarea data-lang="python" id="obj_mixedupcode_10_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        d = Dice(4)
        values = [1, 2, 3, 4]
        self.assertIn(d.roll(), values, "Testing roll value with 4 sides")
        x = d.roll()
        self.assertEqual(len(d.history), 2, "Testing history length")
        self.assertEqual(d.history[-1], x, "Testing save last roll")
        self.assertEqual(d.__str__(), "Dice with 4 sides", "Testing __str__ with 4 sides")
        d = Dice()
        values = [1, 2, 3, 4, 5, 6]
        self.assertIn(d.roll(), values, "Testing roll values with 6 sides")
        self.assertEqual(d.__str__(), "Dice with 6 sides", "Testing __str__ with 6 sides")

myTests().main()
</textarea>
</div>
</div>
</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>