<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>16.15. Write Code Exercises &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/637.css?v=bddd2011b20db76a20ea" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.css?v=7c5388ad4389eae1a318" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/runtime.bundle.js?v=aa7c717c38c5405af9e6"></script>
    <script src="../_static/637.bundle.js?v=1123ad75f9b821cdd636"></script>
    <script src="../_static/runestone.bundle.js?v=9c8e466e8cbed4a3af5a"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="16.16. Group Work - Classes, constructors, attributes, and methods" href="group-classes.html" />
    <link rel="prev" title="16.14. Mixed-Up Code Exercises" href="mixedupcode.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'py4e-int';
    eBookConfig.basecourse = 'py4e-int';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          py4e-int
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=writecode'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">16.14. </span>Mixed-Up Code Exercises" data-toggle="tooltip">
    <a href="mixedupcode.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">16.16. </span>Group Work - Classes, constructors, attributes, and methods' data-toggle="tooltip" >              
    <a href="group-classes.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
  
<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="write-code-exercises">
<h1><span class="section-number">16.15. </span>Write Code Exercises<a class="headerlink" href="#write-code-exercises" title="Permalink to this headline">¶</a></h1>
<ol class="arabic">
<li><div id='obj_writecode1' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=obj_writecode1q data-question_label="16.15.1">
<div id=obj_writecode1q_question class="ac_question col-md-12">
<p>Complete the following code to create an object <code class="docutils literal notranslate"><span class="pre">spot</span></code> from the <code class="docutils literal notranslate"><span class="pre">Dog</span></code> class
with the name “Spot” and update their tricks to include “spin” then “sit”. You can
use <code class="docutils literal notranslate"><span class="pre">print(spot)</span></code> to print the string representation of the object.</p>

</div>
<textarea data-lang="python" id="obj_writecode1q_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Dog:

    def __init__(self, name):
        self.name = name
        self.tricks = []

    def updateTricks(self, trick):
        self.tricks.append(trick)

    def __str__(self):
        return f'Dog(name = {self.name}, tricks = {str(self.tricks)})'

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(spot.tricks, ['spin', 'sit'], 'Checking that sit and spin are added to the list of tricks.')
        self.assertEqual(spot.__str__(), "Dog(name = Spot, tricks = ['spin', 'sit'])", 'Checking that a Spot prints correctly.')
myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=obj_writecode1a data-question_label="16.15.2">
<div id=obj_writecode1a_question class="ac_question col-md-12">
<p>Below the class definition, create an instance of the class Dog, assigning it to the variable <code class="docutils literal notranslate"><span class="pre">spot</span></code>.
Be sure to include a name when constructing the instance! This works similar to a function.
To update the dog’s tricks, use the instance and dot notation to call the method, with the name of the trick as the argument.
This needs to be done for each trick. Then, using the print function, you can print the dog’s tricks and the contents of the <code class="docutils literal notranslate"><span class="pre">__str__</span></code> method.</p>

</div>
<textarea data-lang="python" id="obj_writecode1a_editor" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Dog:

    def __init__(self, name):
        self.name = name
        self.tricks = []    # creates a new empty list for each dog

    def updateTricks(self, trick):
        self.tricks.append(trick)

    def __str__(self):
        return f'Dog(name = {self.name}, tricks = {str(self.tricks)})'

spot = Dog('Spot')
spot.updateTricks('spin')
spot.updateTricks('sit')
print(spot.tricks)
print(spot)

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(spot.tricks, ['spin', 'sit'], 'Checking that sit and spin are added to the list of tricks.')
        self.assertEqual(spot.__str__(), "Dog(name = Spot, tricks = ['spin', 'sit'])", 'Checking that a Spot prints correctly.')
myTests().main()
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=obj_writecode2 data-question_label="16.15.3">
<div id=obj_writecode2_question class="ac_question col-md-12">
<p>Complete the following code to include a method named <code class="docutils literal notranslate"><span class="pre">updateAttacks(attack)</span></code>
which appends the <code class="docutils literal notranslate"><span class="pre">attack</span></code> to the list of attacks.</p>

</div>
<textarea data-lang="python" id="obj_writecode2_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Pokemon:

    def __init__(self, name, kind):
        self.name = name
        self.type = kind
        self.attacks = []

    def __str__(self):
        return f"Pokemon name is {self.name} and type is {self.type}. Attacks include: {str(self.attacks)}"

bulbasaur = Pokemon('Bulbasaur', 'Grass')
bulbasaur.updateAttacks('Vine Whip')
bulbasaur.updateAttacks('Tackle')

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        pikachu = Pokemon('Pikachu','Electric')
        pikachu.updateAttacks('Quick Attack')
        self.assertEqual(pikachu.__str__(), "Pokemon name is Pikachu and type is Electric. Attacks include: ['Quick Attack']", "Testing with Pikachu and Quick Attack")

myTests().main()
</textarea>
</div>
</div>
</li>
<li><div id='obj_writecode3' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=obj_writecode3q data-question_label="16.15.4">
<div id=obj_writecode3q_question class="ac_question col-md-12">
<p>Correct the 7 errors in the following code. The program should create
a class that prints the title and author of a book.</p>

</div>
<textarea data-lang="python" id="obj_writecode3q_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
Class Book:

    def __init__(title, author):
        title = title
        author = author
    def __str__(self):
        return f"Your book is {title} by {author}"

book = new Book("The Odyssey", "Homer")
print(book)

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        gatsby = Book("The Great Gatsby", "F. Scott Fitzgerald")
        self.assertEqual(gatsby.__str__(), "Your book is The Great Gatsby by F. Scott Fitzgerald", "Testing Great Gatsby")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=obj_writecode3a data-question_label="16.15.5">
<div id=obj_writecode3a_question class="ac_question col-md-12">
<ol class="arabic simple">
<li><p><code class="docutils literal notranslate"><span class="pre">class</span></code> is a reserved word and must be lowercase.</p></li>
<li><p>class methods must include the self parameter as a reference to the current instance of the class. <code class="docutils literal notranslate"><span class="pre">def</span> <span class="pre">__init__(self,</span> <span class="pre">title,</span> <span class="pre">author):</span></code></p></li>
<li><p>Use dot notation with the self instance to access the title variable in the Book class. <code class="docutils literal notranslate"><span class="pre">self.title</span> <span class="pre">=</span> <span class="pre">title</span></code></p></li>
<li><p>Use dot notation with the self instance to access the author variable in the Book class. <code class="docutils literal notranslate"><span class="pre">self.author</span> <span class="pre">=</span> <span class="pre">author</span></code></p></li>
<li><p>Use dot notation with the self instance to access the title variable when using it in a string. <code class="docutils literal notranslate"><span class="pre">self.title</span></code></p></li>
<li><p>Use dot notation with the self instance to access the author variable when using it in a string. <code class="docutils literal notranslate"><span class="pre">self.author</span></code></p></li>
<li><p>Use the correct class name, in this case it is <code class="docutils literal notranslate"><span class="pre">Book</span></code> not <code class="docutils literal notranslate"><span class="pre">new</span> <span class="pre">Book</span></code>.</p></li>
</ol>

</div>
<textarea data-lang="python" id="obj_writecode3a_editor" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Book:

    def __init__(self, title, author):
        self.title = title
        self.author = author

    def __str__(self):
        return f"Your book is {self.title} by {self.author}"

book = Book("The Odyssey", "Homer")
print(book)

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        gatsby = Book("The Great Gatsby", "F. Scott Fitzgerald")
        self.assertEqual(gatsby.__str__(), "Your book is The Great Gatsby by F. Scott Fitzgerald", "Testing Great Gatsby")

myTests().main()
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=obj_writecode4 data-question_label="16.15.6">
<div id=obj_writecode4_question class="ac_question col-md-12">
<p>Correct all the errors in the following code. The program should create
a class that prints the name the tricks the dog knows.</p>

</div>
<textarea data-lang="python" id="obj_writecode4_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Dog

    def init(self, name):
        name = self.name
        self.tricks = []

    def updateTricks(self.trick):
        self.tricks.append = trick

    def __str__():
        return f"The new dog is named {self.name} and his tricks are {str(self.tricks)}"

dog = new Dog('Toby')
dog.updateTricks('catch')
dog.updateTrick('spin')

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(dog.tricks, ['catch', 'spin'], 'Checking that catch and spin are added to the list of tricks.')
        self.assertEqual(dog.__str__(), "The new dog is named Toby and his tricks are ['catch', 'spin']", 'Checking that a Toby prints correctly.')
myTests().main()
</textarea>
</div>
</div>
</li>
<li><div id='obj_writecode5' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=obj_writecode5q data-question_label="16.15.7">
<div id=obj_writecode5q_question class="ac_question col-md-12">
<p>Complete the following code to include a method named <code class="docutils literal notranslate"><span class="pre">getTitle</span></code> that returns
the title and a method named <code class="docutils literal notranslate"><span class="pre">getAuthor</span></code> that returns the ‘author’.</p>

</div>
<textarea data-lang="python" id="obj_writecode5q_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Book:

    def __init__(self, title, author):
        self.title = title
        self.author = author

book = Book("The Odyssey", "Homer")
print(book.getTitle())
print(book.getAuthor())

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        gatsby = Book("The Great Gatsby", "F. Scott Fitzgerald")
        self.assertEqual(gatsby.getAuthor(), "F. Scott Fitzgerald", "Testing Great Gatsby author")
        self.assertEqual(gatsby.getTitle(), "The Great Gatsby", "Testing Great Gatsby title")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=obj_writecode5a data-question_label="16.15.8">
<div id=obj_writecode5a_question class="ac_question col-md-12">
<p>Create the <code class="docutils literal notranslate"><span class="pre">getTitle</span></code> and <code class="docutils literal notranslate"><span class="pre">getAuthor</span></code> methods just as you would define a function.
Include the <code class="docutils literal notranslate"><span class="pre">self</span></code> parameter and use the self instance with dot notation to access
title and author, as needed.</p>

</div>
<textarea data-lang="python" id="obj_writecode5a_editor" 
      data-timelimit=25000     data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Book:

    def __init__(self, title, author):
        self.title = title
        self.author = author

    # Create the methods using self to access the attributes
    def getTitle(self):
        return self.title

    def getAuthor(self):
        return self.author


book = Book("The Odyssey", "Homer")
print(book.getTitle())
print(book.getAuthor())
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=obj_writecode6 data-question_label="16.15.9">
<div id=obj_writecode6_question class="ac_question col-md-12">
<p>Complete the following code to include a method named <code class="docutils literal notranslate"><span class="pre">getTricks</span></code> that returns the
tricks list and a method named <code class="docutils literal notranslate"><span class="pre">getName</span></code> that returns the name when called.</p>

</div>
<textarea data-lang="python" id="obj_writecode6_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Dog:

    def __init__(self, name):
        self.name = name
        self.tricks = []

    def updateTricks(self, trick):
        self.tricks.append(trick)

dog = Dog('Frito')
dog.updateTricks('spin')
dog.updateTricks('sit')
print(dog.getName())
print(dog.getTricks())


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        fido = Dog("Fido")
        fido.updateTricks("beg")
        self.assertEqual(fido.getName(), "Fido", "Testing a new dog with a new name")
        self.assertEqual(fido.getTricks(), ['beg'], "Testing that Fido has a new trick")


myTests().main()
</textarea>
</div>
</div>
</li>
<li><div id='obj_writecode7' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=obj_writecode7q data-question_label="16.15.10">
<div id=obj_writecode7q_question class="ac_question col-md-12">
<p>Add a new class named <code class="docutils literal notranslate"><span class="pre">Paperback</span></code> that extends the <code class="docutils literal notranslate"><span class="pre">Book</span></code> class. Add a
method named <code class="docutils literal notranslate"><span class="pre">__str__</span></code> within Paperback that sends a string representation
for the Paperback book, reading <code class="docutils literal notranslate"><span class="pre">&quot;Paperback</span> <span class="pre">book</span> <span class="pre">[TITLE]</span> <span class="pre">was</span> <span class="pre">written</span> <span class="pre">by</span> <span class="pre">[AUTHOR]&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="obj_writecode7q_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Book:

    def __init__(self, title, author):
        self.title = title
        self.author = author

    def getTitle(self):
        return self.title

    def getAuthor(self):
        return self.author

book = Paperback("The Odyssey", "Homer")
print(book)

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        gatsby = Paperback("The Great Gatsby", "F. Scott Fitzgerald")
        self.assertEqual(gatsby.__str__(), "Paperback book The Great Gatsby was written by F. Scott Fitzgerald", "Testing Great Gatsby")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=obj_writecode7a data-question_label="16.15.11">
<div id=obj_writecode7a_question class="ac_question col-md-12">
<p>Create the <code class="docutils literal notranslate"><span class="pre">Paperback</span></code> class just as you did the <code class="docutils literal notranslate"><span class="pre">Book</span></code> class, but
use the <code class="docutils literal notranslate"><span class="pre">Book</span></code> class as the parameter.
Define the <code class="docutils literal notranslate"><span class="pre">__str__</span></code> function as usual, using instances from the parent class.</p>

</div>
<textarea data-lang="python" id="obj_writecode7a_editor" 
      data-timelimit=25000     data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Book:

    def __init__(self, title, author):
        self.title = title
        self.author = author

    def getTitle(self):
        return self.title

    def getAuthor(self):
        return self.author

# Create Paperback class, using Book class
class Paperback(Book):

    def __str__(self):
        return f"Paperback book {self.title} was written by {self.author}"

book = Paperback("The Odyssey", "Homer")
print(book)
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=obj_writecode8 data-question_label="16.15.12">
<div id=obj_writecode8_question class="ac_question col-md-12">
<p>Update the new class named <code class="docutils literal notranslate"><span class="pre">WaterType</span></code> which inherits properties of the <code class="docutils literal notranslate"><span class="pre">Pokemon</span></code> class.
Add the following three methods to WaterType: <code class="docutils literal notranslate"><span class="pre">updateAttacks</span></code> appends the attacks list with
a new attack, <code class="docutils literal notranslate"><span class="pre">getName</span></code> returns the name, and <code class="docutils literal notranslate"><span class="pre">getAttacks</span></code> returns the attacks when called.</p>

</div>
<textarea data-lang="python" id="obj_writecode8_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Pokemon:

    def __init__(self, name):
        self.name = name
        self.attacks = []

class WaterType(Pokemon):


pokemon = WaterType('Gyrados')
pokemon.updateAttacks('Twister')
pokemon.updateAttacks('Whirlpool')
print(pokemon.getName())
print(pokemon.getAttacks())

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        horsea = WaterType("Horsea")
        horsea.updateAttacks("Water Gun")
        horsea.updateAttacks("Twister")
        self.assertEqual(horsea.getName(), "Horsea", "Testing a new Pokemon, Horsea")
        self.assertEqual(horsea.getAttacks(), ['Water Gun', 'Twister'], "Testing that Horsea has new tricks")

myTests().main()
</textarea>
</div>
</div>
</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=obj_writecode10 data-question_label="16.15.13">
<div id=obj_writecode10_question class="ac_question col-md-12">
<p>Add a new class named <code class="docutils literal notranslate"><span class="pre">WaterType</span></code> that inherits from Pokemon class. that takes ‘name’ as initial values, creates
an instance of Pokemon with type as ‘water’ by default and stores the instance
in a list named ‘watertypes’. Also create a method called ‘addPokemons’ which
takes ‘name’ as arguments, creates an instance of Pokemon and stores it in
‘watertypes’. Also create ‘__str__’ that returns the string representation
of the object that includes the ‘watertypes’ list.</p>

</div>
<textarea data-lang="python" id="obj_writecode10_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Pokemon:

    def __init__(self,name):
        self.name = name
        self.type = type

    def __repr__(self)
        return f"Pokemon {self.name} is of type {self.type}"

newWater = WaterType("Magikarp")
newWater.addPokemons("Krabby")
print(newWater)
</textarea>
</div>
</div>
</li>
</ol>
</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 16.14. Mixed-Up Code Exercises' data-toggle="tooltip">
  <a href="mixedupcode.html" aria-label="Previous - 16.14. Mixed-Up Code Exercises">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 16.16. Group Work - Classes, constructors, attributes, and methods' data-toggle="tooltip" >
    <a href="group-classes.html" aria-label="Next - 16.16. Group Work - Classes, constructors, attributes, and methods">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>