<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.17.1: http://docutils.sourceforge.net/" />

    <title>Mixed-Up Code Exercises &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.cd688d1623c9932e.css?v=630908D8" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.bf73008a1eba9f4e.css?v=7DED055E" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/runtime.439d581f758f2b99.bundle.js?v=2E54442"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.51d9dde4e5162190.bundle.js?v=9CAA2D33"></script>
    <script src="../_static/runestone.42b576791a4ed54b.bundle.js?v=B7293199"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    
    eBookConfig.useRunestoneServices = {{use_services}};
    eBookConfig.host = '';
    eBookConfig.app = eBookConfig.host + '/runestone';
    eBookConfig.course = '{{ course_name }}';
    eBookConfig.basecourse = '{{ base_course }}';
    eBookConfig.isLoggedIn = {{ is_logged_in}};
    eBookConfig.email = '{{ user_email }}';
    eBookConfig.isInstructor = {{ is_instructor }};
    eBookConfig.username = '{{ user_id}}';
    eBookConfig.readings = {{ readings|safe}};
    eBookConfig.activities = {{ activity_info|safe }}
    eBookConfig.downloadsEnabled = {{downloads_enabled}};
    eBookConfig.allow_pairs = {{allow_pairs}}
    eBookConfig.enableCompareMe = {{enable_compare_me}};
    // _`new_server_prefix`: Defined by the server to be the path to the new server.
    eBookConfig.new_server_prefix = '{{ new_server_prefix }}';
    
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->

    
        {% if serve_ad is defined and serve_ad and settings.adsenseid %}
         <script data-ad-client="{{settings.adsenseid}}" async src="https://pagead2.googlesyndication.com/pagead/js/adsbygoogle.js"></script>
        {% endif %}
    



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
            
            {{ course_name }}
            
          
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="browsing_warning"></span></li>

        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li id="inst_peer_link"><a href='/runestone/peer/instructor.html'>Peer Instruction (Instructor)</a></li>            
            <li><a href='/runestone/peer/student.html'>Peer Instruction (Student)</a></li>            
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/ns/books/published/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=mixedupcode_orig'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
            <script></script>
<li class="dropdown globaltoc-container">
    <a
        href="../index.html"
        class="dropdown-toggle"
        data-toggle="dropdown"
        aria-label="index"
        >This Chapter<b class="caret"></b
    ></a>
    <ul class="dropdown-menu globaltoc">
         
        {% for row in subchapter_list %}
        <li class="toctree-l1">
            <a class="reference internal" href="{{ row['subchap_uri']}}"
                >{{ row['title']}}</a
            >
        </li>
        {% endfor %} 
        
    </ul>
</li>
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->

    
        {% if settings.num_banners > 0 and settings.show_rs_banner %}
         <div style="width: 100%;">
            <a href="/runestone/default/donate?banner={{banner_num}}">
            <img class="runestone-banner"
                 src="/runestone/static/images/RunestoneBanner{{banner_num}}.png"
                 alt="Please Support Runestone"
                 style="display: block; width: 728px; margin-left: auto; margin-right: auto; margin-top: 10px;">
            </a>
         </div>
        {% endif %}
    


  
  <section id="mixed-up-code-exercises">
<h1>Mixed-Up Code Exercises<a class="headerlink" href="#mixed-up-code-exercises" title="Permalink to this headline">¶</a></h1>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="mixedupcode_pictures_pp1" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">removeGreen(img)</span></code> that sets the green in all pixels in image <code class="docutils literal notranslate"><span class="pre">img</span></code> to 0.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="1"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        from image import *
---
def removeGreen(img):
---
        pixels = img.getPixels()
---
        for p in pixels:
---
                p.setGreen(0)
---
                img.updatePixel(p)
---
        win = ImageWin(img.getWidth(), img.getHeight())
---
        win = ImageWin(img.getHeight(), img.getWidth()) #paired
---
        img.draw(win)
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=mixedupcode_pictures_ac1 data-question_label="2">
<div id=mixedupcode_pictures_ac1_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">removeGreen(img)</span></code> that sets the green in all pixels in image <code class="docutils literal notranslate"><span class="pre">img</span></code> to 0.</p>

</div>
<textarea data-lang="python" id="mixedupcode_pictures_ac1_editor" 
      data-timelimit=25000     
    data-audio=''  data-datafile='kitten.jpg'    
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
              img = Image('kitten.jpg')
                    removeGreen(img)
                    pixels = img.getPixels()
                    for p in pixels:
                              self.assertEqual(p, 0, 'Checking that the green in each pixel was set to 0.')
myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="mixedupcode_pictures_pp2" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">changeColors(img)</span></code> that sets the red in all pixels equal to the blue value in the image <code class="docutils literal notranslate"><span class="pre">img</span></code>. Then, it sets the green and blue in all pixels to 255.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="3"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        from image import *
---
def changeColors(img):
---
        pixels = img.getPixels()
---
        pixels = img.Pixels() #paired
---
        for p in pixels:
---
                b = p.getBlue()
---
                p.setRed(b)
                p.setGreen(255)
                p.setBlue(255)
---
                img.updatePixel(p)
---
        win = ImageWin(img.getWidth(), img.getHeight())
---
        win = ImageWin(img.getHeight(), img.getWidth()) #paired
---
        img.draw(win)
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=mixedupcode_pictures_ac2 data-question_label="4">
<div id=mixedupcode_pictures_ac2_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">changeColors(img)</span></code> that sets the red in all pixels equal to the blue value in the image <code class="docutils literal notranslate"><span class="pre">img</span></code>. Then, it sets the green and blue in all pixels to 255.</p>

</div>
<textarea data-lang="python" id="mixedupcode_pictures_ac2_editor" 
      data-timelimit=25000     
    data-audio=''  data-datafile='beach.jpg'    
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
        def testOne(self):
                img = Image('beach.jpg')
                pixels1 = img.getPixels()
                p = pixels1[-1]
                p2 = pixels[0]
                x1 = p.getBlue()
                x2 = p2.getBlue()

                changeColors(img)
                pixels = img.getPixels()
                p = pixels[-1]
                p2 = pixels[0]
                self.assertEqual(p.getRed(), x1, "Checking that the last pixel's red value was set to the blue value")
                self.assertEqual(p2.getRed(), x2, "Checking that the first pixel's red value was set to the blue value")
                self.assertEqual(p.getBlue(), 255, "Checking that the last blue value was set to 255")
                self.assertEqual(p.getGreen(), 255, "Checking that the last green value was set to 255")
myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="mixedupcode_pictures_pp3" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">reduceGreen(img)</span></code> that reduces the green in all pixels by 75% in image <code class="docutils literal notranslate"><span class="pre">img</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="5"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        from image import *
---
def reduceGreen(img):
---
        pixelList = img.getPixels()
---
        pixelList = img.Pixels() #paired
---
        for p in pixelList:
---
                g = p.getGreen()
---
                p.setGreen(g * 0.25)
---
                p.setGreen(g * 0.75) #paired
---
                img.updatePixel(p)
---
        win = ImageWin(img.getWidth(), img.getHeight())
---
        img.draw(win)
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=mixedupcode_pictures_ac3 data-question_label="6">
<div id=mixedupcode_pictures_ac3_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">reduceGreen(img)</span></code> that reduces the green in all pixels by 75% in image <code class="docutils literal notranslate"><span class="pre">img</span></code>.</p>

</div>
<textarea data-lang="python" id="mixedupcode_pictures_ac3_editor" 
      data-timelimit=25000     
    data-audio=''  data-datafile='kitten.jpg'    
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
        def testOne(self):
                img = Image('kitten.jpg')
                pixels1 = img.getPixels()
                p = pixels1[-1]
                x1 = p.getGreen() * 0.25

                reduceGreen(img)
                pixels = c.getPixels()
                p = pixels[-1]
                c1 = p.getGreen()
                self.assertEqual(c1, x1, "Checking that the green value in the last pixel was reduced by 75%")
myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="mixedupcode_pictures_pp4" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">changeRed(img)</span></code> that sets all red values equal to 1.5x the blue value in the image <code class="docutils literal notranslate"><span class="pre">img</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="7"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        from image import *
---
def changeRed(img):
---
        pixels = img.getPixels()
---
        for p in pixels:
---
                r = p.getRed() #paired
---
                b = p.getBlue()
---
                p.setRed(b * 1.5)
---
                r.setRed(b * 1.5) #paired
---
                img.updatePixel(p)
---
        win = ImageWin(img.getWidth(), img.getHeight())
---
        img.draw(win)
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=mixedupcode_pictures_ac4 data-question_label="8">
<div id=mixedupcode_pictures_ac4_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">changeRed</span></code> that sets all red values equal to 1.5x the blue value in the image <code class="docutils literal notranslate"><span class="pre">img</span></code>.</p>

</div>
<textarea data-lang="python" id="mixedupcode_pictures_ac4_editor" 
      data-timelimit=25000     
    data-audio=''  data-datafile='kitten.jpg'    
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
        def testOne(self):
                img = Image('kitten.jpg')
                pixels1 = img.getPixels()
                p = pixels1[-1]
                x1 = p.getBlue() * 1.5

                changeRed(img)
                pixels = img.getPixels()
                p = pixels[-1]
                c1 = p.getRed()
                self.assertEqual(c1, x1, "Checking that the red value was set to 1.5x the blue value")
myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="mixedupcode_pictures_pp_5" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">modifyColors(img)</span></code> that decreases the red to 60% of its original value, increases the blue by 60% of its original value, and sets the green to 0 in the image <code class="docutils literal notranslate"><span class="pre">img</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="9"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        from image import *
---
def modifyColors(img):
---
        pixels = img.getPixels()
---
        for p in pixels:
---
                r = p.getRed()
                b = p.getBlue()
---
                p.setRed(r * 0.6)
                p.setBlue(b * 1.6)
                p.setGreen(0)
---
                p.setRed(r / 0.6)
                p.setBlue(b * 0.6)
                p.setGreen(0) #paired
---
                img.updatePixel(p)
---
        win = ImageWin(img.getWidth(), img.getHeight())
---
        img.draw(win)
---
        img.draw() #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=mixedupcode_pictures_ac5 data-question_label="10">
<div id=mixedupcode_pictures_ac5_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">modifyColors(img)</span></code> that decreases the red to 60% of its original value, increases the blue by 60% of its original value, and sets the green to 0 in the image <code class="docutils literal notranslate"><span class="pre">img</span></code>.</p>

</div>
<textarea data-lang="python" id="mixedupcode_pictures_ac5_editor" 
      data-timelimit=25000     
    data-audio=''  data-datafile='beach.jpg'    
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
        def testOne(self):
                img = Image("beach.jpg")
                pixels1 = img.getPixels()
                p = pixels1[1]
                x1 = (p.getRed() * 0.6)
                x2 = (p.getBlue() * 1.6)

                modifyColors(img)
                pixels = img.getPixels()
                p = pixels[1]
                c1 = p.getRed()
                c2 = p.getBlue()
                c3 = p.getGreen()
                self.assertEqual(1.2, x1, "Checking that red value was reduced to 60% of original value")
                self.assertEqual(1.6, x2, "Checking that blue value increased by 60% of original value")
                self.assertEqual(c3, 0, "Checking that green value was set to 0")
myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="mixedupcode_pictures_pp6" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">imageToWhite(img)</span></code> that makes the image <code class="docutils literal notranslate"><span class="pre">img</span></code> completely white.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="11"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        from image import *
---
from picture import * #paired
---
def imageToWhite(img):
---
        pixels = img.getPixels()
        for p in pixels:
---
                p.setRed(255)
                p.setGreen(255)
                p.setBlue(255)
---
                p.setRed(0)
                p.setGreen(0)
                p.setBlue(0) #paired
---
                img.updatePixel(p)
---
        win = ImageWin(img.getWidth(), img.getHeight())
        img.draw(win)
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=mixedupcode_pictures_ac6 data-question_label="12">
<div id=mixedupcode_pictures_ac6_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">imageToWhite(img)</span></code> that makes the image <code class="docutils literal notranslate"><span class="pre">img</span></code> completely white.</p>

</div>
<textarea data-lang="python" id="mixedupcode_pictures_ac6_editor" 
      data-timelimit=25000     
    data-audio=''  data-datafile='motorcycle.jpg'    
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
        def testOne(self):
                img = Image("motorcycle.jpg")

                imageToWhite(img)
                pixels = img.getPixels()
                for p in pixels:
                    r = p.getRed()
                    b = p.getBlue()
                    g = p.getGreen()
                    self.assertEqual(r, 255, "Checking that the red value was set to 255")
                    self.assertEqual(b, 255, "Checking that the blue value was set to 255")
                    self.assertEqual(g, 255, "Checking that the green value was set to 255")
myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="mixedupcode_pictures_pp_7" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">alternatingRed(img)</span></code> that sets the red of every other pixel to 0 in the image <code class="docutils literal notranslate"><span class="pre">img</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="13"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        from image import *
---
def alternatingRed(img):
---
        for x in range(0, img.getWidth(), 2):
---
        for x in range(img.getWidth()): #paired
---
                for y in range(0, img.getHeight(), 2):
---
                for y in range(img.getHeight()): #paired
---
                        p = img.getPixel(x,y)
---
                        p.setRed(0)
---
                        img.updatePixel(p)
---
        win = ImageWin(img.getWidth(), img.getHeight())
        img.draw(win)
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=mixedupcode_pictures_ac7 data-question_label="14">
<div id=mixedupcode_pictures_ac7_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">alternatingRed(img)</span></code> that sets the red of every other pixel to 0 in the image <code class="docutils literal notranslate"><span class="pre">img</span></code>.</p>

</div>
<textarea data-lang="python" id="mixedupcode_pictures_ac7_editor" 
      data-timelimit=25000     
    data-audio=''  data-datafile='vangogh.jpg'    
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
        def testOne(self):
                img = Image("vangogh.jpg")
                pixels_orig = img.getPixels()
                r_orig1 = img.getPixel(100,100).getRed()
                r_orig2 = img.getPixel(101,100).getRed()


                alternatingRed(img)
                r1 = img.getPixel(100,100).getRed()
                r2 = img.getPixel(101,100).getRed()
                self.assertEqual(r1, 0, "Checking that the values of the even elements have changed to zero")
                self.assertNotEqual(r_orig1, r1, "Checking that the values of the even elements have changed")
                self.assertEqual(r2, 40, "Checking that the values of the odd elements have stayed the same")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="mixedupcode_pictures_pp8" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">changeQuadrantColors(img)</span></code> that only changes the color of the pixels in the bottom left quadrant of the image <code class="docutils literal notranslate"><span class="pre">img</span></code>.
The code should set the red value to the original blue value, the green value to the original red value, and the blue value to the original green value.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="15"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        from image import *
---
def changeQuadrantColors(img):
---
        halfWidth = (int) (img.getWidth() / 2)
        halfHeight = (int) (img.getHeight() / 2)
---
        for x in range(halfWidth):
---
                for y in range(halfHeight, img.getHeight()):
---
                for y in range(halfHeight): #paired
---
                        p = img.getPixel(x, y)
---
                        r = p.getRed()
                        g = p.getGreen()
                        b = p.getBlue()
---
                        newPixel = Pixel(b, r, g)
---
                        img.setPixel(x, y, newPixel)
---
        win = ImageWin(img.getWidth(), img.getHeight())
        img.draw(win)
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=mixedupcode_pictures_ac8 data-question_label="16">
<div id=mixedupcode_pictures_ac8_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">changeQuadrantColors(img)</span></code> that only changes the color of the pixels in the bottom left quadrant of the image <code class="docutils literal notranslate"><span class="pre">img</span></code>.
The code should set the red value to the original blue value, the green value to the original red value, and the blue value to the original green value.</p>

</div>
<textarea data-lang="python" id="mixedupcode_pictures_ac8_editor" 
      data-timelimit=25000     
    data-audio=''  data-datafile='kitten.jpg'    
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
        def testOne(self):
                img = Image("kitten.jpg")
                pixel_orig = img.getPixel(0,75)
                r_orig = pixel_orig.getRed()
                b_orig = pixel_orig.getBlue()
                g_orig = pixel_orig.getGreen()


                changeQuadrantColors(img)
                pixel_new = img.getPixel(0,75)
                r_new = pixel_new.getRed()
                b_new = pixel_new.getBlue()
                g_new = pixel_new.getGreen()
                self.assertEqual(r_new, b_orig, "Checking that the red values are set to original blue values")
                self.assertEqual(g_new, r_orig, "Checking that the green values are set to original red values")
                self.assertEqual(b_new, g_orig, "Checking that the blue values are set to original green values")
myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="mixedupcode_pictures_pp9" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">copyRightSide(img)</span></code> that copies the right side of the image onto the left side in the image <code class="docutils literal notranslate"><span class="pre">img</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="17"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        from image import *
---
def copyRightSide(img):
---
        halfway = (int) (img.getWidth() / 2)
---
        for x in range(halfway, img.getWidth()):
---
        for x in range(halfway): #paired
---
                for y in range(img.getHeight()):
---
                        p = img.getPixel(x, y)
---
                        img.setPixel(x - halfway, y, p)
---
                        img.setPixel(halfway + x, y, p) #paired
---
        win = ImageWin(img.getWidth(), img.getHeight())
        img.draw(win)
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=mixedupcode_pictures_ac9 data-question_label="18">
<div id=mixedupcode_pictures_ac9_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">copyRightSide(img)</span></code> that copies the right side of the image onto the left side in the image <code class="docutils literal notranslate"><span class="pre">img</span></code>.</p>

</div>
<textarea data-lang="python" id="mixedupcode_pictures_ac9_editor" 
      data-timelimit=25000     
    data-audio=''  data-datafile='motorcycle.jpg'    
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
        def testOne(self):
                img = Image("motorcycle.jpg")
                pixel_right = img.getPixel(195,50)

                copyRightSide(img)
                pixel_new = img.getPixel(97,50)
                self.assertEqual(list(pixel_right), list(pixel_new), "Checking that the right side of the image is copied onto the left side")
myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="mixedupcode_pictures_pp10" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">copyTopQuarter(img)</span></code> that copies the pixels from the top quarter of the y-axis to the bottom quarter of the y-axis in the image <code class="docutils literal notranslate"><span class="pre">img</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="19"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        from image import *
---
def copyTopQuarter(img):
---
        quarterHeight = (int) (img.getHeight() / 4)
---
        for x in range(img.getWidth()):
---
                for y in range(quarterHeight):
---
                        p = img.getPixel(x, y)
---
                        img.setPixel(x, quarterHeight * 3 + y, p)
---
        win = ImageWin(img.getWidth(), img.getHeight())
        img.draw(win)
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=mixedupcode_pictures_ac10 data-question_label="20">
<div id=mixedupcode_pictures_ac10_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">copyTopQuarter(img)</span></code> that copies the pixels from the top quarter of the y-axis to the bottom quarter of the y-axis in the image <code class="docutils literal notranslate"><span class="pre">img</span></code>.</p>

</div>
<textarea data-lang="python" id="mixedupcode_pictures_ac10_editor" 
      data-timelimit=25000     
    data-audio=''  data-datafile='vangogh.jpg'    
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
        def testOne(self):
                img = Image("vangogh.jpg")
                pixels1 = img.getPixel(0,0)
                pixels2 = img.getPixel(50, 50)
                pixels3 = img.getPixel(0,111)


                copyTopQuarter(img)
                pixels_new1 = img.getPixel(0,111)
                pixels_new2 = img.getPixel(50,50)
                self.assertEqual(list(pixels1), list(pixels_new1), "Check that the top quarter is copied to bottom quarter")
                self.assertEqual(list(pixels2), list(pixels_new2), "Checking that the pixels not in the bottom quarter stayed the same")
myTests().main()
</textarea>
</div>
</div>
</section>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>



  
    {% if settings.google_ga %}
      <script>
        var _gaq = _gaq || [];
        _gaq.push(['_setAccount', '{{ settings.google_ga }}']);
        _gaq.push(['_trackPageview']);

        (function() {
          var ga = document.createElement('script'); ga.async = true;
          ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
          var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
        })();
      </script>
    {% endif %}
  

  
    <script type="application/json" class="js-hypothesis-config">
    {
      "showHighlights": false
    }
    </script>

    <script async src="https://hypothes.is/embed.js"></script>
    <!--provide space to prevent hypothes.is sidebar from obscuring content-->
    <style>@media screen and (max-width:767px) { body { padding-right: 21px; }}</style>

  


<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>