<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.17.1: http://docutils.sourceforge.net/" />

    <title>Mixed-Up Code Questions &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.cd688d1623c9932e.css?v=630908D8" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.bf73008a1eba9f4e.css?v=7DED055E" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/runtime.439d581f758f2b99.bundle.js?v=2E54442"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.51d9dde4e5162190.bundle.js?v=9CAA2D33"></script>
    <script src="../_static/runestone.42b576791a4ed54b.bundle.js?v=B7293199"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    
    eBookConfig.useRunestoneServices = {{use_services}};
    eBookConfig.host = '';
    eBookConfig.app = eBookConfig.host + '/runestone';
    eBookConfig.course = '{{ course_name }}';
    eBookConfig.basecourse = '{{ base_course }}';
    eBookConfig.isLoggedIn = {{ is_logged_in}};
    eBookConfig.email = '{{ user_email }}';
    eBookConfig.isInstructor = {{ is_instructor }};
    eBookConfig.username = '{{ user_id}}';
    eBookConfig.readings = {{ readings|safe}};
    eBookConfig.activities = {{ activity_info|safe }}
    eBookConfig.downloadsEnabled = {{downloads_enabled}};
    eBookConfig.allow_pairs = {{allow_pairs}}
    eBookConfig.enableCompareMe = {{enable_compare_me}};
    // _`new_server_prefix`: Defined by the server to be the path to the new server.
    eBookConfig.new_server_prefix = '{{ new_server_prefix }}';
    
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->

    
        {% if serve_ad is defined and serve_ad and settings.adsenseid %}
         <script data-ad-client="{{settings.adsenseid}}" async src="https://pagead2.googlesyndication.com/pagead/js/adsbygoogle.js"></script>
        {% endif %}
    



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
            
            {{ course_name }}
            
          
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="browsing_warning"></span></li>

        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li id="inst_peer_link"><a href='/runestone/peer/instructor.html'>Peer Instruction (Instructor)</a></li>            
            <li><a href='/runestone/peer/student.html'>Peer Instruction (Student)</a></li>            
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/ns/books/published/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=mixedupcode_orig'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
            <script></script>
<li class="dropdown globaltoc-container">
    <a
        href="../index.html"
        class="dropdown-toggle"
        data-toggle="dropdown"
        aria-label="index"
        >This Chapter<b class="caret"></b
    ></a>
    <ul class="dropdown-menu globaltoc">
         
        {% for row in subchapter_list %}
        <li class="toctree-l1">
            <a class="reference internal" href="{{ row['subchap_uri']}}"
                >{{ row['title']}}</a
            >
        </li>
        {% endfor %} 
        
    </ul>
</li>
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->

    
        {% if settings.num_banners > 0 and settings.show_rs_banner %}
         <div style="width: 100%;">
            <a href="/runestone/default/donate?banner={{banner_num}}">
            <img class="runestone-banner"
                 src="/runestone/static/images/RunestoneBanner{{banner_num}}.png"
                 alt="Please Support Runestone"
                 style="display: block; width: 728px; margin-left: auto; margin-right: auto; margin-top: 10px;">
            </a>
         </div>
        {% endif %}
    


  
  <section id="mixed-up-code-questions">
<h1>Mixed-Up Code Questions<a class="headerlink" href="#mixed-up-code-questions" title="Permalink to this headline">¶</a></h1>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="re_mixed1" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">greeting_search</span></code> that takes in a list of strings, <code class="docutils literal notranslate"><span class="pre">lst</span></code>, as a parameter and returns <code class="docutils literal notranslate"><span class="pre">True</span></code> if the list
has “Good Morning!” in it and <code class="docutils literal notranslate"><span class="pre">False</span></code> otherwise. For example, <code class="docutils literal notranslate"><span class="pre">greeting_search([&quot;Hello!&quot;,</span> <span class="pre">&quot;Good</span> <span class="pre">Morning!&quot;,</span> <span class="pre">&quot;good</span> <span class="pre">morning!&quot;])</span></code>
should return <code class="docutils literal notranslate"><span class="pre">True</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="1"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        import re
---
def greeting_search(lst):
---
    for item in lst:
---
        if re.search(item, "Good Morning!"):
---
        if re.search(lst, "Good Morning!"): #paired
---
            return True
---
    return False
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=re_mixed1_ac data-question_label="2">
<div id=re_mixed1_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">greeting_search</span></code> that takes in a list of strings, <code class="docutils literal notranslate"><span class="pre">lst</span></code>, as a parameter and returns <code class="docutils literal notranslate"><span class="pre">True</span></code> if the list
has “Good Morning!” in it and <code class="docutils literal notranslate"><span class="pre">False</span></code> otherwise. For example, <code class="docutils literal notranslate"><span class="pre">greeting_search([&quot;Hello!&quot;,</span> <span class="pre">&quot;Good</span> <span class="pre">Morning!&quot;,</span> <span class="pre">&quot;good</span> <span class="pre">morning!&quot;])</span></code>
should return <code class="docutils literal notranslate"><span class="pre">True</span></code>.</p>

</div>
<textarea data-lang="python" id="re_mixed1_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
import re

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(greeting_search(['Hello!', 'Good Morning!', 'good morning!']), True, "greeting_search(['Hello!', 'Good Morning!', 'good morning!'])")
        self.assertEqual(greeting_search(['Hello!', 'good morning!']), False, "greeting_search(['Hello!', 'good morning!'])")
        self.assertEqual(greeting_search(['Good Morning!', 'Hello!', 'good morning!']), True, "greeting_search(['Good Morning!', 'Hello!', 'good morning!'])")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="re_mixed2" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">search_h</span></code> that takes in a list of strings, <code class="docutils literal notranslate"><span class="pre">lst</span></code>, and uses regular expressions to add all the words
starting with an “H” or an “h” to a new list and returns that new list. For example, <code class="docutils literal notranslate"><span class="pre">search_h([&quot;Hello!&quot;,</span> <span class="pre">&quot;hello.&quot;,</span> <span class="pre">&quot;Morning!&quot;,</span> <span class="pre">&quot;hi&quot;])</span></code>
should return <code class="docutils literal notranslate"><span class="pre">['Hello!',</span> <span class="pre">'hello.',</span> <span class="pre">'hi']</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="3"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        import re
---
def search_h(lst):
---
    new_list = []
---
    for item in lst:
---
        if re.search('^[Hh]', item):
---
        if re('^H,h'), item): #paired
---
            new_list.append(item)
---
    return new_list
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=re_mixed2_ac data-question_label="4">
<div id=re_mixed2_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">search_h</span></code> that takes in a list of strings, <code class="docutils literal notranslate"><span class="pre">lst</span></code>, and uses regular expressions to add all the words
starting with an “H” or an “h” to a new list and returns that new list. For example, <code class="docutils literal notranslate"><span class="pre">search_h([&quot;Hello!&quot;,</span> <span class="pre">&quot;hello.&quot;,</span> <span class="pre">&quot;Morning!&quot;,</span> <span class="pre">&quot;hi&quot;])</span></code>
should return <code class="docutils literal notranslate"><span class="pre">['Hello!',</span> <span class="pre">'hello.',</span> <span class="pre">'hi']</span></code>.</p>

</div>
<textarea data-lang="python" id="re_mixed2_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
import re

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(search_h(['Hello!', 'hello.', 'Morning!', 'hi']), ['Hello!', 'hello.', 'hi'], "search_h(['Hello!', 'hello.', 'Morning!', 'hi'])")
        self.assertEqual(search_h(['Morning!']), [], "search_h(['Morning!'])")
        self.assertEqual(search_h(['fresh', 'Kahoot']), [], "search_h(['fresh', 'Kahoot'])")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="re_mixed3" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">search_o</span></code> that takes in a list of strings, <code class="docutils literal notranslate"><span class="pre">lst</span></code>, and uses regular expressions to add all the words
that have a lowercase “o” to a new list and returns that new list. For example, <code class="docutils literal notranslate"><span class="pre">search_o(['Hello',</span> <span class="pre">'Good</span> <span class="pre">Morning!',</span> <span class="pre">'hi'])</span></code> should return
<code class="docutils literal notranslate"><span class="pre">['Hello',</span> <span class="pre">'Good</span> <span class="pre">Morning!']</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="5"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        import re
---
def search_o(lst):
---
    new_list = []
---
    for item in lst:
---
        if re.search('o', item):
---
        if re.search('^o', item): #paired
---
            new_list.append(item)
---
    return new_list
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=re_mixed3_ac data-question_label="6">
<div id=re_mixed3_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">search_o</span></code> that takes in a list of strings, <code class="docutils literal notranslate"><span class="pre">lst</span></code>, and uses regular expressions to add all the words
that have a lowercase “o” to a new list and returns that new list. For example, <code class="docutils literal notranslate"><span class="pre">search_o(['Hello',</span> <span class="pre">'Good</span> <span class="pre">Morning!',</span> <span class="pre">'hi'])</span></code> should return
<code class="docutils literal notranslate"><span class="pre">['Hello',</span> <span class="pre">'Good</span> <span class="pre">Morning!']</span></code>.</p>

</div>
<textarea data-lang="python" id="re_mixed3_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
import re

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(search_o(['Hello', 'Good Morning!', 'hi']), ['Hello', 'Good Morning!'], "search_o(['Hello', 'Good Morning!', 'hi'])")
        self.assertEqual(search_o(['Other', 'open', 'Petroleum']), ['open', 'Petroleum'], "search_o(['Other', 'open', 'Petroleum'])")
        self.assertEqual(search_o(['hi', 'bye']), [], "search_o(['hi', 'bye'])")


myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="re_mixed4_v2" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">match_word(word,</span> <span class="pre">str)</span></code> that returns True if the string <code class="docutils literal notranslate"><span class="pre">str</span></code> contains the word <code class="docutils literal notranslate"><span class="pre">word</span></code>, but not if it is part of another word.  For example, <code class="docutils literal notranslate"><span class="pre">match_word('is',</span> <span class="pre">&quot;This</span> <span class="pre">was</span> <span class="pre">bad&quot;)</span></code> would return False and <code class="docutils literal notranslate"><span class="pre">match_word('is',</span> <span class="pre">&quot;This</span> <span class="pre">is</span> <span class="pre">good&quot;)</span></code> would return True.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="7"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        import re
---
import regex #paired
---
def match_word(word, str):
---
    if re.search(r"\b" + word + r"\b", str):
---
    if re.search("\b" + word + "\b", str): #paired
---
        return True
---
    else:
---
        return False
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=re_mixed4_v2_ac data-question_label="8">
<div id=re_mixed4_v2_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">match_word(word,</span> <span class="pre">str)</span></code> that returns True if the string <code class="docutils literal notranslate"><span class="pre">str</span></code> contains the word <code class="docutils literal notranslate"><span class="pre">word</span></code>, but not if it is part of another word.  For example, <code class="docutils literal notranslate"><span class="pre">match_word('is',</span> <span class="pre">&quot;This</span> <span class="pre">was</span> <span class="pre">bad&quot;)</span></code> would return False and <code class="docutils literal notranslate"><span class="pre">match_word('is',</span> <span class="pre">&quot;This</span> <span class="pre">is</span> <span class="pre">good&quot;)</span></code> would return True.</p>

</div>
<textarea data-lang="python" id="re_mixed4_v2_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
import re


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(match_word('is','This was bad'),False,"match_word('is','This was bad')")
        self.assertEqual(match_word('is','This is good'),True,"match_word('is','This is good')")
        self.assertEqual(match_word('good','This is good'),True,"match_word('good','This is good')")
        self.assertEqual(match_word('go','This is good'),False,"match_word('go','This is good')")
        self.assertEqual(match_word('This','This is good'),True,"match_word('This','This is good')")


myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="re_mixed5_v2" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">test_pattern</span></code> that takes in one string parameter, <code class="docutils literal notranslate"><span class="pre">sequence</span></code>, and uses regular expressions to see if <code class="docutils literal notranslate"><span class="pre">sequence</span></code> has the pattern
of at least one letter, number, or underscore, at least one space, and at least one letter, number, or underscore again. Return <code class="docutils literal notranslate"><span class="pre">&quot;Match!&quot;</span></code> if <code class="docutils literal notranslate"><span class="pre">sequence</span></code> matches the pattern, and
<code class="docutils literal notranslate"><span class="pre">&quot;Not</span> <span class="pre">a</span> <span class="pre">match!&quot;</span></code> otherwise. For example, <code class="docutils literal notranslate"><span class="pre">test_pattern('Sincerely1</span>&#160;&#160;&#160;&#160;&#160;&#160; <span class="pre">Molly')</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;Match!&quot;</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="9"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        import re
---
def test_pattern(sequence):
---
    pattern = "\w+\s+\w+"
---
    pattern = "\w+\S+\w+" #paired
---
    if re.search(pattern, sequence):
---
        return "Match!"
---
    else:
---
        return "Not a match!"
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=re_mixed5_ac data-question_label="10">
<div id=re_mixed5_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">test_pattern</span></code> that takes in one string parameter, <code class="docutils literal notranslate"><span class="pre">sequence</span></code>, and uses regular expressions to see if <code class="docutils literal notranslate"><span class="pre">sequence</span></code> has the pattern
of at least one letter, number, or underscore, at least one space, and at least one letter, number, or underscore again. Return <code class="docutils literal notranslate"><span class="pre">&quot;Match!&quot;</span></code> if <code class="docutils literal notranslate"><span class="pre">sequence</span></code> matches the pattern, and
<code class="docutils literal notranslate"><span class="pre">&quot;Not</span> <span class="pre">a</span> <span class="pre">match!&quot;</span></code> otherwise. For example, <code class="docutils literal notranslate"><span class="pre">test_pattern('Sincerely1</span>&#160;&#160;&#160;&#160;&#160;&#160; <span class="pre">Molly')</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;Match!&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="re_mixed5_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
import re

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(test_pattern('Sincerely       Molly'), 'Match!', "test_pattern('Sincerely       Molly')")
        self.assertEqual(test_pattern('Sincerely1       Molly'), 'Match!', "test_pattern('Sincerely1       Molly')")
        self.assertEqual(test_pattern('Sincerely   Molly'), 'Match!', "test_pattern('Sincerely   Molly')")
        self.assertEqual(test_pattern('1 M'), 'Match!', "test_pattern('1 M')")
        self.assertEqual(test_pattern('1 1'), 'Match!', "test_pattern('1 1')")
        self.assertEqual(test_pattern('S M'), 'Match!', "test_pattern('S M')")
        self.assertEqual(test_pattern('1S 1M'), 'Match!', "test_pattern('1S 1M')")
        self.assertEqual(test_pattern('_ M_9'), 'Match!', "test_pattern('_ M_9')")
        self.assertEqual(test_pattern('SM'), 'Not a match!', "test_pattern('SM')")
        self.assertEqual(test_pattern(' SM '),'Not a match!', "test_pattern(' SM ')")
        self.assertEqual(test_pattern(' SM'), 'Not a match!', "test_pattern(' SM')")
        self.assertEqual(test_pattern('S, M'), 'Not a match!', "test_pattern('S, M')")
        self.assertEqual(test_pattern('SM '), 'Not a match!', "test_pattern('SM ')")


myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="re_mixed6_v2" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">first_price</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">string</span></code> and uses regular expressions to find and return the first price in
the <code class="docutils literal notranslate"><span class="pre">string</span></code> if there is one. A price must have at least one digit after a ‘$’ symbol and can optionally have a period followed by two digits.
If there isn’t a price in the <code class="docutils literal notranslate"><span class="pre">string</span></code>, return <code class="docutils literal notranslate"><span class="pre">'No</span> <span class="pre">price'</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">first_price('We</span> <span class="pre">just</span> <span class="pre">received</span> <span class="pre">$2098.10</span> <span class="pre">for</span> <span class="pre">cookies</span> <span class="pre">and</span> <span class="pre">$209</span> <span class="pre">for</span> <span class="pre">brownies.')</span></code>
should return <code class="docutils literal notranslate"><span class="pre">&quot;$2098.10&quot;</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="11"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        import re
---
def first_price(string):
    out = []
---
    pattern = '\$\d+(?:\.\d{2})?'
---
    pattern = '$\d+(?:\.\d{2})?' #paired
---
    out = re.findall(pattern,string)
---
    if len(out) > 0:
---
        return out[0]
---
        return out #paired
---
    else:
---
        return 'No price'
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=re_mixed6_ac data-question_label="12">
<div id=re_mixed6_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">first_price</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">string</span></code> and uses regular expressions to find and return the first price in
the <code class="docutils literal notranslate"><span class="pre">string</span></code> if there is one. A price must have at least one digit after a ‘$’ symbol and can optionally have a period followed by two digits.
If there isn’t a price in the <code class="docutils literal notranslate"><span class="pre">string</span></code>, return <code class="docutils literal notranslate"><span class="pre">'No</span> <span class="pre">price'</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">first_price('We</span> <span class="pre">just</span> <span class="pre">received</span> <span class="pre">$2098.10</span> <span class="pre">for</span> <span class="pre">cookies</span> <span class="pre">and</span> <span class="pre">$209</span> <span class="pre">for</span> <span class="pre">brownies.')</span></code>
should return <code class="docutils literal notranslate"><span class="pre">&quot;$2098.10&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="re_mixed6_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
import re

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(first_price('We just received $10.00 for cookies.'), '$10.00', "first_price('We just received $10.00 for cookies.')")
        self.assertEqual(first_price('We just received $2098.10 for cookies and $209 for brownies.'), '$2098.10', "first_price('We just received $2098.10 for cookies and $209 for brownies.')")
        self.assertEqual(first_price('We just received $5.00 for cookies.'), '$5.00', "first_price('We just received $5.00 for cookies.')")
        self.assertEqual(first_price('$5.35 for cookies'), '$5.35', "first_price('$5.35 for cookies')")
        self.assertEqual(first_price('cookies for $5.05'), '$5.05', "first_price('cookies for $5.05')")
        self.assertEqual(first_price('cookies for $5'), '$5', "first_price('cookies for $5')")
        self.assertEqual(first_price('cookies'), 'No price', "first_price('cookies')")


myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="re_mixed7_v2" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">search_email</span></code> that takes in a string parameter, <code class="docutils literal notranslate"><span class="pre">string</span></code>, and uses regular expressions to return
a list with all the emails in the string. The format for the email is at least one letter, number or underscore, the “&#64;” symbol,
and the email domain (which includes a period). For example, <code class="docutils literal notranslate"><span class="pre">search_email('His</span> <span class="pre">email</span> <span class="pre">is</span> <span class="pre">pyth_on&#64;umich.edu</span> <span class="pre">and</span> <span class="pre">her</span> <span class="pre">email</span> <span class="pre">is</span> <span class="pre">java&#64;css.')</span></code> should return
<code class="docutils literal notranslate"><span class="pre">['pyth_on&#64;umich.edu']</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="13"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        import re
---
def search_email(string):
---
    return re.findall('\w+@\w+\.\w+', string)
---
    return re.findall('\S+@S+', string) #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=re_mixed7_ac data-question_label="14">
<div id=re_mixed7_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">search_email</span></code> that takes in a string parameter, <code class="docutils literal notranslate"><span class="pre">string</span></code>, and uses regular expressions to return
a list with all the emails in the string. The format for the email is at least one letter, number or underscore, the “&#64;” symbol,
and the email domain (which includes a period). For example, <code class="docutils literal notranslate"><span class="pre">search_email('His</span> <span class="pre">email</span> <span class="pre">is</span> <span class="pre">pyth_on&#64;umich.edu</span> <span class="pre">and</span> <span class="pre">her</span> <span class="pre">email</span> <span class="pre">is</span> <span class="pre">java&#64;css.')</span></code> should return
<code class="docutils literal notranslate"><span class="pre">['pyth_on&#64;umich.edu']</span></code>.</p>

</div>
<textarea data-lang="python" id="re_mixed7_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
import re

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(search_email('python@umich.edu'), ['python@umich.edu'], "search_email('python@umich.edu')")
        self.assertEqual(search_email('python@umich.edu, java@umich.edu'), ['python@umich.edu', 'java@umich.edu'], "search_email('python@umich.edu, java@umich.edu')")
        self.assertEqual(search_email('His email is pyth_on@umich.edu and her email is java@css.'),['pyth_on@umich.edu'], "search_email('His email is pyth_on@umich.edu and her email is java@css.')")
        self.assertEqual(search_email('His email is python@umich.edu and her email is jav12a@css.com.'), ['python@umich.edu', 'jav12a@css.com'], "search_email('His email is python@umich.edu and her email is jav12a@css.com.')")
        self.assertEqual(search_email('python'), [], "search_email('python')")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="re_mixed8" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">start_from`</span></code> that takes in a list of strings, <code class="docutils literal notranslate"><span class="pre">lst</span></code>, and adds into a new list each item that starts
with “From:” and is followed by one or more characters and an “&#64;” sign. For example, <code class="docutils literal notranslate"><span class="pre">start_from(['From:</span> <span class="pre">Kelly&#64;umich.edu','From:</span> <span class="pre">Kelly&#64;',':</span> <span class="pre">Kelly&#64;'])</span></code>
should return <code class="docutils literal notranslate"><span class="pre">['From:</span> <span class="pre">Kelly&#64;umich.edu',</span> <span class="pre">'From:</span> <span class="pre">Kelly&#64;']</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="15"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        import re
---
def start_from(lst):
---
    new_list = []
---
    for item in lst:
---
        if re.search('^From:.+@', item):
---
        if re.search('^F..m:', item): #paired
---
            new_list.append(item)
---
    return new_list
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=re_mixed8_ac data-question_label="16">
<div id=re_mixed8_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">start_from`</span></code> that takes in a list of strings, <code class="docutils literal notranslate"><span class="pre">lst</span></code>, and adds into a new list each item that starts
with “From:” and is followed by one or more characters and an “&#64;” sign. For example, <code class="docutils literal notranslate"><span class="pre">start_from(['From:</span> <span class="pre">Kelly&#64;umich.edu','From:</span> <span class="pre">Kelly&#64;',':</span> <span class="pre">Kelly&#64;'])</span></code>
should return <code class="docutils literal notranslate"><span class="pre">['From:</span> <span class="pre">Kelly&#64;umich.edu',</span> <span class="pre">'From:</span> <span class="pre">Kelly&#64;']</span></code>.</p>

</div>
<textarea data-lang="python" id="re_mixed8_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
import re

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(start_from(['From: Kelly@umich.edu']), ['From: Kelly@umich.edu'], "start_from(['From: Kelly@umich.edu'])")
        self.assertEqual(start_from(['From: Kelly@umich.edu','From: Kelly@',': Kelly@']), ['From: Kelly@umich.edu', 'From: Kelly@'], "start_from(['From: Kelly@umich.edu','From: Kelly@',': Kelly@'])")
        self.assertEqual(start_from(['From: Kelly','From Kelly','rom Kelly']),[], "start_from(['From: Kelly','From Kelly','rom Kelly'])")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="re_mixed9_v2" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">grab_domain</span></code> that takes in a list of strings, <code class="docutils literal notranslate"><span class="pre">lst</span></code>, as a parameter and returns a new list with the domains of the emails, if they exist.
For example, <code class="docutils literal notranslate"><span class="pre">grab_domain(['python&#64;umich.edu','This</span> <span class="pre">is</span> <span class="pre">Kelly&#64;umich.gov','123','java&#64;css','jav12a&#64;css.com'])</span></code> would return <code class="docutils literal notranslate"><span class="pre">['umich.edu',</span> <span class="pre">'umich.gov',</span> <span class="pre">'css.com']</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="17"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        import re
---
def grab_domain(lst):
---
    out = []
---
    for str in lst:
---
        curr = re.findall('\w+@(\w+\.\w+)', str)
---
        curr = re.findall('\w+@(\w+)', str) #paired
---
        out.extend(curr)
---
        out.append(curr) #paired
---
    return out
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=re_mixed9_ac data-question_label="18">
<div id=re_mixed9_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">grab_domain</span></code> that takes in a list of strings, <code class="docutils literal notranslate"><span class="pre">lst</span></code>, as a parameter and returns a new list with the domains of the emails, if they exist.
For example, <code class="docutils literal notranslate"><span class="pre">grab_domain(['python&#64;umich.edu','This</span> <span class="pre">is</span> <span class="pre">Kelly&#64;umich.gov','123','java&#64;css','jav12a&#64;css.com'])</span></code> would return <code class="docutils literal notranslate"><span class="pre">['umich.edu',</span> <span class="pre">'umich.gov',</span> <span class="pre">'css.com']</span></code>.</p>

</div>
<textarea data-lang="python" id="re_mixed9_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
import re

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(grab_domain(['python@umich.edu','This is Kelly@umich.gov','123','java@css','jav12a@css.com']), ['umich.edu', 'umich.gov', 'css.com'], "grab_domain(['python@umich.edu','This is Kelly@umich.gov','123','java@css','jav12a@css.com'])")
        self.assertEqual(grab_domain(['123','java@css']), [], "grab_domain(['123','java@css'])")
        self.assertEqual(grab_domain(['123@google.com, be@facebook.com']), ['google.com', 'facebook.com'], "grab_domain(['123@google.com, be@facebook.com'])")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="re_mixed10_v2" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">vowels_in_mid</span></code> that takes in one string parameter, <code class="docutils literal notranslate"><span class="pre">string</span></code>, and returns <code class="docutils literal notranslate"><span class="pre">string</span></code> if there is an area in the string
that contains two to four vowels, and it does not start nor end with a vowel. Otherwise, return <code class="docutils literal notranslate"><span class="pre">'Does</span> <span class="pre">not</span> <span class="pre">exist'</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">vowels_in_mid('chEEEEEYYEErry')</span></code>
should return <code class="docutils literal notranslate"><span class="pre">'chEEEEEYYEErry'</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="19"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        import re
---
def vowels_in_mid(string):
---
    p = "[^AEIOU]+[AEIOU]{2,4}[^AEIOU]+"
---
    p = ".+[AEIOU]{2,4}[^AEIOU]+$" #paired
---
    if re.search(p, string):
---
        return string
---
    else:
---
        return 'Does not exist'
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=re_mixed10_ac data-question_label="20">
<div id=re_mixed10_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">vowels_in_mid</span></code> that takes in one string parameter, <code class="docutils literal notranslate"><span class="pre">string</span></code>, and returns <code class="docutils literal notranslate"><span class="pre">string</span></code> if there is an area in the string
that contains two to four vowels, and it does not start nor end with a vowel. Otherwise, return <code class="docutils literal notranslate"><span class="pre">'Does</span> <span class="pre">not</span> <span class="pre">exist'</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">vowels_in_mid('chEEEEEYYEErry')</span></code>
should return <code class="docutils literal notranslate"><span class="pre">'chEEEEEYYEErry'</span></code>.</p>

</div>
<textarea data-lang="python" id="re_mixed10_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
import re

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(vowels_in_mid('chErry'), 'Does not exist', "vowels_in_mid('chErry')")
        self.assertEqual(vowels_in_mid('chEErry'), 'chEErry', "vowels_in_mid('chEErry')")
        self.assertEqual(vowels_in_mid('chEEErry'), 'chEEErry', "vowels_in_mid('chEEErry')")
        self.assertEqual(vowels_in_mid('chEEEErry'), 'chEEEErry', "vowels_in_mid('chEEEErry')")
        self.assertEqual(vowels_in_mid('chEEEEErry'), 'Does not exist', "vowels_in_mid('chEEEEErry')")
        self.assertEqual(vowels_in_mid('chEEEEEYYEErry'), 'chEEEEEYYEErry', "vowels_in_mid('chEEEEEYYEErry')")

myTests().main()
</textarea>
</div>
</div>
</section>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>



  
    {% if settings.google_ga %}
      <script>
        var _gaq = _gaq || [];
        _gaq.push(['_setAccount', '{{ settings.google_ga }}']);
        _gaq.push(['_trackPageview']);

        (function() {
          var ga = document.createElement('script'); ga.async = true;
          ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
          var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
        })();
      </script>
    {% endif %}
  

  
    <script type="application/json" class="js-hypothesis-config">
    {
      "showHighlights": false
    }
    </script>

    <script async src="https://hypothes.is/embed.js"></script>
    <!--provide space to prevent hypothes.is sidebar from obscuring content-->
    <style>@media screen and (max-width:767px) { body { padding-right: 21px; }}</style>

  


<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>