<?xml version="1.0"?>
<section xml:id="inheritance_unit-tests">
  <title>Unit Tests</title>
  <introduction>
  <p>Now that you have learned about classes and inheritance we can explain
            how to write your own unit tests.</p>
        </introduction>
  <subsection xml:id="inheritance_testing-your-code">
    <title>Testing your Code</title>
    <p>How do you test that your program worked?  You may put in print statements about what you expect
                the output to be and compare that to what you actually get to see if your code is correct.</p>
    <p>In fact in test-driven development you should write test cases <term>before</term> you
                start writing any code. That way you will know when your code works.</p>
    <exercise label="ut_example_test_total_even">
      <statement>
        <p>Run the code below to see what it prints.</p>
      </statement>
      <program xml:id="ut_example_test_total_even_editor" interactive="activecode" language="python">
        <input>
# Return the total of the even numbers in a list of numbers called nums
def total_even(nums):
    total = 0
    for num in nums:
        if num % 2 == 0:
            total += num
    return total

# Test cases for total_even
def test():
    print(f"Testing total_even([1, 2]) should return 2 and returns {total_even([1, 2])}")
    print(f"Testing total_even([1, 3]) should return 0 and returns {total_even([1, 3])}")
    print(f"Testing total_even([2, 6]) should return 8 and returns {total_even([2, 6])}")
    print(f"Testing total_even([2, 3, 4]) should return 6 and returns {total_even([2, 3, 4])}")
    print(f"Testing total_even([]) should return 0 and returns {total_even([])}")
    print(f"Testing total_even([1]) should return 0 and returns {total_even([1])}")
    print(f"Testing total_even([2]) should return 2 and returns {total_even([2])}")

# function call
test()
        </input>
      </program>
    </exercise>
    <p>This type of test is better than just printing out actual output because
                then you have to think about what the correct value should be.  It lets you
                quickly compare the expected value and the actual value.  However, you have
                to visually compare the two, which takes some time and you could miss something that doesn't match.</p>
    <note>
      <p>The tests above are using a f-string which is an easier way to create a string
                    with expression results in the string.  Just add <c>f</c> before the start of the string and <c>{expression}</c> in the string.</p>
    </note>
    <p xml:id="inheritance_index-1">Another way to write the code above would be to use a <term>list comprehension</term>.  These are a quick way to return
                a new list by applying some logic to each element of the list.</p>
    <pre>result_list = [returned_value for item in list logic]</pre>
    <exercise label="ut_example_test_total_even_listc">
      <statement>
        <p>Run the code below to see what it prints.</p>
      </statement>
      <program xml:id="ut_example_test_total_even_listc_editor" interactive="activecode" language="python">
        <input>
# Return the total of the even numbers in a list of numbers called nums
def total_even(nums):
    evens = [num for num in nums if num % 2 == 0]
    return sum(evens)

# Test cases for total_even
def test():
    print(f"Testing total_even([1, 2]) should return 2 and returns {total_even([1, 2])}")
    print(f"Testing total_even([1, 3]) should return 0 and returns {total_even([1, 3])}")
    print(f"Testing total_even([2, 6]) should return 8 and returns {total_even([2, 6])}")
    print(f"Testing total_even([2, 3, 4]) should return 6 and returns {total_even([2, 3, 4])}")
    print(f"Testing total_even([]) should return 0 and returns {total_even([])}")
    print(f"Testing total_even([1]) should return 0 and returns {total_even([1])}")
    print(f"Testing total_even([2]) should return 2 and returns {total_even([2])}")

# function call
test()
        </input>
      </program>
    </exercise>
  </subsection>
  <subsection xml:id="inheritance_creating-test-cases">
    <title>Creating Test Cases</title>
    <p>What types of tests should you create?  You should create tests that check
                that the function works for a variety of normal input.  You should also think about how
                to handle unusual input such as checking that <c>total_even</c> works even for an empty list or
                a list with only one item.</p>
  </subsection>
  <subsection xml:id="inheritance_understanding-unit-tests">
    <title>Understanding Unit Tests</title>
    <p xml:id="inheritance_index-2">Another way to test your code is to write unit tests. You have already been
                using hidden unit tests in this ebook to check your code as shown below.  You
                probably quickly realized that you wanted all the tests to <q>pass</q> which is
                highlighted in green. This gives you an even quicker visual indication that your code
                is working than having to compare the expected output to the actual.</p>
    <p>Let's look at the unit tests that have usually been hidden.</p>
    <exercise label="ut-ex-perimeterq-with-unit-test-vis">
      <statement>
        <p>Run the code below to see the unit tests again. See the unit tests after the code. The last unit test compares the result of <c>recPerimeter(3.0, 2)</c> to wrong value.  Fix the value so that all of the tests pass.</p>
      </statement>
      <program xml:id="ut-ex-perimeterq-with-unit-test-vis_editor" interactive="activecode" language="python">
        <input>
def recPerimeter(length, width):
    perimeter = 2 * (length + width)
    return perimeter

from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(recPerimeter(10, 20),60,"Tested recPerimeter on inputs 10 and 20")
        self.assertEqual(recPerimeter(1, 2),6,"Tested recPerimeter on inputs 1 and 2")
        self.assertEqual(recPerimeter(23.2, 12.3),71,"Tested recPerimeter on inputs 23.2 and 12.3")
        self.assertEqual(recPerimeter(3.0, 2),12,"Tested recPerimeter on inputs 3.0 and 2")

myTests().main()
        </input>
      </program>
    </exercise>
    <p>Look at the unit tests above and answer the questions below.</p>
    <exercise>
      <statement>
        <p>Q-4: What is the class that is defined in the unit test code? <var/>  </p>
      </statement>
      <setup>
        <var>
          <condition string="myTests">
            <feedback>
              <p>This defines the class myTests</p>
            </feedback>
          </condition>
          <condition string=".*">
            <feedback>
              <p>What is the name after the class keyword?</p>
            </feedback>
          </condition>
        </var>
      </setup>
    </exercise>
    <exercise>
      <statement>
        <p>Q-5: What class does the defined class inherit from? <var/>  </p>
      </statement>
      <setup>
        <var>
          <condition string="TestCaseGui">
            <feedback>
              <p>The parent is TestCaseGui</p>
            </feedback>
          </condition>
          <condition string=".*">
            <feedback>
              <p>What is the name in the () after the class name?</p>
            </feedback>
          </condition>
        </var>
      </setup>
    </exercise>
    <exercise label="ut-ex-compare-valuep-ac">
      <statement>
        <p>Run the code below to see what the unit tests do and then add another test to
                test the result when x is higher than y.</p>
      </statement>
      <program xml:id="ut-ex-compare-valuep-ac_editor" interactive="activecode" language="python">
        <input>
def checkValue(x, y):
    if x &lt; y:
        return "lower"
    elif x == y:
        return "equal"
    else:
        return "higher"

from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def test_checkValue(self):
        self.assertEqual(checkValue(3, 5), "lower", 'checkValue(3, 5)')
        self.assertEqual(checkValue(5, 5), "equal", 'checkValue(5, 5)')

myTests().main()
        </input>
      </program>
    </exercise>
    <exercise label="ut_get_middle_ac_add_test">
      <statement>
        <p>Fix the function below to return the middle characters from the passed string.
                If the string has an odd length then return the middle character.  If the string has an even
                length return the two middle characters.  For example, <c>get_middle('abc')</c> returns
                <c>'b'</c> and <c>get_middle('abcd')</c> returns <c>'bc'</c>.  Also add tests cases to
                test the result when <c>str</c> has only one or two characters in it.</p>
      </statement>
      <program xml:id="ut_get_middle_ac_add_test_editor" interactive="activecode" language="python">
        <input>
def get_middle(str):
    len_s = len(str)
    mid = len_s // 2
    if len_s % 2 == 0:
        return str[mid - 1: mid + 1]
    else:
        return str

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(get_middle('abc'), 'b', "get_middle('abc')")
        self.assertEqual(get_middle('abcd'), 'bc', "get_middle('abcd')")
        self.assertEqual(get_middle('12345'), '3', "get_middle('12345')")
        self.assertEqual(get_middle('123456'), '34', "get_middle('123456')")

myTests().main()
        </input>
      </program>
    </exercise>
    <p>Add tests to check the result when <c>str</c> has one or two letters.</p>
    <note>
      <p>Unit tests in this ebook must include the line <c>from unittest.gui import TestCaseGui</c>
                    and inherit from <c>TestCaseGui</c>.  Outside of this
                    ebook you should include the line <c>import unittest</c> and inherit from <c>unittest.TestCase</c>.</p>
    </note>
    <p>Unit tests inherit from a class that includes several methods. As you can see
                from the code above one of the methods is <c>assertEqual</c> which returns the result of <c>a == b</c>.</p>
    <exercise label="ut_method_def_dnd">
      <statement>
        <p>Drag each method to what it returns.</p>
      </statement>
      <feedback>
        <p>Read this page and try again.</p>
      </feedback>
      <matches>
        <match order="1">
          <premise>assertTrue(x)</premise>
          <response> x == True</response>
        </match>
        <match order="2">
          <premise>assertIn(a, b)</premise>
          <response>a in b</response>
        </match>
        <match order="3">
          <premise>assertIsInstance(a, b)</premise>
          <response>isinstance(a, b)</response>
        </match>
        <match order="4">
          <premise>assertEqual(a, b)</premise>
          <response>a == b</response>
        </match>
        <match order="5">
          <premise>assertNotEqual(a, b)</premise>
          <response>a != b</response>
        </match>
        <match order="6">
          <premise>assertFalse(x)</premise>
          <response> x == False</response>
        </match>
      </matches>
    </exercise>
    <note>
      <p>All of the assert methods can take a third parameter which is a message (string) that explains what was tested.</p>
    </note>
  </subsection>
  <subsection xml:id="inheritance_writing-unit-tests">
    <title>Writing Unit Tests</title>
    <p>To write a unit test you need to do the following.</p>
    <p>
      <ol marker="1">
        <li>
          <p>Import the library (<c>from unittest.gui import TestCaseGui</c> in this ebook and <c>import unittest</c> elsewhere)</p>
        </li>
        <li>
          <p>Create a class that inherits from the correct class (<c>TestCaseGui</c> in this ebook and <c>unittest.TestCase</c> elsewhere )</p>
        </li>
        <li>
          <p>Optionally create a <c>setUp</c> method which is called before <em>every</em> test method is run.</p>
        </li>
        <li>
          <p>Create one or more methods that start with <c>test</c> to test different things in your code</p>
          <p>
            <ul>
              <li>
                <p>Test one or more results using <c>assert</c> methods such as <c>self.assertEqual</c></p>
              </li>
            </ul>
          </p>
        </li>
        <li>
          <p>Optionally create a <c>tearDown</c> method which is called after <em>every</em> test method is run</p>
        </li>
        <li>
          <p>Run the tests using <c>ClassName.main()</c> in the ebook and otherwise <c>unittest.main()</c></p>
        </li>
      </ol>
    </p>
    <note>
      <p>The order that the tests run is not the order they are defined.</p>
    </note>
    <exercise label="ut-ex-car-class-unit-tests">
      <statement>
        <p>Fix the code below to run and notice the test results.</p>
      </statement>
      <program xml:id="ut-ex-car-class-unit-tests_editor" interactive="activecode" language="python">
        <input>
class Car:

    def __init__(self, make, model, color):
        self.make = make
        model = model
        self.color = color

    def __str__(self)
        return f"{self.make} {self.model} {self.color}"

    def set_color(self, new_color):
        self.color = new_color

from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def setUp():
        self.c1 = Car("Ford", "Volt", "blue")
        self.c2 = Car("Toyota", "Prius", "red")

    def testInitAndStr(self):
        self.assertEqual(c1.__str__(), 'Ford Volt blue', "Test of init and str")
        self.assertEqual(self.c2.__str__(), 'Toyota Prius red', "Test of init and str")

    def testSetColor(self):
        # try to set the color
        self.c1.set_color("white")

        # check that it changed to the new color
        self.assertEqual(self.c1.color, "white", "Test of set color")

myTests().main()
        </input>
      </program>
    </exercise>
    <p>Modify the code above to add a method in the <c>Car</c> class to set (change) the model name and add a test method to test that the new method works.</p>
    <exercise label="ut-pp-test-lower" numbered="yes" adaptive="yes" indentation="show" language="python">
      <statement>
        <p>Put the code in order to create tests for a Person class initials method that returns the first letter
                from the first name followed by the first letter of the last name. The Person class initializer
                takes the person's first name and last name.  First define a <c>setUp</c> method to
                create two Person objects and then define the <c>test_initials</c> method.</p>
      </statement>
      <blocks>
        <block order="1">
          <cline>import unittest</cline>
        </block>
        <block order="6">
          <choice correct="yes">
            <cline>class TestPerson(unittest.TestCase):</cline>
          </choice>
          <choice>
            <cline>class TestPerson(TestCase): </cline>
          </choice>
        </block>
        <block order="4">
          <choice correct="yes">
            <cline>    def setUp(self):</cline>
          </choice>
          <choice>
            <cline>    def setUp(): </cline>
          </choice>
        </block>
        <block order="10">
          <cline>        self.p1 = Person("Lionel", "Roberts")</cline>
        </block>
        <block order="8">
          <choice correct="yes">
            <cline>        self.p2 = Person("Kayla", "Wig")</cline>
          </choice>
          <choice>
            <cline>        p2 = Person("Kayla", "Wig") </cline>
          </choice>
        </block>
        <block order="7">
          <cline>    def test_initials(self):</cline>
        </block>
        <block order="3">
          <cline>        p1i = self.p1.initials()</cline>
          <cline>        self.assertEqual(p1i, "LR", "p1.initials()")</cline>
          <cline>        p2i = self.p2.initials()</cline>
          <cline>        self.assertEqual(p2i, "KW", "p2.initials()")</cline>
        </block>
      </blocks>
    </exercise>
  </subsection>
</section>
