<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.17.1: http://docutils.sourceforge.net/" />

    <title>19.10. Group Work with Multiple Classes &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.cd688d1623c9932e.css?v=630908D8" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.bf73008a1eba9f4e.css?v=7DED055E" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/runtime.439d581f758f2b99.bundle.js?v=2E54442"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.51d9dde4e5162190.bundle.js?v=9CAA2D33"></script>
    <script src="../_static/runestone.42b576791a4ed54b.bundle.js?v=B7293199"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="19.11. Multiple Classes Practice" href="MultipleClassesPractice.html" />
    <link rel="prev" title="19.9. Instances as Return Values" href="InstancesasReturnValues.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    
    eBookConfig.useRunestoneServices = {{use_services}};
    eBookConfig.host = '';
    eBookConfig.app = eBookConfig.host + '/runestone';
    eBookConfig.course = '{{ course_name }}';
    eBookConfig.basecourse = '{{ base_course }}';
    eBookConfig.isLoggedIn = {{ is_logged_in}};
    eBookConfig.email = '{{ user_email }}';
    eBookConfig.isInstructor = {{ is_instructor }};
    eBookConfig.username = '{{ user_id}}';
    eBookConfig.readings = {{ readings|safe}};
    eBookConfig.activities = {{ activity_info|safe }}
    eBookConfig.downloadsEnabled = {{downloads_enabled}};
    eBookConfig.allow_pairs = {{allow_pairs}}
    eBookConfig.enableCompareMe = {{enable_compare_me}};
    // _`new_server_prefix`: Defined by the server to be the path to the new server.
    eBookConfig.new_server_prefix = '{{ new_server_prefix }}';
    
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->

    
        {% if serve_ad is defined and serve_ad and settings.adsenseid %}
         <script data-ad-client="{{settings.adsenseid}}" async src="https://pagead2.googlesyndication.com/pagead/js/adsbygoogle.js"></script>
        {% endif %}
    



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
            
            {{ course_name }}
            
          
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="browsing_warning"></span></li>

        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li id="inst_peer_link"><a href='/runestone/peer/instructor.html'>Peer Instruction (Instructor)</a></li>            
            <li><a href='/runestone/peer/student.html'>Peer Instruction (Student)</a></li>            
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/ns/books/published/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=MultipleClasses'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
            <script></script>
<li class="dropdown globaltoc-container">
    <a
        href="../index.html"
        class="dropdown-toggle"
        data-toggle="dropdown"
        aria-label="index"
        >This Chapter<b class="caret"></b
    ></a>
    <ul class="dropdown-menu globaltoc">
         
        {% for row in subchapter_list %}
        <li class="toctree-l1">
            <a class="reference internal" href="{{ row['subchap_uri']}}"
                >{{ row['title']}}</a
            >
        </li>
        {% endfor %} 
        
    </ul>
</li>
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">19.9. </span>Instances as Return Values" data-toggle="tooltip">
    <a href="InstancesasReturnValues.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">19.11. </span>Multiple Classes Practice' data-toggle="tooltip" >              
    <a href="MultipleClassesPractice.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->

    
        {% if settings.num_banners > 0 and settings.show_rs_banner %}
         <div style="width: 100%;">
            <a href="/runestone/default/donate?banner={{banner_num}}">
            <img class="runestone-banner"
                 src="/runestone/static/images/RunestoneBanner{{banner_num}}.png"
                 alt="Please Support Runestone"
                 style="display: block; width: 728px; margin-left: auto; margin-right: auto; margin-top: 10px;">
            </a>
         </div>
        {% endif %}
    


  
  <section id="group-work-with-multiple-classes">
<span id="index-0"></span><h1><span class="section-number">19.10. </span>Group Work with Multiple Classes<a class="headerlink" href="#group-work-with-multiple-classes" title="Permalink to this headline">¶</a></h1>
<p>It is best to use a POGIL approach with the following. In POGIL students work
in groups on activities and each member has an assigned role.  For more information see <a class="reference external" href="https://cspogil.org/Home">https://cspogil.org/Home</a>.</p>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>If you work in a group, have only one member of the group fill in the answers on this page.  That person will be able to share answers with the group at the bottom of the page.</p>
</div>
<p><strong>Learning Objectives</strong></p>
<p>Students will know and be able to do the following.</p>
<p><em>Content Objectives:</em></p>
<ul class="simple">
<li><p>Explain how objects of one class can be used in another</p></li>
<li><p>Identify object attributes and methods</p></li>
<li><p>Identify class attributes</p></li>
<li><p>Define the parts of a UML class diagram</p></li>
<li><p>Explain how to use a description of a system to find classes, attributes, and methods.</p></li>
</ul>
<p><em>Process Objectives:</em></p>
<ul class="simple">
<li><p>Modify code that involves multiple classes</p></li>
<li><p>Predict output from code with multiple classes</p></li>
</ul>
<section id="multiple-classes">
<h2><span class="section-number">19.10.1. </span>Multiple Classes<a class="headerlink" href="#multiple-classes" title="Permalink to this headline">¶</a></h2>
<p>Object-oriented programs typically have many classes.  Each class describes
the data (attributes) that objects of that class have and their behaviors (methods). One class can use objects from another class.</p>
<p>Let’s use the Point class that we created earlier to define a Rectangle class.
A rectangle can be defined by two points if we assume that one point is the top left
and the other is the bottom right.  That means the width of a rectangle is the difference
between the x values and the height is the difference between the y values.
The area of a rectangle is the width times the height.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=mult_class_point_and_rectangle data-question_label="19.10.1.1">
<div id=mult_class_point_and_rectangle_question class="ac_question col-md-12">
<p>Run the code below to see what it prints.  You can also use the “Show in CodeLens” button to step through the code.</p>

</div>
<textarea data-lang="python" id="mult_class_point_and_rectangle_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Point:

    def __init__(self, initX, initY):
        """ Create a new point at the given coordinates. """
        self.x = initX
        self.y = initY

    def distanceFromOrigin(self):
        return ((self.x ** 2) + (self.y ** 2)) ** 0.5

    def __str__(self):
        return f"x = {self.x}, y = {self.y}"

    def halfway(self, target):
        mx = (self.x + target.x) / 2
        my = (self.y + target.y) / 2
        return Point(mx, my)


class Rectangle:

    def __init__(self, p1, p2):
        """ Create a new rectangle with the given points """
        self.p1 = p1
        self.p2 = p2

    def area(self):
        """ Return the area of the rectangle """
        width = abs(self.p1.x - self.p2.x)
        height = abs(self.p1.y - self.p2.y)
        return width * height

point1 = Point(0,0)
point2 = Point(10,5)
r1 = Rectangle(point1, point2)
print(r1.area())
point3 = Point(-5,-10)
r2 = Rectangle(point2, point3)
print(r2.area())
</textarea>
</div>
</div>

    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="19.10.1.2" data-multipleanswers="true"  id="mult_class_rectangle_object_methods"  style="visibility: hidden;">
    <p>Q-2: Look at the code above.  What are the object methods in the <code class="docutils literal notranslate"><span class="pre">Rectangle</span></code> class?  Pick all that apply.</p>
<li data-component="answer"  id="mult_class_rectangle_object_methods_opt_a"><p>self.p1</p>
</li><li data-component="feedback" id="mult_class_rectangle_object_methods_opt_a">
<p>This is an object attribute in the Rectangle class.</p></li><li data-component="answer" data-correct id="mult_class_rectangle_object_methods_opt_b"><p>self.__init__</p>
</li><li data-component="feedback" id="mult_class_rectangle_object_methods_opt_b">
<p>This is an object method in the Rectangle class.</p></li><li data-component="answer"  id="mult_class_rectangle_object_methods_opt_c"><p>self.p2</p>
</li><li data-component="feedback" id="mult_class_rectangle_object_methods_opt_c">
<p>This is an object attribute in the Rectangle class.</p></li><li data-component="answer" data-correct id="mult_class_rectangle_object_methods_opt_d"><p>self.area</p>
</li><li data-component="feedback" id="mult_class_rectangle_object_methods_opt_d">
<p>This is an object method in the Rectangle class.</p></li>

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="19.10.1.3" data-multipleanswers="false"  id="mult_class_point_rect_mcq2"  style="visibility: hidden;">
    <p>Q-3: Look at the code above.  What type of thing (class) is <code class="docutils literal notranslate"><span class="pre">p1</span></code> in the <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method in the <code class="docutils literal notranslate"><span class="pre">Rectangle</span></code> class?</p>
<li data-component="answer" data-correct id="mult_class_point_rect_mcq2_opt_a"><p>Point</p>
</li><li data-component="feedback" id="mult_class_point_rect_mcq2_opt_a">
<p>It was created by the Point class and passed to the __init__ method in the Rectangle class</p></li><li data-component="answer"  id="mult_class_point_rect_mcq2_opt_b"><p>Rectangle</p>
</li><li data-component="feedback" id="mult_class_point_rect_mcq2_opt_b">
<p>It is set in the __init__ method in the Rectangle class, but it is not a rectangle</p></li><li data-component="answer"  id="mult_class_point_rect_mcq2_opt_c"><p>Tuple</p>
</li><li data-component="feedback" id="mult_class_point_rect_mcq2_opt_c">
<p>It is not a tuple.  Point(0,0) creates an object of the Point class and calls the __init__ method in that class.</p></li>

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="19.10.1.4" data-multipleanswers="true"  id="mult_class_point_attributes"  style="visibility: hidden;">
    <p>Q-4: Look at the code above.  What are the object attributes in the <code class="docutils literal notranslate"><span class="pre">Point</span></code> class?  Pick all that apply.</p>
<li data-component="answer" data-correct id="mult_class_point_attributes_opt_a"><p>self.x</p>
</li><li data-component="feedback" id="mult_class_point_attributes_opt_a">
<p>This stores the x position for the point. It is an object attribute.</p></li><li data-component="answer"  id="mult_class_point_attributes_opt_b"><p>self.distanceFromOrigin</p>
</li><li data-component="feedback" id="mult_class_point_attributes_opt_b">
<p>This is an object method in the Point class.</p></li><li data-component="answer" data-correct id="mult_class_point_attributes_opt_c"><p>self.y</p>
</li><li data-component="feedback" id="mult_class_point_attributes_opt_c">
<p>This stores the y position for the point. It is an object attribute.</p></li><li data-component="answer"  id="mult_class_point_attributes_opt_d"><p>self.halfway</p>
</li><li data-component="feedback" id="mult_class_point_attributes_opt_d">
<p>This is an object method in the Point class.</p></li>

    </ul>
    </div>
    </section>
<section id="object-oriented-analysis-and-design">
<h2><span class="section-number">19.10.2. </span>Object-Oriented Analysis and Design<a class="headerlink" href="#object-oriented-analysis-and-design" title="Permalink to this headline">¶</a></h2>
<p>How do you determine the classes that you need and the attributes and methods that objects of the class should have?
In object-oriented analysis and design you often walk through a scenario of how the system will work.  Write down all the nouns and verbs that are mentioned.  Nouns are potential classes or attributes.  Verbs are potential methods. If a noun represents simple data like a number or string it is usually just an attribute of a class. If a noun has data associated with it or behavior then it is likely a class.</p>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>One approach to object-oriented analysis and design is to use index cards to represent classes.  Write the class name at the top of the card. Put the data or attributes that each object needs to keep track of below that and the methods or behaviors that objects need to be able to do below that.</p>
</div>
<p>How would you create software for a simple card game?  Play the Aces and Twos card game at <a class="reference external" href="https://www.mathsisfun.com/games/card-match-game.html">https://www.mathsisfun.com/games/card-match-game.html</a></p>
<p>When the game starts all of the cards are all face down in rows and columns.  You click on a card to select it and it turns over to show the suit and rank (value).  Then you click another card to select it and it also turns face up.  If the two cards match (have the same rank and value) then the cards stay facing up.  Otherwise they both turn down and you take another turn.  You win when you have found all of the matches.  The game keeps track of the total time it took you to find all of the matches.</p>

    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="19.10.2.1" data-multipleanswers="true"  id="mult_class_card_game_classes"  style="visibility: hidden;">
    <p>Q-5: Look at explanation of the game above.  Which of the following nouns could be classes?</p>
<li data-component="answer" data-correct id="mult_class_card_game_classes_opt_a"><p>Game</p>
</li><li data-component="feedback" id="mult_class_card_game_classes_opt_a">
<p>Game is a noun and you can play a game so it has behavior so it is a good candidate class.</p></li><li data-component="answer" data-correct id="mult_class_card_game_classes_opt_b"><p>Card</p>
</li><li data-component="feedback" id="mult_class_card_game_classes_opt_b">
<p>Card is a noun and you can turn a card so it has behavior.  A card has a suit and rank so it has attributes.</p></li><li data-component="answer"  id="mult_class_card_game_classes_opt_c"><p>Row</p>
</li><li data-component="feedback" id="mult_class_card_game_classes_opt_c">
<p>Row is a noun, but here it just describes how to layout the cards when you start the game. It doesn’t need to be a class.</p></li><li data-component="answer"  id="mult_class_card_game_classes_opt_d"><p>Suit</p>
</li><li data-component="feedback" id="mult_class_card_game_classes_opt_d">
<p>Suit is a noun, but it can be an integer representing one of the four suits.  It will be an attribute of a Card object.</p></li>

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="19.10.2.2" data-multipleanswers="true"  id="mult_class_card_game_bevaiors"  style="visibility: hidden;">
    <p>Q-6: Look at explanation of the game above.  Which of the following verbs could be methods (behaviors)?</p>
<li data-component="answer" data-correct id="mult_class_card_game_bevaiors_opt_a"><p>select</p>
</li><li data-component="feedback" id="mult_class_card_game_bevaiors_opt_a">
<p>You select a card so it could be a method on the Card class.</p></li><li data-component="answer" data-correct id="mult_class_card_game_bevaiors_opt_b"><p>turn</p>
</li><li data-component="feedback" id="mult_class_card_game_bevaiors_opt_b">
<p>You turn a card over so it could be a method on the Card class.</p></li><li data-component="answer"  id="mult_class_card_game_bevaiors_opt_c"><p>show</p>
</li><li data-component="feedback" id="mult_class_card_game_bevaiors_opt_c">
<p>While this is a verb, it happens when you turn over a card. It probably doesn’t need to be another method.</p></li><li data-component="answer" data-correct id="mult_class_card_game_bevaiors_opt_d"><p>match</p>
</li><li data-component="feedback" id="mult_class_card_game_bevaiors_opt_d">
<p>In the game you need to determine if two cards match, so you do need a method to support this.</p></li>

    </ul>
    </div>
    </section>
<section id="uml-diagrams">
<h2><span class="section-number">19.10.3. </span>UML Diagrams<a class="headerlink" href="#uml-diagrams" title="Permalink to this headline">¶</a></h2>
<p>UML means Unified Modeling Language.  It is a standard way to visualize the design of an object-oriented program.  See <a class="reference external" href="https://en.wikipedia.org/wiki/Unified_Modeling_Language">https://en.wikipedia.org/wiki/Unified_Modeling_Language</a> for more information.</p>
<p>The figure below shows a possible class diagram for a Card class.</p>
<figure class="align-default">
<img alt="Class diagram for a Card class" src="../_images/card.png" />
</figure>
<p>In UML each class is shown in a box with the class name on top optionally followed by a line then the attributes and then optionally followed by a line and then the behaviors (methods).</p>
<p>To play the card game you need two of each type of card.  So you will need to two decks.  A deck has 52 cards in the four suits from Ace to King.</p>
<figure class="align-default">
<img alt="Class diagram for Card and Deck" src="../_images/cardAndDeck.png" />
</figure>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=mult_classes_card_and_deck data-question_label="19.10.3.1">
<div id=mult_classes_card_and_deck_question class="ac_question col-md-12">
<p>Run the code below to see what it prints.  You can also use the “Show in CodeLens” button to step through the code.</p>

</div>
<textarea data-lang="python" id="mult_classes_card_and_deck_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Card:

    suit_names = ["Clubs", "Diamonds", "Hearts", "Spades"]
    rank_names = ["Ace", "2", "3", "4", "5", "6", "7", "8", "9",  "10", "Jack", "Queen", "King"]

    def __init__(self, suit=0, rank = 0):
        self.suit = suit
        self.rank = rank

    def __str__(self):
        return Card.rank_names[self.rank] + " of " + Card.suit_names[self.suit]

class Deck:

    def __init__(self):
        self.card_list = []
        for i in range(len(Card.suit_names)):
            for j in range(len(Card.rank_names)):
                self.card_list.append(Card(i, j))

    def __str__(self):
        output = ""
        for card in self.card_list:
            output += card.__str__() + ", "
        return output


def main():
    my_deck = Deck()
    print(my_deck)

if __name__ == "__main__":
    main()
</textarea>
</div>
</div>
<p>Notice that we have defined a list of <code class="docutils literal notranslate"><span class="pre">suit_names</span></code> and <code class="docutils literal notranslate"><span class="pre">rank_names</span></code> in the <code class="docutils literal notranslate"><span class="pre">Card</span></code> class.  These will be created in the class <code class="docutils literal notranslate"><span class="pre">Card</span></code> and not in each object of the class <code class="docutils literal notranslate"><span class="pre">Card</span></code>.  They are called <code class="docutils literal notranslate"><span class="pre">class</span> <span class="pre">attributes</span></code>.
Each object of the <code class="docutils literal notranslate"><span class="pre">Card</span></code> class doesn’t need to have its own copy of these lists.  That would be a waste of space. Instead they are created in the class and all objects of the class have access to it.</p>

        <div class="runestone ">
        <div data-component="fillintheblank" data-question_label="19.10.3.2" id="mult_class_attribute_access_fitb"  style="visibility: hidden;">
            <p>Q-8: Use <input type="text">.rank_names to access the class attribute <code class="docutils literal notranslate"><span class="pre">rank_names</span></code> in the <code class="docutils literal notranslate"><span class="pre">Card</span></code> class.</p>

        <script type="application/json">
            [[{"regex": "^\\s*Card\\s*$", "regexFlags": "", "feedback": "<p>Use the name of the class to access a class attribute</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Look at the __init__ method in the Deck class</p>\n"}]]
        </script>

        </div>
        </div>
            
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="19.10.3.3" data-multipleanswers="true"  id="mult-classes-card-attributes-vs-class"  style="visibility: hidden;">
    <p>Q-9: Look at class <code class="docutils literal notranslate"><span class="pre">Card</span></code> above.  Which of the following are object attributes?  Select all that apply.</p>
<li data-component="answer"  id="mult-classes-card-attributes-vs-class_opt_a"><p>suit_names</p>
</li><li data-component="feedback" id="mult-classes-card-attributes-vs-class_opt_a">
<p>This is defined in the class not on an object.  It is a class attribute.</p></li><li data-component="answer"  id="mult-classes-card-attributes-vs-class_opt_b"><p>rank_names</p>
</li><li data-component="feedback" id="mult-classes-card-attributes-vs-class_opt_b">
<p>This is defined in the class not on an object.  It is a class attribute.</p></li><li data-component="answer" data-correct id="mult-classes-card-attributes-vs-class_opt_c"><p>self.suit</p>
</li><li data-component="feedback" id="mult-classes-card-attributes-vs-class_opt_c">
<p>This is an object attribute.  It is set in every card object.</p></li><li data-component="answer" data-correct id="mult-classes-card-attributes-vs-class_opt_d"><p>self.rank</p>
</li><li data-component="feedback" id="mult-classes-card-attributes-vs-class_opt_d">
<p>This is an object attribute.  It is set in every card object.</p></li>

    </ul>
    </div>
    
<div class="runestone">
<div data-component="shortanswer" data-question_label="19.10.3.4" class="journal" id=mult_class_diff_class_and_object_attribute  >
<p>Q-10: How can you tell if an attribute is a class or object attribute?</p>

</div>
</div> <!-- end of runestone div -->

        <div class="runestone ">
        <div data-component="fillintheblank" data-question_label="19.10.3.5" id="mult_class_deck_who_creates_cards_fitb"  style="visibility: hidden;">
            <p>Q-11: The <input type="text"> function in class <input type="text"> creates all of the card objects in the deck object.</p>

        <script type="application/json">
            [[{"regex": "^\\s*__init__\\s*$", "regexFlags": "", "feedback": "<p>The __init__ function creates all 52 cards</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>An object of a class is created with Class()</p>\n"}], [{"regex": "^\\s*Deck\\s*$", "regexFlags": "", "feedback": "<p>The Deck class creates the 52 cards when you create a new deck</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Which class creates multiple cards?</p>\n"}]]
        </script>

        </div>
        </div>
            <p>Use the following UML class diagram to answer the next few questions.</p>
<figure class="align-default">
<img alt="class diagram of a Flight and Airport class" src="../_images/flightAndAirport.png" />
</figure>

    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="19.10.3.6" data-multipleanswers="false"  id="mult-classes-flight-airport-mcq1"  style="visibility: hidden;">
    <p>Q-12: Look at class diagram for <code class="docutils literal notranslate"><span class="pre">Flight</span></code> and <code class="docutils literal notranslate"><span class="pre">Airport</span></code> above.  How many object attributes does a <code class="docutils literal notranslate"><span class="pre">Flight</span></code> object have?</p>
<li data-component="answer"  id="mult-classes-flight-airport-mcq1_opt_a"><p>3</p>
</li><li data-component="feedback" id="mult-classes-flight-airport-mcq1_opt_a">
<p>How many items are listed after the first line in the class diagram for Flight?</p></li><li data-component="answer" data-correct id="mult-classes-flight-airport-mcq1_opt_b"><p>5</p>
</li><li data-component="feedback" id="mult-classes-flight-airport-mcq1_opt_b">
<p>A Flight object has 5 object attributes: (number, departure_date, departure_time, departure_airport, and arrival_airport)</p></li><li data-component="answer"  id="mult-classes-flight-airport-mcq1_opt_c"><p>7</p>
</li><li data-component="feedback" id="mult-classes-flight-airport-mcq1_opt_c">
<p>This would be true if Flight inherited from Airport, but it does not.</p></li>

    </ul>
    </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=classes_flight_and_airport data-question_label="19.10.3.7">
<div id=classes_flight_and_airport_question class="ac_question col-md-12">
<p>Given the <code class="docutils literal notranslate"><span class="pre">Airport</span></code> class shown below and the class diagram shown above, write the <code class="docutils literal notranslate"><span class="pre">Flight</span></code> class with an <code class="docutils literal notranslate"><span class="pre">__init__</span></code> method that takes the object attributes in order from top to bottom in the class diagram (number, departure_date, departure_time, departure_airport, and arrival_airport).</p>

</div>
<textarea data-lang="python" id="classes_flight_and_airport_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Airport:

    def __init__(self, name, code):
        self.name = name
        self.code = code

    def __str__(self):
        return f"{self.name}: {self.code}"

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        a1 = Airport("Atlanta", "ATL")
        a2 = Airport("Detroit", "DTW")
        f1 = Flight(190, "10/3/2021", "16:00", a1, a2)
        self.assertEqual(f1.number, 190, "test of f1.number")
        self.assertEqual(f1.departure_date, "10/3/2021", "test of f1.departure_date")
        self.assertEqual(f1.departure_time, "16:00", "test of f1.departure_time")
        self.assertEqual(f1.departure_airport, a1, "test of f1.departure_airport")
        self.assertEqual(f1.arrival_airport, a2, "test of f1.arrival_airport")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=classes_order_and_item data-question_label="19.10.3.8">
<div id=classes_order_and_item_question class="ac_question col-md-12">
<p>Given the <code class="docutils literal notranslate"><span class="pre">Item</span></code>  and <code class="docutils literal notranslate"><span class="pre">Order</span></code> classes shown below, write the <code class="docutils literal notranslate"><span class="pre">get_total</span></code> method in the <code class="docutils literal notranslate"><span class="pre">Order</span></code> class that returns a total of all of the prices of the items in the order.</p>

</div>
<textarea data-lang="python" id="classes_order_and_item_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
class Item:

    def __init__(self, name, price):
        self.name = name
        self.price = price

    def __str__(self):
        return f"{self.name}: {self.price}"

class Order:

    def __init__(self, items):
        self.items = items


====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        i1 = Item("Burger", 6.99)
        i2 = Item("Fries", 2.99)
        i3 = Item("Coke", 1.99)
        o1 = Order([i1, i2])
        o2 = Order([i1, i3])
        o3 = Order([i2, i3])
        self.assertAlmostEqual(o1.get_total(), 9.98, 2, "test of order 1 total")
        self.assertAlmostEqual(o2.get_total(), 8.98, 2, "test of order 2 total")
        self.assertAlmostEqual(o3.get_total(), 4.98, 2, "test of order 3 total")


myTests().main()
</textarea>
</div>
</div>

<div class="runestone">
<div data-component="shortanswer" data-question_label="19.10.3.9" class="journal" id=mult_class_relationship_between_classes_sa  >
<p>Q-15: What is the relationship between the example classes on this page, such as <code class="docutils literal notranslate"><span class="pre">Point</span></code> and <code class="docutils literal notranslate"><span class="pre">Rectangle</span></code> or <code class="docutils literal notranslate"><span class="pre">Item</span></code> and <code class="docutils literal notranslate"><span class="pre">Order</span></code>?  How do they relate to each other?</p>

</div>
</div> <!-- end of runestone div -->
<p>If you worked in a group, you can copy the answers from this page to the other group members.  Select the group members below and click the button to share the answers.</p>

<div class="runestone sqcontainer %(optclass)s">
    <div data-component="groupsub" id=mult_classes_groupsub data-size_limit=4>
        <div class="col-sm-6">
            <select id="assignment_group" multiple class="assignment_partner_select" style="width: 100%">
            </select>
        </div>
        <div id="groupsub_button" class="col-sm-6">
        </div>
        <p>The Submit Group button will submit the answer for each each question
        on this page for each member of your group. It also logs you as the official
        group submitter.</p>
    </div>
</div>
</section>
</section>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 19.9. Instances as Return Values' data-toggle="tooltip">
  <a href="InstancesasReturnValues.html" aria-label="Previous - 19.9. Instances as Return Values">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 19.11. Multiple Classes Practice' data-toggle="tooltip" >
    <a href="MultipleClassesPractice.html" aria-label="Next - 19.11. Multiple Classes Practice">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>



  
    {% if settings.google_ga %}
      <script>
        var _gaq = _gaq || [];
        _gaq.push(['_setAccount', '{{ settings.google_ga }}']);
        _gaq.push(['_trackPageview']);

        (function() {
          var ga = document.createElement('script'); ga.async = true;
          ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
          var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
        })();
      </script>
    {% endif %}
  

  
    <script type="application/json" class="js-hypothesis-config">
    {
      "showHighlights": false
    }
    </script>

    <script async src="https://hypothes.is/embed.js"></script>
    <!--provide space to prevent hypothes.is sidebar from obscuring content-->
    <style>@media screen and (max-width:767px) { body { padding-right: 21px; }}</style>

  


<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>