<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>4.13. Write Code Questions &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/637.css?v=bddd2011b20db76a20ea" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.css?v=7c5388ad4389eae1a318" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/runtime.bundle.js?v=aa7c717c38c5405af9e6"></script>
    <script src="../_static/637.bundle.js?v=1123ad75f9b821cdd636"></script>
    <script src="../_static/runestone.bundle.js?v=9c8e466e8cbed4a3af5a"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="4.14. Group Work - Conditionals and Logic" href="pogil.html" />
    <link rel="prev" title="4.12. Mixed-up Code Questions" href="MixedupCode.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'py4e-int';
    eBookConfig.basecourse = 'py4e-int';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          py4e-int
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=WriteCode'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">4.12. </span>Mixed-up Code Questions" data-toggle="tooltip">
    <a href="MixedupCode.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">4.14. </span>Group Work - Conditionals and Logic' data-toggle="tooltip" >              
    <a href="pogil.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
  
<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="write-code-questions">
<h1><span class="section-number">4.13. </span>Write Code Questions<a class="headerlink" href="#write-code-questions" title="Permalink to this headline">¶</a></h1>
<ol class="arabic">
<li><div id='cndtnl-exercise-five' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=cndtnl-wc-fiveq data-question_label="4.13.1">
<div id=cndtnl-wc-fiveq_question class="ac_question col-md-12">
<p>Fix the errors in the code, and change it to use only <strong>one</strong> if statement. The code should print
“The number is 5” when the number is 5, and should print “The number is NOT 5” when it is not.</p>

</div>
<textarea data-lang="python" id="cndtnl-wc-fiveq_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
x = 5
if x == 5:
print("The number is 5")
if x != 5:
print("The number is NOT 5")
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=cndtnl-wc-fivea data-question_label="4.13.2">
<div id=cndtnl-wc-fivea_question class="ac_question col-md-12">
<p>Fix the errors in the code, and change it to use only <strong>one</strong> if statement. The code should print
“The number is 5” when the number is 5, and should print “The number is NOT 5” when it is not.</p>

</div>
<textarea data-lang="python" id="cndtnl-wc-fivea_editor" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
# Initialize x and set condition for when x = 5
x = 5
if x == 5:
    # Indent the body of if statement
    print("The number is 5")
# Use else for the secondary condition
else:
    # Indent the body of else statement
    print("The number is NOT 5")
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=cndtnl-wc-hiq data-question_label="4.13.3">
<div id=cndtnl-wc-hiq_question class="ac_question col-md-12">
<p>Complete the conditional below so the word “Hi” is printed if x does not
equal 3 and “Hello” is printed otherwise.</p>

</div>
<textarea data-lang="python" id="cndtnl-wc-hiq_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
x = 4
if x
</textarea>
</div>
</div>
</li>
<li><div id='cndtnl-exercise-pay' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=cndtnl-wc-payq data-question_label="4.13.4">
<div id=cndtnl-wc-payq_question class="ac_question col-md-12">
<p>Complete the pay computation to give the employee 1.5 times the hourly rate for hours worked
above 40 hours, if the regular pay rate is $10 an hour. Then set <code class="docutils literal notranslate"><span class="pre">grossPay</span></code> equal to the
amount an employee would be paid for working 45 hours.</p>

</div>
<textarea data-lang="python" id="cndtnl-wc-payq_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
hours = 45
rate = 10
overtimeRate =  #your code here
grossPay = 0
if hours  #your code here
    grossPay = hours * rate
else:
    overTime =  #your code here
    grossPay =  #your code here
print(grossPay)

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(grossPay, 475.0)

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=cndtnl-wc-paya data-question_label="4.13.5">
<div id=cndtnl-wc-paya_question class="ac_question col-md-12">
<p>Complete the pay computation to give the employee 1.5 times the hourly rate for hours worked
above 40 hours, if the regular pay rate is $10 an hour. Then set <code class="docutils literal notranslate"><span class="pre">grossPay</span></code> equal to the
amount an employee would be paid for working 45 hours.</p>

</div>
<textarea data-lang="python" id="cndtnl-wc-paya_editor" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
# Initializing variables
hours = 45
rate = 10
# overtimeRate is 1.5 the rate amount
overtimeRate = rate * 1.5
grossPay = 0
# Begin conditional to see if hours are within regular pay
if hours <= 40:
    #if within 40 hours, pay will be hours * rate
    grossPay = hours * rate
# Else statement for when hours are greater than 40
else:
    # Create variable for overtime hours
    overTime = hours % 40
    # Pay will equal the regular rate for 40 hours,
    # plus the overtime rate for the extra hours
    grossPay = (rate * 40) + (overTime * overtimeRate)
# Print the final pay
print(grossPay)
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=cndtnl-wc-pay2q data-question_label="4.13.6">
<div id=cndtnl-wc-pay2q_question class="ac_question col-md-12">
<p>Rewrite your pay program using <code class="docutils literal notranslate"><span class="pre">try</span></code> and <code class="docutils literal notranslate"><span class="pre">except</span></code> so that your program handles non-numeric
input gracefully by printing a message and exiting the program. The following shows two
executions of the program:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">Enter</span> <span class="n">Hours</span><span class="p">:</span> <span class="mi">20</span>
<span class="n">Enter</span> <span class="n">Rate</span><span class="p">:</span> <span class="n">nine</span>
<span class="n">Error</span><span class="p">,</span> <span class="n">please</span> <span class="n">enter</span> <span class="n">numeric</span> <span class="nb">input</span>
</pre></div>
</div>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">Enter</span> <span class="n">Hours</span><span class="p">:</span> <span class="n">forty</span>
<span class="n">Error</span><span class="p">,</span> <span class="n">please</span> <span class="n">enter</span> <span class="n">numeric</span> <span class="nb">input</span>
</pre></div>
</div>

</div>
<textarea data-lang="python" id="cndtnl-wc-pay2q_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

</textarea>
</div>
</div>
</li>
<li><div id='cndtnl-exercise-distance' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=cndtnl-wc-distanceq data-question_label="4.13.7">
<div id=cndtnl-wc-distanceq_question class="ac_question col-md-12">
<p>Write the code to calculate and print the cost of a 14 mile cab ride. If the distance traveled
is less than or equal to 12 miles the cost is $2.00 a mile, and if the distance traveled is more
than 12 miles the cost is $1.50 a mile. Assign the final cost to the variable <code class="docutils literal notranslate"><span class="pre">total</span></code>.</p>

</div>
<textarea data-lang="python" id="cndtnl-wc-distanceq_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
distance = 14

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(total, 21.0)

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=cndtnl-wc-distancea data-question_label="4.13.8">
<div id=cndtnl-wc-distancea_question class="ac_question col-md-12">
<p>Write the code to calculate and print the cost of a 14 mile cab ride. If the distance traveled
is less than or equal to 12 miles the cost is $2.00 a mile, and if the distance traveled is more
than 12 miles the cost is $1.50 a mile. Assign the final cost to the variable <code class="docutils literal notranslate"><span class="pre">total</span></code>.</p>

</div>
<textarea data-lang="python" id="cndtnl-wc-distancea_editor" 
      data-timelimit=25000     data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
# distance variable given
distance = 14
# Set conditions
if distance <= 12:
    # rate will be 2 if the distance is less than or equal to 12
    rate = 2.00
if distance > 12:
    # rate will be 1.5 if the distance is greater than 12
    rate = 1.50
# Calculate trip cost
total = distance * rate
# Print final statement
print("Total cost of trip: " + str(total))
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=cndtnl-wc-scoreq data-question_label="4.13.9">
<div id=cndtnl-wc-scoreq_question class="ac_question col-md-12">
<p>Write a program to prompt for a score between 0.0 and 1.0. If the score is out of range,
print an error message. If the score is between 0.0 and 1.0, print a grade using the
following table:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span> <span class="n">Score</span>   <span class="n">Grade</span>
<span class="o">&gt;=</span> <span class="mf">0.9</span>     <span class="n">A</span>
<span class="o">&gt;=</span> <span class="mf">0.8</span>     <span class="n">B</span>
<span class="o">&gt;=</span> <span class="mf">0.7</span>     <span class="n">C</span>
<span class="o">&gt;=</span> <span class="mf">0.6</span>     <span class="n">D</span>
 <span class="o">&lt;</span> <span class="mf">0.6</span>     <span class="n">F</span>
</pre></div>
</div>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">Enter</span> <span class="n">score</span><span class="p">:</span> <span class="mf">0.95</span>
<span class="n">A</span>
</pre></div>
</div>

</div>
<textarea data-lang="python" id="cndtnl-wc-scoreq_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

</textarea>
</div>
</div>
</li>
<li><div id='cndtnl-exercise-weight' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=cndtnl-wc-weightq data-question_label="4.13.10">
<div id=cndtnl-wc-weightq_question class="ac_question col-md-12">
<p>Fix the example such that the cost of frozen yogurt is 0 if you pour exactly 1 lb. in your cup.</p>

</div>
<textarea data-lang="python" id="cndtnl-wc-weightq_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
weight = 0.5
if weight < 1:
    price = 1.45
if weight > 1:
    price = 1.15
total = weight * price
print(weight)
print(price)
print(total)
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=cndtnl-wc-weighta data-question_label="4.13.11">
<div id=cndtnl-wc-weighta_question class="ac_question col-md-12">
<p>Fix the example such that the cost of frozen yogurt is 0 if you pour exactly 1 lb. in your cup.</p>

</div>
<textarea data-lang="python" id="cndtnl-wc-weighta_editor" 
      data-timelimit=25000     data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
# Initial weight is given
weight = 0.5
# Set conditionals
if weight < 1:
    # Set price to 1.45 when weight is less than 1
    price = 1.45
# Create condition for when weight is 1
if weight == 1:
    # Set price to 0 when weight is 1
    price = 0
if weight > 1:
    # Set price to 1.15 when weight is greater than 1
    price = 1.15
# calculate total using newly assigned variables
total = weight * price
# Print variables
print(weight)
print(price)
print(total)
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=cndtnl-wc-walletq data-question_label="4.13.12">
<div id=cndtnl-wc-walletq_question class="ac_question col-md-12">
<p>Write a procedure that takes 2 ints, total price, and amount in wallet. Print “You have enough
money” if the difference between the wallet and price is 0 or greater; otherwise,
print “Get more money”.</p>

</div>
<textarea data-lang="python" id="cndtnl-wc-walletq_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

</textarea>
</div>
</div>
</li>
<li><div id='cndtnl-exercise-leapYear' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=cndtnl-wc-leapYearq data-question_label="4.13.13">
<div id=cndtnl-wc-leapYearq_question class="ac_question col-md-12">
<p>3 criteria must be taken into account to identify leap years:</p>
<ul class="simple">
<li><p>The year is evenly divisible by 4;</p></li>
<li><p>If the year can be evenly divided by 100, it is NOT a leap year, unless;</p></li>
<li><p>The year is also evenly divisible by 400. Then it is a leap year.</p></li>
</ul>
<p>Write a program that takes a year as a parameter and sets <code class="docutils literal notranslate"><span class="pre">leapYear</span></code> equal to <code class="docutils literal notranslate"><span class="pre">True</span></code> if the year is a
leap year, <code class="docutils literal notranslate"><span class="pre">False</span></code> otherwise. (use a few different years to test your work)</p>

</div>
<textarea data-lang="python" id="cndtnl-wc-leapYearq_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=cndtnl-wc-leapYeara data-question_label="4.13.14">
<div id=cndtnl-wc-leapYeara_question class="ac_question col-md-12">
<p>3 criteria must be taken into account to identify leap years:</p>
<ul class="simple">
<li><p>The year is evenly divisible by 4;</p></li>
<li><p>If the year can be evenly divided by 100, it is NOT a leap year, unless;</p></li>
<li><p>The year is also evenly divisible by 400. Then it is a leap year.</p></li>
</ul>
<p>Write a program that takes a year as a parameter and sets <code class="docutils literal notranslate"><span class="pre">leapYear</span></code> equal to <code class="docutils literal notranslate"><span class="pre">True</span></code> if the year is a
leap year, <code class="docutils literal notranslate"><span class="pre">False</span></code> otherwise. (use a few different years to test your work)</p>

</div>
<textarea data-lang="python" id="cndtnl-wc-leapYeara_editor" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
year = 1900 #TEST WITH ANY YEAR
# Initializing leapYear to False since most years are not leap years
leapYear = False

# Use modulus to see if the year is divisible by 4
if year % 4 == 0:
    # If divisible by 4, use modulus to see if divisible by 100
    if year % 100 == 0:
        # if divisible by 4 and 100, use modulus to see if divisible by 400
        if year % 400 == 0:
            # if divisible by 4, 100, and 400 it is a leap year
            leapYear = True
    else: # if divisible by 4 and not 100 it is a leap year
        leapYear = True
print(leapYear)
''' Since leapYear started as False, we do not need to reset it to False
    for the instances that the year is not a leap year '''

# ANOTHER OPTION
year = 1900
leapYear = False
if year % 400 == 0:
    leapYear = True
elif year % 4 == 0 and year % 100 != 0:
    leapYear = True
print(leapYear)
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=cndtnl-wc-evenq data-question_label="4.13.15">
<div id=cndtnl-wc-evenq_question class="ac_question col-md-12">
<p>Finish the following code. It first sets <code class="docutils literal notranslate"><span class="pre">n</span></code> to a number input by a user.  Convert the number from a string to an integer and set <code class="docutils literal notranslate"><span class="pre">result</span></code> to <code class="docutils literal notranslate"><span class="pre">True</span></code> if the
number is an <strong>even number</strong> (evenly divisible by two) and <code class="docutils literal notranslate"><span class="pre">False</span></code> if it is <strong>odd</strong>. Note: use the modulo operator.</p>

</div>
<textarea data-lang="python" id="cndtnl-wc-evenq_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
n = input("Enter a number")
</textarea>
</div>
</div>
</li>
</ol>
</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 4.12. Mixed-up Code Questions' data-toggle="tooltip">
  <a href="MixedupCode.html" aria-label="Previous - 4.12. Mixed-up Code Questions">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 4.14. Group Work - Conditionals and Logic' data-toggle="tooltip" >
    <a href="pogil.html" aria-label="Next - 4.14. Group Work - Conditionals and Logic">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>