<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Questions for Cognitive Load Experiment - B &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/637.css?v=bddd2011b20db76a20ea" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.css?v=7c5388ad4389eae1a318" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/runtime.bundle.js?v=aa7c717c38c5405af9e6"></script>
    <script src="../_static/637.bundle.js?v=1123ad75f9b821cdd636"></script>
    <script src="../_static/runestone.bundle.js?v=9c8e466e8cbed4a3af5a"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'py4e-int';
    eBookConfig.basecourse = 'py4e-int';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          py4e-int
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=ex-cogload-B'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
  
<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="questions-for-cognitive-load-experiment-b">
<h1>Questions for Cognitive Load Experiment - B<a class="headerlink" href="#questions-for-cognitive-load-experiment-b" title="Permalink to this headline">¶</a></h1>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=exp1_q1_write data-question_label="1">
<div id=exp1_q1_write_question class="ac_question col-md-12">
<p>Finish the function <code class="docutils literal notranslate"><span class="pre">has22</span></code> below to return <code class="docutils literal notranslate"><span class="pre">True</span></code> if there are at least two items in the list <code class="docutils literal notranslate"><span class="pre">nums</span></code> that are adjacent and both equal to 2, otherwise return <code class="docutils literal notranslate"><span class="pre">False</span></code>.  For example, return <code class="docutils literal notranslate"><span class="pre">True</span></code> for <code class="docutils literal notranslate"><span class="pre">has22([1,</span> <span class="pre">2,</span> <span class="pre">2])</span></code> since there are two adjacent items equal to 2 (at index 1 and 2) and <code class="docutils literal notranslate"><span class="pre">False</span></code> for <code class="docutils literal notranslate"><span class="pre">has22([2,</span> <span class="pre">1,</span> <span class="pre">2])</span></code> since the 2’s are not adjacent.</p>

</div>
<textarea data-lang="python" id="exp1_q1_write_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def has22(nums):

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(has22([1, 2, 2]), True, 'has22([1, 2, 2])')
        self.assertEqual(has22([1, 2, 1, 2]), False, 'has22([1, 2, 1, 2])')
        self.assertEqual(has22([2, 1, 2]), False, 'has22([2, 1, 2])')
        self.assertEqual(has22([2, 2, 1]), True, 'has22([2, 2, 1])')
        self.assertEqual(has22([3, 4, 2]), False, 'has22([3, 4, 2])')
        self.assertEqual(has22([2]), False, 'has22([2])')
        self.assertEqual(has22([]), False, 'has22([])')
        self.assertEqual(has22([3, 3, 1]), False, 'has22([3, 3, 1])')
        self.assertEqual(has22([1, 4, 4]), False, 'has22([1, 4, 4])')

myTests().main()
</textarea>
</div>
</div>

<div class="runestone">
<ul data-component="poll" id=exp1_q1_write_poll data-comment class='alert alert-warning' data-results='instructor' data-question_label="2" >
In solving the preceding problem I invested:

<li>1. Very, very low mental effort</li>

<li>2. Very low mental effort</li>

<li>3. Low mental effort</li>

<li>4. Rather low mental effort</li>

<li>5. Neither low nor high mental effort</li>

<li>6. Rather high mental effort</li>

<li>7. High mental effort</li>

<li>8. Very high mental effort</li>

<li>9. Very, very high mental effort</li>
</ul></div>
        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="Count_Target_In_Range_Order" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Put the code in order to define <code class="docutils literal notranslate"><span class="pre">countInRange</span></code> that returns a count of the number of times that a <code class="docutils literal notranslate"><span class="pre">target</span></code> value appears in a list between the <code class="docutils literal notranslate"><span class="pre">start</span></code> and <code class="docutils literal notranslate"><span class="pre">end</span></code> indices (inclusive).  For example, <code class="docutils literal notranslate"><span class="pre">countInRange(1,2,4,[1,</span> <span class="pre">2,</span> <span class="pre">1,</span> <span class="pre">1,</span> <span class="pre">1,</span> <span class="pre">1])</span></code> should return <code class="docutils literal notranslate"><span class="pre">3</span></code> since there are three 1’s between index 2 and 4 inclusive.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="3"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def countInRange(target, start, end, numList):
---
    count = 0
---
    count = 1 #paired
---
    for index in range(start, end+1):
---
    for index in range(start, end): #paired
---
        current = numList[index]
---
        current = numList[start] #paired
---
        if current == target:
---
        if index == target: #paired
---
            count = count + 1
---
            count++ #paired
---
    return count
        </pre>
        </div>
        </div>
    
<div class="runestone">
<ul data-component="poll" id=Count_Target_In_Range_Order_poll data-comment class='alert alert-warning' data-results='instructor' data-question_label="4" >
In solving the preceding problem I invested:

<li>1. Very, very low mental effort</li>

<li>2. Very low mental effort</li>

<li>3. Low mental effort</li>

<li>4. Rather low mental effort</li>

<li>5. Neither low nor high mental effort</li>

<li>6. Rather high mental effort</li>

<li>7. High mental effort</li>

<li>8. Very high mental effort</li>

<li>9. Very, very high mental effort</li>
</ul></div>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=exp1_q3_write data-question_label="5">
<div id=exp1_q3_write_question class="ac_question col-md-12">
<p>Finish the function <code class="docutils literal notranslate"><span class="pre">diffMaxMin</span></code> to return the difference between the largest and smallest value in the passed list of numbers (<code class="docutils literal notranslate"><span class="pre">nums</span></code>). For example, <code class="docutils literal notranslate"><span class="pre">diffMaxMin([1,2,3])</span></code> should return 2 since the difference between 3 and 1 is 2.</p>

</div>
<textarea data-lang="python" id="exp1_q3_write_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def diffMaxMin(nums):

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(diffMaxMin([1, 2, 3]), 2, 'diffMaxMin([1, 2, 3])')
        self.assertEqual(diffMaxMin([10, 20, 30]), 20, 'diffMaxMin([10, 20, 30])')
        self.assertEqual(diffMaxMin([1, 10]), 9, 'diffMaxMin([1 , 10])')
        self.assertEqual(diffMaxMin([10, 1]), 9, 'diffMaxMin([-10 , 1])')
        self.assertEqual(diffMaxMin([4, 1, 10]), 9, 'diffMaxMin([4, 1 , 10])')
        self.assertEqual(diffMaxMin([4, -10, 1]), 14, 'diffMaxMin([4, -10 , 1])')



myTests().main()
</textarea>
</div>
</div>

<div class="runestone">
<ul data-component="poll" id=exp1_q3_write_poll data-comment class='alert alert-warning' data-results='instructor' data-question_label="6" >
In solving the preceding problem I invested:

<li>1. Very, very low mental effort</li>

<li>2. Very low mental effort</li>

<li>3. Low mental effort</li>

<li>4. Rather low mental effort</li>

<li>5. Neither low nor high mental effort</li>

<li>6. Rather high mental effort</li>

<li>7. High mental effort</li>

<li>8. Very high mental effort</li>

<li>9. Very, very high mental effort</li>
</ul></div>
        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="Total_Dict_Values_PP" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Put the code in order to define <code class="docutils literal notranslate"><span class="pre">total_values</span></code> that takes a dictionary (<code class="docutils literal notranslate"><span class="pre">dict</span></code>) and returns the total of the values in the dictionary.  For example, <code class="docutils literal notranslate"><span class="pre">total_values({'red':</span> <span class="pre">3,</span> <span class="pre">'blue':</span> <span class="pre">2,</span> <span class="pre">'green':</span> <span class="pre">20})</span></code> would return <code class="docutils literal notranslate"><span class="pre">25</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="7"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def total_values(dict):
---
def total_values(): #paired
---
    total = 0
---
    for key in dict:
---
    for key in dict #paired
---
        total += dict[key]
---
        total += key #paired
---
    return total
        </pre>
        </div>
        </div>
    
<div class="runestone">
<ul data-component="poll" id=Total_Dict_Values_PP_poll data-comment class='alert alert-warning' data-results='instructor' data-question_label="8" >
In solving the preceding problem I invested:

<li>1. Very, very low mental effort</li>

<li>2. Very low mental effort</li>

<li>3. Low mental effort</li>

<li>4. Rather low mental effort</li>

<li>5. Neither low nor high mental effort</li>

<li>6. Rather high mental effort</li>

<li>7. High mental effort</li>

<li>8. Very high mental effort</li>

<li>9. Very, very high mental effort</li>
</ul></div>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=exp1_q5_write data-question_label="9">
<div id=exp1_q5_write_question class="ac_question col-md-12">
<p>Finish the function <code class="docutils literal notranslate"><span class="pre">get_names</span></code> that takes a list of dictionaries and returns a list of strings with the names from the dictionaries.  The key for the first name is ‘first’ and the key for the last name is ‘last’.  Return a list of the full names (first last) as a string.  If the ‘first’ or ‘last’ key is missing in the dictionary use ‘Unknown’.  For example, <code class="docutils literal notranslate"><span class="pre">[{'first':</span> <span class="pre">'Ann',</span> <span class="pre">'last':</span> <span class="pre">'Brown'},</span> <span class="pre">{'first':</span> <span class="pre">'Darius'}]</span></code> should return <code class="docutils literal notranslate"><span class="pre">['Ann</span> <span class="pre">Brown',</span> <span class="pre">'Darius</span> <span class="pre">Unknown']</span></code>..</p>

</div>
<textarea data-lang="python" id="exp1_q5_write_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def get_names(d_list):

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        d1 = {'first': 'Tsi', 'last': 'Choi'}
        d2 = {'first': 'Joy', 'last': 'Ali'}
        d3 = {'first': 'Jose', 'last': 'Capo'}
        d4 = {'first': 'Malik'}
        d5 = {'last': 'Owens'}
        d_l1 = [d1, d2]
        d_l2 = [d3]
        d_l3 = [d4]
        d_l4 = [d5]
        self.assertEqual(get_names(d_l1), ['Tsi Choi', 'Joy Ali'], "get_names([{'first': 'Tsi', 'last': 'Choi'}, {'first': 'Joy', 'last': 'Ali'}])")
        self.assertEqual(get_names(d_l2), ['Jose Capo'], "get_names([{'first': 'Jose', 'last': 'Capo'}])")
        self.assertEqual(get_names(d_l3), ['Malik Unknown'], "get_names([{'first': 'Malik'}])")
        self.assertEqual(get_names(d_l4), ['Unknown Owens'], "get_names([{'last': 'Owens'}])")


myTests().main()
</textarea>
</div>
</div>

<div class="runestone">
<ul data-component="poll" id=exp1_q5_write_poll data-comment class='alert alert-warning' data-results='instructor' data-question_label="10" >
In solving the preceding problem I invested:

<li>1. Very, very low mental effort</li>

<li>2. Very low mental effort</li>

<li>3. Low mental effort</li>

<li>4. Rather low mental effort</li>

<li>5. Neither low nor high mental effort</li>

<li>6. Rather high mental effort</li>

<li>7. High mental effort</li>

<li>8. Very high mental effort</li>

<li>9. Very, very high mental effort</li>
</ul></div></div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>