<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Mixed-up Code Questions &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/637.css?v=bddd2011b20db76a20ea" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.css?v=7c5388ad4389eae1a318" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/runtime.bundle.js?v=aa7c717c38c5405af9e6"></script>
    <script src="../_static/637.bundle.js?v=1123ad75f9b821cdd636"></script>
    <script src="../_static/runestone.bundle.js?v=9c8e466e8cbed4a3af5a"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'py4e-int';
    eBookConfig.basecourse = 'py4e-int';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          py4e-int
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=MixedupCode_orig'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
  
<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="mixed-up-code-questions">
<h1>Mixed-up Code Questions<a class="headerlink" href="#mixed-up-code-questions" title="Permalink to this headline">¶</a></h1>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="list_MUC_password_creator" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">password_maker(word_list)</span></code> that takes a list of words, <code class="docutils literal notranslate"><span class="pre">word_list</span></code> and returns a string with the first and last character of
each word in a list. For example, <code class="docutils literal notranslate"><span class="pre">password_maker(&quot;show&quot;,</span> <span class="pre">&quot;me&quot;,</span> <span class="pre">&quot;the&quot;,</span> <span class="pre">&quot;money&quot;)</span></code> would return <code class="docutils literal notranslate"><span class="pre">swmetemy</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="1"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def password_maker(word_list):
---
    password = ""
---
    for word in word_list:
---
        first_char = word[0]
---
        last_char = word[-1]
---
        password += first_char
---
        password += last_char
---
    return password
---
    return word_list #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=list_MUC_password_creator_ac data-question_label="2">
<div id=list_MUC_password_creator_ac_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">password_maker(word_list)</span></code> that takes a a list of words, <code class="docutils literal notranslate"><span class="pre">word_list</span></code> and returns a string with the first and last character of each
word in a list. For example, <code class="docutils literal notranslate"><span class="pre">password_maker([&quot;show&quot;,</span> <span class="pre">&quot;me&quot;,</span> <span class="pre">&quot;the&quot;,</span> <span class="pre">&quot;money&quot;])</span></code> would return <code class="docutils literal notranslate"><span class="pre">swmetemy</span></code>.</p>

</div>
<textarea data-lang="python" id="list_MUC_password_creator_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):

        self.assertEqual(password_maker(["show", "me", "the", "money"]), "swmetemy", 'password_maker("show", "me", "the", "money")')
        self.assertEqual(password_maker(["Bye", "Now"]), "BeNw", 'password_maker(["Bye", "Now"])')
        self.assertEqual(password_maker(["Beautiful"]), "Bl", 'password_maker(["Beautiful"])')
        self.assertEqual(password_maker(["You", "complete", "me"]), "Yuceme", 'password_maker(["You", "complete", "me"])')

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="list_MUC_rainfall_v2" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">avg_rainfall(day_list)</span></code> that takes a list with the amount of rain for each day.
There are some problems with the equipment, so ignore days that have a negative number for the amount of rain. Calculate the average rainfall as the sum of the non-negative values divided by the number of non-negative values. If at least one day has a non-negative rainfall return <code class="docutils literal notranslate"><span class="pre">Average:</span> <span class="pre">(average)</span></code> and otherwise return <code class="docutils literal notranslate"><span class="pre">No</span> <span class="pre">rain</span></code>.  For example,  <code class="docutils literal notranslate"><span class="pre">avg_rainfall([30,-2,</span> <span class="pre">0])</span></code> should return <code class="docutils literal notranslate"><span class="pre">Average:</span> <span class="pre">15</span></code> and <code class="docutils literal notranslate"><span class="pre">avg_rainfall([-3,</span> <span class="pre">-2])</span></code> should return <code class="docutils literal notranslate"><span class="pre">No</span> <span class="pre">rain</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="3"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def avg_rainfall(day_list):
---
    sum = 0
    count = 0
---
    for day in day_list:
---
        if day >= 0:
---
        if day > 0: #paired
---
            sum = sum + day
            count = count + 1
---
    if count > 0:
---
        avg = sum / count
---
        return "Average: " + str(avg)
---
        return "Average: " + avg #paired
---
    else:
---
        return "No rain"
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=list_MUC_rainfall_v2_ac data-question_label="4">
<div id=list_MUC_rainfall_v2_ac_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">avg_rainfall(day_list)</span></code> that takes a list with the amount of rain for each day.
There are some problems with the equipment, so ignore days that have a negative number for the amount of rain. Calculate the average rainfall as the sum of the non-negative values divided by the number of non-negative values. If at least one day has a non-negative rainfall return <code class="docutils literal notranslate"><span class="pre">Average:</span> <span class="pre">(average)</span></code> and otherwise return <code class="docutils literal notranslate"><span class="pre">No</span> <span class="pre">rain</span></code>.  For example,  <code class="docutils literal notranslate"><span class="pre">avg_rainfall([30,-2,</span> <span class="pre">0])</span></code> should return <code class="docutils literal notranslate"><span class="pre">Average:</span> <span class="pre">15</span></code> and <code class="docutils literal notranslate"><span class="pre">avg_rainfall([-3,</span> <span class="pre">-2])</span></code> should return <code class="docutils literal notranslate"><span class="pre">No</span> <span class="pre">rain</span></code>.</p>

</div>
<textarea data-lang="python" id="list_MUC_rainfall_v2_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(avg_rainfall([30, -2, 0]), "Average: 15.0", "rainfall([30, -2, 0])")
        self.assertEqual(avg_rainfall([-2, -3]), "No rain", "rainfall([-2, -3])")
        self.assertEqual(avg_rainfall([4, 0, -3, 20]), "Average: 8.0", "avg_rainfall([4, 0, -3, 20])")
        self.assertEqual(avg_rainfall([20, 4]), "Average: 12.0", "rainfall([20, 4])")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="list_MUC_mix_items" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">mix</span></code> that takes a list <code class="docutils literal notranslate"><span class="pre">list1</span></code> and returns a new list with the first and last items from the original list swapped.  For example, <code class="docutils literal notranslate"><span class="pre">mix([3,</span> <span class="pre">2,</span> <span class="pre">1,</span> <span class="pre">4]</span></code>
should return <code class="docutils literal notranslate"><span class="pre">[4,</span> <span class="pre">2,</span> <span class="pre">1,</span> <span class="pre">3]</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="5"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def mix(list1):
---
    first = list1[0]
    last = list1[-1]
---
    out = [last]
---
    out.extend(list1[1:-1])
---
    out.extend(list1[1:-2]) #paired
---
    out.append(first)
---
    out.extend(first) #paired
---
    return out
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=list_MUC_mix_items_ac data-question_label="6">
<div id=list_MUC_mix_items_ac_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">mix</span></code> that takes a list <code class="docutils literal notranslate"><span class="pre">list1</span></code> and returns a new list with the first and last items from the original list swapped.  For example, <code class="docutils literal notranslate"><span class="pre">mix([3,</span> <span class="pre">2,</span> <span class="pre">1,</span> <span class="pre">4])</span></code>
should return <code class="docutils literal notranslate"><span class="pre">[4,</span> <span class="pre">2,</span> <span class="pre">1,</span> <span class="pre">3]</span></code>.</p>

</div>
<textarea data-lang="python" id="list_MUC_mix_items_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(mix([3, 2, 1, 4]), [4, 2, 1, 3], "mix([3, 2, 1, 4])")
        self.assertEqual(mix(['a', 'b', 'c']), ['c', 'b', 'a'], "mix(['a', 'b', 'c'])")
        self.assertEqual(mix(['a', 'b']), ['b', 'a'], "mix(['a', 'b'])")
        self.assertEqual(mix([2, 2, 1, 1]), [1, 2, 1, 2], "mix([2, 2, 1, 1])")


myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="list_MUC_discount" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">discount(price_list)</span></code> that takes a list of prices and returns a new list with the original price reduced by half.
Limit each price to two digits after the decimal by rounding to the closest value.  Hint, you can use the <code class="docutils literal notranslate"><span class="pre">round</span></code> function to do this.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="7"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def discount(price_list):
---
    discounts = []
---
    for price in price_list:
---
        new_price = round(price / 2, 2)
---
        new_price = price / 2 #paired
---
        discounts.append(new_price)
---
        discounts.extend(new_price) #paired
---
    return discounts
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=list_MUC_discount_ac data-question_label="8">
<div id=list_MUC_discount_ac_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">discount(price_list)</span></code> that takes a list of prices and returns a new list with the original price reduced by half.
Limit each price to two digits after the decimal by rounding to the closest value.  Hint, you can use the <code class="docutils literal notranslate"><span class="pre">round</span></code> function to do this.</p>

</div>
<textarea data-lang="python" id="list_MUC_discount_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(discount([20.00, 15.60, 22.99]), [10.00, 7.80, 11.50], "discount([20.00, 15.60, 22.99])")
        self.assertEqual(discount([14.50]), [7.25], "discount([14.50])")
        self.assertEqual(discount([12.22]), [6.11], "discount([12.22])")
        self.assertEqual(discount([12.31]), [6.16], "discount([12.31])")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="list_MUC_append_if_a" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">filter_a(word_list)</span></code> that takes a list of words, <code class="docutils literal notranslate"><span class="pre">word_list</span></code> and returns a new list with only the words in <code class="docutils literal notranslate"><span class="pre">word_list</span></code> that start with the letter <code class="docutils literal notranslate"><span class="pre">a</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">filter_a(['a',</span> <span class="pre">'bye',</span> <span class="pre">'above'])</span></code> should return <code class="docutils literal notranslate"><span class="pre">['a',</span> <span class="pre">'above']</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="9"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def filter_a(word_list):
---
    res = []
---
    for word in word_list:
---
        if word[0] == 'a':
---
        if word[0] = 'a': #paired
---
            res.append(word)
---
            word_list.append(res) #paired
---
    return res
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=list_MUC_append_if_a_ac data-question_label="10">
<div id=list_MUC_append_if_a_ac_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">filter_a(word_list)</span></code> that takes a list of words, <code class="docutils literal notranslate"><span class="pre">word_list</span></code> and returns a new list with only the words in <code class="docutils literal notranslate"><span class="pre">word_list</span></code> that start with the letter <code class="docutils literal notranslate"><span class="pre">a</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">filter_a(['a',</span> <span class="pre">'bye',</span> <span class="pre">'above'])</span></code> should return <code class="docutils literal notranslate"><span class="pre">['a',</span> <span class="pre">'above']</span></code>.</p>

</div>
<textarea data-lang="python" id="list_MUC_append_if_a_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(filter_a(['a', 'bye', 'above']), ['a', 'above'], "filter_a(['a', 'bye', 'above'])")
        self.assertEqual(filter_a(['bye']), [], "filter_a(['bye'])")
        self.assertEqual(filter_a(['bye', 'a', 'around']), ['a', 'around'], "filter_a(['bye', 'a', 'around'])")
        self.assertEqual(filter_a(['bye', 'above']), ['above'], "filter_a(['bye', 'above'])")



myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="list_MUC_add_vocab" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">vocab_list</span></code> that takes two lists: <code class="docutils literal notranslate"><span class="pre">terms</span></code> and <code class="docutils literal notranslate"><span class="pre">vocab</span></code>. It should loop through the list of terms and add each item to
<code class="docutils literal notranslate"><span class="pre">vocab</span></code> if it is not already in that list. It should return the list <code class="docutils literal notranslate"><span class="pre">vocab</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="11"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def vocab_list(terms, vocab):
---
    for word in terms:
---
        </pre>
        </div>
        </div>
    <div class="system-message">
<p class="system-message-title">System Message: WARNING/2 (<span class="docutils literal">/Users/barbarer/Documents/GitHub/py4e-int/_sources/lists/MixedupCode_orig.rst</span>, line 256)</p>
<p>Explicit markup ends without a blank line; unexpected unindent.</p>
</div>
<dl>
<dt>&lt;&lt;&lt;&lt;&lt;&lt;&lt; HEAD</dt><dd><blockquote>
<div><p>if response not in wordList:</p>
</div></blockquote>
<div class="system-message">
<p class="system-message-title">System Message: WARNING/2 (<span class="docutils literal">/Users/barbarer/Documents/GitHub/py4e-int/_sources/lists/MixedupCode_orig.rst</span>, line 258)</p>
<p>Block quote ends without a blank line; unexpected unindent.</p>
</div>
<div class="system-message">
<p class="system-message-title">System Message: SEVERE/4 (<span class="docutils literal">/Users/barbarer/Documents/GitHub/py4e-int/_sources/lists/MixedupCode_orig.rst</span>, line 258)</p>
<p>Unexpected section title or transition.</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">=====</span>
</pre></div>
</div>
</div>
<blockquote>
<div><p>if response in wordList: #paired</p>
</div></blockquote>
<div class="system-message">
<p class="system-message-title">System Message: WARNING/2 (<span class="docutils literal">/Users/barbarer/Documents/GitHub/py4e-int/_sources/lists/MixedupCode_orig.rst</span>, line 260)</p>
<p>Block quote ends without a blank line; unexpected unindent.</p>
</div>
<div class="system-message">
<p class="system-message-title">System Message: SEVERE/4 (<span class="docutils literal">/Users/barbarer/Documents/GitHub/py4e-int/_sources/lists/MixedupCode_orig.rst</span>, line 260)</p>
<p>Unexpected section title or transition.</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">=====</span>
</pre></div>
</div>
</div>
<blockquote>
<div><p>vocabulary.append(response)</p>
</div></blockquote>
</dd>
</dl>
<div class="system-message">
<p class="system-message-title">System Message: WARNING/2 (<span class="docutils literal">/Users/barbarer/Documents/GitHub/py4e-int/_sources/lists/MixedupCode_orig.rst</span>, line 262)</p>
<p>Definition list ends without a blank line; unexpected unindent.</p>
</div>
<div class="system-message">
<p class="system-message-title">System Message: SEVERE/4 (<span class="docutils literal">/Users/barbarer/Documents/GitHub/py4e-int/_sources/lists/MixedupCode_orig.rst</span>, line 262)</p>
<p>Missing matching underline for section title overline.</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">=======</span>
            <span class="k">if</span> <span class="n">word</span> <span class="ow">not</span> <span class="ow">in</span> <span class="n">vocab</span><span class="p">:</span>
    <span class="o">=====</span>
</pre></div>
</div>
</div>
<blockquote>
<div><blockquote>
<div><p>if word !in vocab: #paired</p>
</div></blockquote>
<div class="system-message">
<p class="system-message-title">System Message: WARNING/2 (<span class="docutils literal">/Users/barbarer/Documents/GitHub/py4e-int/_sources/lists/MixedupCode_orig.rst</span>, line 266)</p>
<p>Block quote ends without a blank line; unexpected unindent.</p>
</div>
<div class="system-message">
<p class="system-message-title">System Message: SEVERE/4 (<span class="docutils literal">/Users/barbarer/Documents/GitHub/py4e-int/_sources/lists/MixedupCode_orig.rst</span>, line 266)</p>
<p>Unexpected section title or transition.</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">=====</span>
</pre></div>
</div>
</div>
<blockquote>
<div><p>vocab.append(word)</p>
</div></blockquote>
</div></blockquote>
<div class="system-message">
<p class="system-message-title">System Message: WARNING/2 (<span class="docutils literal">/Users/barbarer/Documents/GitHub/py4e-int/_sources/lists/MixedupCode_orig.rst</span>, line 268)</p>
<p>Block quote ends without a blank line; unexpected unindent.</p>
</div>
<dl>
<dt>&gt;&gt;&gt;&gt;&gt;&gt;&gt; upstream/master</dt><dd><div class="system-message">
<p class="system-message-title">System Message: SEVERE/4 (<span class="docutils literal">/Users/barbarer/Documents/GitHub/py4e-int/_sources/lists/MixedupCode_orig.rst</span>, line 269)</p>
<p>Unexpected section title or transition.</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">=====</span>
</pre></div>
</div>
</div>
<blockquote>
<div><p>word.append(vocab) #paired</p>
</div></blockquote>
<div class="system-message">
<p class="system-message-title">System Message: WARNING/2 (<span class="docutils literal">/Users/barbarer/Documents/GitHub/py4e-int/_sources/lists/MixedupCode_orig.rst</span>, line 271)</p>
<p>Block quote ends without a blank line; unexpected unindent.</p>
</div>
<div class="system-message">
<p class="system-message-title">System Message: SEVERE/4 (<span class="docutils literal">/Users/barbarer/Documents/GitHub/py4e-int/_sources/lists/MixedupCode_orig.rst</span>, line 271)</p>
<p>Unexpected section title or transition.</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">=====</span>
</pre></div>
</div>
</div>
<blockquote>
<div><p>return vocab</p>
</div></blockquote>
</dd>
</dl>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=list_MUC_add_vocab_ac data-question_label="12">
<div id=list_MUC_add_vocab_ac_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">vocab_list</span></code> that takes two lists: <code class="docutils literal notranslate"><span class="pre">terms</span></code> and <code class="docutils literal notranslate"><span class="pre">vocab</span></code>. It should loop through the list of terms and add each item to
<code class="docutils literal notranslate"><span class="pre">vocab</span></code> if it is not already in that list. It should return the list <code class="docutils literal notranslate"><span class="pre">vocab</span></code>.</p>

</div>
<textarea data-lang="python" id="list_MUC_add_vocab_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(vocab_list(['hi', 'bye'], ['hi', 'bye']), ['hi', 'bye'], "vocab_list(['hi', 'bye'], ['hi', 'bye'])")
        self.assertEqual(vocab_list(['hi', 'bye', 'run'], ['hi', 'bye']), ['hi', 'bye', 'run'], "vocab_list(['hi', 'bye', 'run'], ['hi', 'bye'])")
        self.assertEqual(vocab_list(['hi', 'bye'], []), ['hi', 'bye'], "vocab_list(['hi', 'bye'], [])")
        self.assertEqual(vocab_list(['an', 'anniversary'], ['anniversary']), ['anniversary', 'an'] , "vocab_list(['an', 'anniversary'], ['anniversary'])")


myTests().main()
</textarea>
</div>
</div>
</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>