<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Mixed-up code Questions &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/637.css?v=bddd2011b20db76a20ea" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.css?v=7c5388ad4389eae1a318" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/runtime.bundle.js?v=aa7c717c38c5405af9e6"></script>
    <script src="../_static/637.bundle.js?v=1123ad75f9b821cdd636"></script>
    <script src="../_static/runestone.bundle.js?v=9c8e466e8cbed4a3af5a"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'py4e-int';
    eBookConfig.basecourse = 'py4e-int';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          py4e-int
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=mixedupcode_orig'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
  
<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="mixed-up-code-questions">
<h1>Mixed-up code Questions<a class="headerlink" href="#mixed-up-code-questions" title="Permalink to this headline">¶</a></h1>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="ch5ex1muc" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create the function, <code class="docutils literal notranslate"><span class="pre">list_to_5(n)</span></code>, so that while the integer <code class="docutils literal notranslate"><span class="pre">n</span></code> is less than <code class="docutils literal notranslate"><span class="pre">5</span></code>, add <code class="docutils literal notranslate"><span class="pre">n</span></code> to a list to return and increment <code class="docutils literal notranslate"><span class="pre">n</span></code> by <code class="docutils literal notranslate"><span class="pre">1</span></code>.
Then return the list. For example, if <code class="docutils literal notranslate"><span class="pre">n</span></code> is 0, then the function should return <code class="docutils literal notranslate"><span class="pre">[0,</span> <span class="pre">1,</span> <span class="pre">2,</span> <span class="pre">3,</span> <span class="pre">4]</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="1"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def list_to_5(n):
---
    list = []
---
    while n < 5:
---
    while n > 5: #paired
---
        list.append(n)
---
        n += 1
---
    return list
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ch5ex1muc-ac data-question_label="2">
<div id=ch5ex1muc-ac_question class="ac_question col-md-12">
<p>Write the function, <code class="docutils literal notranslate"><span class="pre">list_to_5(n)</span></code>, so that while the integer <code class="docutils literal notranslate"><span class="pre">n</span></code> is less than 5, add <code class="docutils literal notranslate"><span class="pre">n</span></code> to a list to return and increment <code class="docutils literal notranslate"><span class="pre">n</span></code> by <code class="docutils literal notranslate"><span class="pre">1</span></code>.
Then return the list. For example, if <code class="docutils literal notranslate"><span class="pre">n</span></code> is 0, then the function should return <code class="docutils literal notranslate"><span class="pre">[0,</span> <span class="pre">1,</span> <span class="pre">2,</span> <span class="pre">3,</span> <span class="pre">4]</span></code>.</p>

</div>
<textarea data-lang="python" id="ch5ex1muc-ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def list_to_5(n):
    # write code here

====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(list_to_5(2), [2, 3, 4], 'list_to_5(2)')
        self.assertEqual(list_to_5(6), [], 'list_to_5(6)')
        self.assertEqual(list_to_5(4), [4], 'list_to_5(4)')

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="ch5ex2muc" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create the function, <code class="docutils literal notranslate"><span class="pre">numbers(accum)</span></code>, so that <code class="docutils literal notranslate"><span class="pre">accum</span></code>, with <code class="docutils literal notranslate"><span class="pre">accum</span></code> being any integer less than 6, increases its
value by 2. Then, return <code class="docutils literal notranslate"><span class="pre">accum</span></code>. For example, if <code class="docutils literal notranslate"><span class="pre">accum</span></code> is 0, the function should increase its value by 2 and return 6.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="3"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def numbers(accum):
---
    while accum < 6:
---
    while accum < 7: #paired
---
        accum = accum + 2
---
    return accum
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ch5ex2muc-ac data-question_label="4">
<div id=ch5ex2muc-ac_question class="ac_question col-md-12">
<p>Write the function, <code class="docutils literal notranslate"><span class="pre">numbers(accum)</span></code>, so that <code class="docutils literal notranslate"><span class="pre">accum</span></code>, with <code class="docutils literal notranslate"><span class="pre">accum</span></code> being any integer less than 6, increases its
value by 2. Then, return <code class="docutils literal notranslate"><span class="pre">accum</span></code>. For example, if <code class="docutils literal notranslate"><span class="pre">accum</span></code> is 0, the function should increase its value by 2 and return 6.</p>

</div>
<textarea data-lang="python" id="ch5ex2muc-ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def numbers(accum):
    # write code here

====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(numbers(4), 6, 'numbers(4)')
        self.assertEqual(numbers(6), 6, 'numbers(6)')
        self.assertEqual(numbers(7), 7, 'numbers(7)')
        self.assertEqual(numbers(0), 6, 'numbers(0)')

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="ch5ex3muc" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create the following function, <code class="docutils literal notranslate"><span class="pre">skipCount(start,</span> <span class="pre">increment,</span> <span class="pre">stop)</span></code>, so that <code class="docutils literal notranslate"><span class="pre">start</span></code> is increased by <code class="docutils literal notranslate"><span class="pre">increment</span></code> until it hits
<code class="docutils literal notranslate"><span class="pre">stop</span></code> with <code class="docutils literal notranslate"><span class="pre">start</span></code> being the starting integer, <code class="docutils literal notranslate"><span class="pre">increment</span></code> being the number the integers are incrementing by, and <code class="docutils literal notranslate"><span class="pre">stop</span></code> being the ending integer.
Then return an integer for the last number. For example, if <code class="docutils literal notranslate"><span class="pre">start</span></code> is 0, <code class="docutils literal notranslate"><span class="pre">increment</span></code> is 2, and <code class="docutils literal notranslate"><span class="pre">stop</span></code> is 8, the function should increase by 2 and
then return 8.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="5"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def skipCount(start, increment, stop):
---
    counter = start
---
    while counter < stop:
---
        counter += increment
---
    return counter
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ch5ex3muc-ac data-question_label="6">
<div id=ch5ex3muc-ac_question class="ac_question col-md-12">
<p>Write the following function, <code class="docutils literal notranslate"><span class="pre">skipCount(start,</span> <span class="pre">increment,</span> <span class="pre">stop)</span></code>, so that <code class="docutils literal notranslate"><span class="pre">start</span></code> is increased by <code class="docutils literal notranslate"><span class="pre">increment</span></code> until it hits
<code class="docutils literal notranslate"><span class="pre">stop</span></code> with <code class="docutils literal notranslate"><span class="pre">start</span></code> being the starting integer, <code class="docutils literal notranslate"><span class="pre">increment</span></code> being the number the integers are incrementing by, and <code class="docutils literal notranslate"><span class="pre">stop</span></code> being the ending integer.
Then return an integer for the last number. For example, if <code class="docutils literal notranslate"><span class="pre">start</span></code> is 0, <code class="docutils literal notranslate"><span class="pre">increment</span></code> is 2, and <code class="docutils literal notranslate"><span class="pre">stop</span></code> is 8, the function should increase by 2 and
then return 8.</p>

</div>
<textarea data-lang="python" id="ch5ex3muc-ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def skipCount(start, increment, stop):
    counter = start

    # write code here

====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(skipCount(0,2,8), 8, 'skipCount(0,2,8)')
        self.assertEqual(skipCount(0,2,9), 8, 'skipCount(0,2,9)')
        self.assertEqual(skipCount(0,2,7), 6, 'skipCount(0,2,7)')

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="ch5ex4muc" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create the following function, <code class="docutils literal notranslate"><span class="pre">loopMe(first,</span> <span class="pre">second)</span></code>, to infinitely loop using the variables <code class="docutils literal notranslate"><span class="pre">first</span></code> and <code class="docutils literal notranslate"><span class="pre">second</span></code> with one being greater
than the other and then return a boolean, ‘True’. For example, if <code class="docutils literal notranslate"><span class="pre">first</span></code> is 4 and <code class="docutils literal notranslate"><span class="pre">second</span></code> is 3, then the function should infintely loop
and return ‘True’.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="7"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def loopMe(first, second):
---
    while first > second:
---
    while first == second: #paired
---
        return True
---
loopMe(7, 4)
---
loopMe(1, 3) #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ch5ex4muc-ac data-question_label="8">
<div id=ch5ex4muc-ac_question class="ac_question col-md-12">
<p>Write the following function, <code class="docutils literal notranslate"><span class="pre">loopMe(first,</span> <span class="pre">second)</span></code>, to infinitely loop using the variables <code class="docutils literal notranslate"><span class="pre">first</span></code> and <code class="docutils literal notranslate"><span class="pre">second</span></code> with one being greater
than the other and then return a boolean, ‘True’. For example, if <code class="docutils literal notranslate"><span class="pre">first</span></code> is 4 and <code class="docutils literal notranslate"><span class="pre">second</span></code> is 3, then the function should infintely loop
and return ‘True’.</p>

</div>
<textarea data-lang="python" id="ch5ex4muc-ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def loopMe(first,second):
    # write code here

====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(loopMe(4,3), 'True', 'loopMe(4,3)')

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="ch5ex5muc" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create the following function, <code class="docutils literal notranslate"><span class="pre">calculation(numbers)</span></code>, to return a float of the average of the list of numbers by taking a list <code class="docutils literal notranslate"><span class="pre">numbers</span></code>,
summing the numbers, and then dividing the sum by the length of the list, using a for loop. For example, if <code class="docutils literal notranslate"><span class="pre">numbers</span></code> is [90, 94, 85, 78, 87, 98], then the
function should return the float 88.66666666666667.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="9"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def calculation(numbers):
---
    sum = 0
---
    for number in numbers:
---
        sum = sum + number
---
        sum = sum - number #paired
---
    return sum/len(numbers)
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ch5ex5muc-ac data-question_label="10">
<div id=ch5ex5muc-ac_question class="ac_question col-md-12">
<p>Write the following function, <code class="docutils literal notranslate"><span class="pre">calculation(numbers)</span></code>, to return a float of the average of the list of numbers by taking a list <code class="docutils literal notranslate"><span class="pre">numbers</span></code>,
summing the numbers, and then dividing the sum by the length of the list, using a for loop. For example, if <code class="docutils literal notranslate"><span class="pre">numbers</span></code> is [90, 94, 85, 78, 87, 98], then the
function should return the float 88.66666666666667.</p>

</div>
<textarea data-lang="python" id="ch5ex5muc-ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def calculation(numbers):
    sum = 0
    # write code here

====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertAlmostEqual(calculation([90, 94, 85, 78, 87, 98]), 88.66666666666667, 2, 'calculation([90, 94, 85, 78, 87, 98])')
        self.assertAlmostEqual(calculation([65, 69, 84, 99, 88, 98]), 83.8333333333333, 2, 'calculation([65, 69, 84, 99, 88, 98])')

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="ch5ex6muc" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create the following function, <code class="docutils literal notranslate"><span class="pre">sum_of_range(start,</span> <span class="pre">end)</span></code>, to return an integer for the sum of all numbers between two numbers of your choosing using by using the
variables <code class="docutils literal notranslate"><span class="pre">start</span></code> followed by <code class="docutils literal notranslate"><span class="pre">end</span></code>. For example, if <code class="docutils literal notranslate"><span class="pre">start</span></code> is 0 and <code class="docutils literal notranslate"><span class="pre">end</span></code> is 5, the returned sum should be (0 + 1 + 2 + 3 + 4 + 5) which is 15.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="11"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def sum_of_range(start, end):
---
    sum = 0
---
    for number in range(start, end):
---
        sum = sum + number
---
        sum = sum + numbers #paired
---
    return(sum)
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ch5ex6muc-ac data-question_label="12">
<div id=ch5ex6muc-ac_question class="ac_question col-md-12">
<p>Write the following function, <code class="docutils literal notranslate"><span class="pre">sum_of_range(start,</span> <span class="pre">end)</span></code>, to return an integer for the sum of all numbers between two numbers of your choosing using by using the
variables <code class="docutils literal notranslate"><span class="pre">start</span></code> followed by <code class="docutils literal notranslate"><span class="pre">end</span></code>. For example, if <code class="docutils literal notranslate"><span class="pre">start</span></code> is 0 and <code class="docutils literal notranslate"><span class="pre">end</span></code> is 5, the returned sum should be (0 + 1 + 2 + 3 + 4 + 5) which is 15.</p>

</div>
<textarea data-lang="python" id="ch5ex6muc-ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def sum_of_range(start, end):

    # write code here
====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(sum_of_range(0,5), 15, 'sum_of_range(0,5)')
        self.assertEqual(sum_of_range(2,5), 14, 'sum_of_range(2,5)')
        self.assertEqual(sum_of_range(0,31), 465, 'sum_of_range(0,31)')
        self.assertEqual(sum_of_range(0,20), 190, 'sum_of_range(0,20)')
        self.assertEqual(sum_of_range(0,101), 5050, 'sum_of_range(0,101)')

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="ch5ex7muc" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create the following function, <code class="docutils literal notranslate"><span class="pre">odd_sum(start,</span> <span class="pre">increment,</span> <span class="pre">end)</span></code>, to return the sum of all odd numbers between two numbers of your choosing
using the variables <code class="docutils literal notranslate"><span class="pre">start</span></code>, <code class="docutils literal notranslate"><span class="pre">increment</span></code>, and <code class="docutils literal notranslate"><span class="pre">end</span></code>. For example, if <code class="docutils literal notranslate"><span class="pre">start</span></code> is 1, <code class="docutils literal notranslate"><span class="pre">increment</span></code> is 2, and <code class="docutils literal notranslate"><span class="pre">end</span></code> is 6, the returned sum
should be (1 + 3 + 5) which is 9.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="13"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def odd_sum(start, increment, end):
---
    sum = 0
---
    for number in range(start, end, increment):
---
    for numbers in range(start, increment, end): #paired
---
        sum = sum + number
---
    return sum
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ch5ex7muc-ac data-question_label="14">
<div id=ch5ex7muc-ac_question class="ac_question col-md-12">
<p>Write the following function, <code class="docutils literal notranslate"><span class="pre">odd_sum(start,</span> <span class="pre">increment,</span> <span class="pre">end)</span></code>, to return the sum of all odd numbers between two numbers of your choosing
using the variables <code class="docutils literal notranslate"><span class="pre">start</span></code>, <code class="docutils literal notranslate"><span class="pre">increment</span></code>, and <code class="docutils literal notranslate"><span class="pre">end</span></code>. For example, if <code class="docutils literal notranslate"><span class="pre">start</span></code> is 1, <code class="docutils literal notranslate"><span class="pre">increment</span></code> is 2, and <code class="docutils literal notranslate"><span class="pre">end</span></code> is 5, the returned sum
should be (1 + 3 + 5) which is 9</p>

</div>
<textarea data-lang="python" id="ch5ex7muc-ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def odd_sum(start, increment, end):
    sum = 0
    # write code here
====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(odd_sum(1, 2, 5), 9, 'odd_sum(1, 2, 5)')
        self.assertEqual(odd_sum(1, 2, 30), 225, 'odd_sum(1, 2, 30)')
        self.assertEqual(odd_sum(2, 3, 100), 1650, 'odd_sum(2, 3, 100)')
        self.assertEqual(odd_sum(-1, 4, 30), 104, 'odd_sum(-1, 4, 30)')

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="ch5ex8muc" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create the following function, <code class="docutils literal notranslate"><span class="pre">multiple_three(num)</span></code>, to return the sum of every multiple of 3 between 3 and and a number of your choosing using <code class="docutils literal notranslate"><span class="pre">num</span></code> and
return an integer for the sum. For example, if <code class="docutils literal notranslate"><span class="pre">num</span></code> is 37, the function should return <code class="docutils literal notranslate"><span class="pre">234</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="15"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def multiple_three(num):
---
    sum = 0
    numbers = range(3, num, 3)
---
    for number in numbers:
---
        sum = sum + number
---
        sum = sum - number #paired
---
    return sum
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ch5ex8muc-ac data-question_label="16">
<div id=ch5ex8muc-ac_question class="ac_question col-md-12">
<p>Write the following function, <code class="docutils literal notranslate"><span class="pre">multiple_three(num)</span></code>, to return the sum of every multiple of 3 between 3 and and a number of your choosing using <code class="docutils literal notranslate"><span class="pre">num</span></code> and
return an integer for the sum. For example, if <code class="docutils literal notranslate"><span class="pre">num</span></code> is 37, the function should return <code class="docutils literal notranslate"><span class="pre">234</span></code>.</p>

</div>
<textarea data-lang="python" id="ch5ex8muc-ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def multiple_three(num):
    sum = 0
    # write code here
====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(multiple_three(37), 234, 'multiple_three(37)')
        self.assertEqual(multiple_three(101), 1683, 'multiple_three(101)')
        self.assertEqual(multiple_three(19), 63, 'multiple_three(19)')

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="ch5ex9muc" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create the following function, <code class="docutils literal notranslate"><span class="pre">water_quality(pHvalues)</span></code>, to return a float for the average pH of water samples by using a list
<code class="docutils literal notranslate"><span class="pre">pHvalues</span></code>. For example, if <code class="docutils literal notranslate"><span class="pre">pHvalues</span></code> is [7.0, 8.2, 6.7, 7.5, 8.0, 7.2], then the function should return 7.433333333333334.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="17"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def water_quality(pHvalues):
---
    total = 0
---
    for number in pHvalues:
---
        total = total + number
---
        total = total * number #paired
---
    average = total / len(pHvalues)
---
    return average
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ch5ex9muc-ac data-question_label="18">
<div id=ch5ex9muc-ac_question class="ac_question col-md-12">
<p>Write the following function, <code class="docutils literal notranslate"><span class="pre">water_quality(pHvalues)</span></code>, to return a float for the average pH of water samples by using a list
<code class="docutils literal notranslate"><span class="pre">pHvalues</span></code>. For example, if <code class="docutils literal notranslate"><span class="pre">pHvalues</span></code> is [7.0, 8.2, 6.7, 7.5, 8.0, 7.2], then the function should return 7.433333333333334.</p>

</div>
<textarea data-lang="python" id="ch5ex9muc-ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def water_quality(pHvalues):
    total = 0
    # write code here
====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertAlmostEqual(water_quality([7.0, 8.2, 6.7, 7.5, 8.0, 7.2]), 7.433333333333334, 2, 'water_quality([7.0, 8.2, 6.7, 7.5, 8.0, 7.2], 6)')
        self.assertAlmostEqual(water_quality([7, 8, 6, 7, 8, 7]), 7.166666666667, 2, 'water_quality([7, 8, 6, 7, 8, 7])')


myTests().main()
</textarea>
</div>
</div>

        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="ch5ex10muc" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create the following function, <code class="docutils literal notranslate"><span class="pre">countdown(counter)</span></code>, to start with a number of your choosing using the variable
<code class="docutils literal notranslate"><span class="pre">counter</span></code>, decrement it, and then return the last number. For example, if <code class="docutils literal notranslate"><span class="pre">counter</span></code> is 5, then the function
should countdown from 5 and return 0.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="19"   data-adaptive="true"      data-numbered="left"  style="visibility: hidden;">
        def countdown(counter):
---
    while counter > 0:
---
        counter = counter - 1
---
        counter = counter + 1 #paired
---
    return counter
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ch5ex10muc-ac data-question_label="20">
<div id=ch5ex10muc-ac_question class="ac_question col-md-12">
<p>Write the following function, <code class="docutils literal notranslate"><span class="pre">countdown(counter)</span></code>, to start with a number of your choosing using the variable
<code class="docutils literal notranslate"><span class="pre">counter</span></code>, decrement it, and then return the last number. For example, if <code class="docutils literal notranslate"><span class="pre">counter</span></code> is 5, then the function
should countdown from 5 and return 0.</p>

</div>
<textarea data-lang="python" id="ch5ex10muc-ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def countdown(counter):
    # write code here
====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(countdown(5), 0, 'countdown(5)')
        self.assertEqual(countdown(3), 0, 'countdown(3)')
        self.assertEqual(countdown(7), 0, 'countdown(7)')

myTests().main()
</textarea>
</div>
</div>
</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>