<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <title>5.11. Write Code Questions &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=DB047029" />
    <link rel="stylesheet" type="text/css" href="https://cdn.jsdelivr.net/npm/handsontable@7.2.2/dist/handsontable.full.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-3.4.1/css/bootstrap.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/presenter_mode.css?v=9FF1320F" />
    <link rel="stylesheet" type="text/css" href="../_static/jquery-ui-1.10.3.custom.min.css?v=97E0178A" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/user-highlights.css?v=A2326E52" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone-custom-sphinx-bootstrap.css?v=5.0.8" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/language_data.js"></script>
    <script src="../_static/jquery.highlight.js?v=F1C496C8"></script>
    <script src="../_static/sql-wasm.js?v=44FE0709"></script>
    <script src="https://cdn.jsdelivr.net/npm/handsontable@7.2.2/dist/handsontable.full.js"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/pytutor-embed.bundle.js?v=268B5C04"></script>
    <script src="../_static/html4css1.css?v=4D8F3FB4"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/Blob.js?v=AEC2447"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.emitter.bidi.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.emitter.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.fallbacks.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.messagestore.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.parser.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.language.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega@4.0.0-rc.2/build/vega.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega-lite@2.5.0/build/vega-lite.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega-embed@3.14.0/build/vega-embed.js"></script>
    <script src="../_static/runestone.js?v=3EAD674C"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js?v=2B07FE4F"></script>
    <script src="../_static/bootstrap-3.4.1/js/bootstrap.min.js"></script>
    <script src="../_static/jquery-fix.js?v=1226591B"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/jquery.idle-timer.js?v=E942DAF1"></script>
    <script src="../_static/presenter_mode.js?v=2DF16FB4"></script>
    <script src="../_static/theme.js?v=D899FB97"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="6. Strings" href="../strings/toctree.html" />
    <link rel="prev" title="5.10. Mixed-up code Questions" href="mixedupcode.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'py4e-int';
    eBookConfig.basecourse = 'py4e-int';
    eBookConfig.isLoggedIn = false;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;


</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar" aria-label="navbar toggle" class="visuallyhidden">toggle navbar</span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
          <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          py4e-int
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span aria-label="Social" class="visuallyhidden">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span aria-label="Search" class="visuallyhidden">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span aria-label="User" class="visuallyhidden">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:ACFactory.toggleScratchActivecode()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span aria-label="Scratch Activecode" class="visuallyhidden">Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span aria-label="Help" class="visuallyhidden">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=WriteCode'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">5.10. </span>Mixed-up code Questions" data-toggle="tooltip">
    <a href="mixedupcode.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">6. </span>Strings' data-toggle="tooltip" >              
    <a href="../strings/toctree.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

      </div> <!-- navbar -->







    </div> <!-- container -->
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
  
  <div class="section" id="write-code-questions">
<h1><span class="section-number">5.11. </span>Write Code Questions<a class="headerlink" href="#write-code-questions" title="Permalink to this headline">¶</a></h1>
<div class="admonition warning">
<p class="admonition-title">Warning</p>
<p>Be careful not to create infinite loops, as they will cause the page to freeze.</p>
</div>
<ol class="arabic">
<li><div id='itr-ex-countdown' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<p>Fix the 5 syntax errors in the code below to print a countdown of the numbers from 10 to 0.</p>

<div data-childcomponent="itr-ex-countdownq" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=itr-ex-countdownq data-lang="python" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static>
def countdown():
    counter = 10
    while Counter > 10:
        Print(counter)
        counter = counter + 1

countdown()

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(recPerimeter(10, 20),60,"Tested recPerimeter on inputs 10 and 20")
        self.assertEqual(recPerimeter(1, 2),6,"Tested recPerimeter on inputs 1 and 2")
        self.assertEqual(recPerimeter(23.2, 12.3),71,"Tested recPerimeter on inputs 23.2 and 12.3")
        self.assertEqual(recPerimeter(3.0, 2),10.0,"Tested recPerimeter on inputs 3.0 and 2")

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="itr-ex-countdowna" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=itr-ex-countdowna data-lang="python" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static>
def countdown():
    counter = 10
    while counter >= 0:
        print(counter)
        counter = counter - 1

countdown()
</textarea>
</div>
</div>
    </div>

</li>
<li>
<div data-childcomponent="itr-ex-100Loopq" class="runestone explainer ac_section alert alert-warning">
<p>The following code will loop too many times. Change one line that
will to make it loop only 5 times.</p>

<textarea data-component="activecode" id=itr-ex-100Loopq data-lang="python" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static>
x = 5
while x < 100:
    print(x)
    x = x + 1
</textarea>
</div>
</li>
<li><div id='itr-ex-count' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<p>Make 5 changes to the code below to correctly print a count up from -10 to 0.</p>

<div data-childcomponent="itr-ex-countq" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=itr-ex-countq data-lang="python" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static>
output = ""
x = -10
while x < 0
    x = x - 1
output = output + str(x) + " "
print(output)

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(output,"-10 -9 -8 -7 -6 -5 -4 -3 -2 -1 0 ")

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="itr-ex-counta" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=itr-ex-counta data-lang="python" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static>
output = ""
x = -11
while x < 0:
    x = x + 1
    output = output + str(x) + " "
print(output)
</textarea>
</div>
</div>
    </div>

</li>
<li>
<div data-childcomponent="itr-ex-allValuesq" class="runestone explainer ac_section alert alert-warning">
<p>Finish lines 1 and 5 so that the following code correct prints all the values from -5 to -1.</p>

<textarea data-component="activecode" id=itr-ex-allValuesq data-lang="python" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static>
output =
x = -5
while x < 0:
    output = output + str(x) + " "
    x =
print(output)

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(output,"-5 -4 -3 -2 -1 ")

myTests().main()
</textarea>
</div>
</li>
<li><div id='itr-ex-six' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<p>Complete the code on lines 4 and 6 so that it prints the number 6.</p>

<div data-childcomponent="itr-ex-sixq" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=itr-ex-sixq data-lang="python" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static>
x = 3
i = 0
while i < 3:
    x =
    i = i + 1
print()

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(x,6)

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="itr-ex-sixa" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=itr-ex-sixa data-lang="python" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static>
x = 3
i = 0
while i < 3:
    x = x + 1
    i = i + 1
print(x)
</textarea>
</div>
</div>
    </div>

</li>
<li>
<div data-childcomponent="itr-ex-rootq" class="runestone explainer ac_section alert alert-warning">
<p>The code below is supposed to print an estimate of the square root. But,
the indention is wrong on 4 lines. Fix these errors to find the estimate.</p>

<textarea data-component="activecode" id=itr-ex-rootq data-lang="python" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static>
target = 6
    guess = 2
guessSquared = guess * guess
while abs(target-guessSquared) > 0.01:
    closer = target / guess
guess = (guess + closer) / 2.0
        guessSquared = guess * guess
    print("Square root of", target, "is", guess)
</textarea>
</div>
</li>
<li><div id='int-ex-inclusive' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<p>The function currently takes a start and stop argument and uses a for
loop to find the sum of all the numbers between them (inclusive). Change
the for loop to a while loop while still using the parameters.</p>

<div data-childcomponent="int-ex-inclusiveq" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=int-ex-inclusiveq data-lang="python" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static>
def sumFunc(start, stop):
    sum = 0
    for num in range(start, stop + 1):
        sum = sum + num
    return sum

print(sumFunc(1,10))

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sumFunc(1, 10),55,"Tested sumFunc on inputs 1 and 10")
        self.assertEqual(sumFunc(10, 3),0,"Tested sumFunc on inputs 10 and 3")
        self.assertEqual(sumFunc(-5, 0),-15,"Tested sumFunc on inputs 20 and 50")
        self.assertEqual(sumFunc(-3, 12),72,"Tested sumFunc on inputs -3 and 12")

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="int-ex-inclusivea" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=int-ex-inclusivea data-lang="python" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static>
def sumFunc(start, stop):
    sum = 0
    num = start
    while num <= stop:
        sum = sum + num
        num += 1
    return sum

print(sumFunc(1,10))
</textarea>
</div>
</div>
    </div>

</li>
<li>
<div data-childcomponent="int-ex-timesq" class="runestone explainer ac_section alert alert-warning">
<p>The program below is supposed to print the times tables for 1 to 3, but
there are 5 errors.  Fix the errors.</p>

<textarea data-component="activecode" id=int-ex-timesq data-lang="python" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static>
for x in range(1,3):
    for y in range(1,10)
        print(str(x) + " * " str(y) + " = " x*y)
</textarea>
</div>
</li>
<li><div id='int-ex-timesWhile' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<p>The following code prints the times tables for 1 to 3 using two for loops.
Rewrite the program to use a while and a for loop instead.</p>

<div data-childcomponent="int-ex-timesWhileq" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=int-ex-timesWhileq data-lang="python" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static>
for x in range(1,4):
    for y in range(1,11):
        print(str(x) + " * " + str(y) + " = " + str(x*y))
</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="int-ex-timesWhileA" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=int-ex-timesWhileA data-lang="python" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static>
# There are a few different ways this can be done
# One is shown here

for x in range(1,4):
    y = 1
    while y < 11:
        print(str(x) + " * " + str(y) + " = " + str(x*y))
        y += 1
</textarea>
</div>
</div>
    </div>

</li>
<li>
<div data-childcomponent="int-ex-whileq" class="runestone explainer ac_section alert alert-warning">
<p>Rewrite the following code to use a while loop instead of a for loop.</p>

<textarea data-component="activecode" id=int-ex-whileq data-lang="python" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static>
product = 1  # Start out with nothing
numbers = range(1,11)
for number in numbers:
    product = product * number
print(product)

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(product, 3628800)

myTests().main()
</textarea>
</div>
</li>
</ol>
</div>


  <div id=scprogresscontainer>
    You have attempted <span id=scprogresstotal></span> of <span id=scprogressposs></span> activities on this page <div id="subchapterprogress"></div>
  </div>
  
      
<ul>
  <li id="relations-prev" style="opacity:0.4" onmouseover="mouseOverPrev()" onmouseout="mouseOutPrev()"class="navLink" title='Previous Section - 5.10. Mixed-up code Questions' data-toggle="tooltip">
    <a href="mixedupcode.html" aria-label="5.10. Mixed-up code Questions">
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">
            <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
        </div>
    </a>
  </li>
</ul>

  
<ul>
  <li id="relations-next" style="opacity:0.4" onmouseover="mouseOverNext()" onmouseout="mouseOutNext()" class="navLink" title='Next Section - 6. Strings' data-toggle="tooltip" >
    <a href="../strings/toctree.html" aria-label="6. Strings">
      <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">
          <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
      </div>
    </a>
  </li>
</ul>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});

</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>




<script> getSwitch(); </script>


  </body>
</html>