<?xml version="1.0"?>
<section xml:id="iterations_write-code-questions">
  <title>Write Code Questions</title>
  <warning>
    <p>Be careful not to create infinite loops, as they will cause the page to freeze.</p>
  </warning>
  <exercise>
    <statement>
      <p>Fix the 5 errors in the code below to return a countdown of the numbers from 10 to 0, including 0.</p>
    </statement>
    <program interactive="activecode" language="python" xml:id="itr-ex-countdownq_editor">
      <code>
        def countdown():
            counter = 10
            answer = ""
            while Counter &gt; 0:
                answer = answer + str(counter) + " "
                counter = counter + 1
            Return answer
        countDown()
      </code>

      
      <tests>
        from unittest.gui import TestCaseGui

        class myTests(TestCaseGui):
            
            def testOne(self):
                output = countdown()
                self.assertEqual(output,"10 9 8 7 6 5 4 3 2 1 0 ")
        myTests().main()
      </tests>
    </program>
    <solution>
      <p>Fix the 5 errors in the code below to return a countdown of the numbers from 10 to 0, including 0.</p>
      <program xml:id="itr-ex-countdowna_editor" language="python">
        <code>
def countdown():
    counter = 10
    answer = ""

    # Keep running loop until counter equals 0 (use &lt;=)
    # Use correct variable name (counter is lowercase)
    while counter &gt;= 0:
        answer = answer + str(counter) + " "

        # Decrement to lower the counter
        counter = counter - 1

    # Keyword return is lowercase
    return answer

# Call correct function name (countdown)
countdown()
        </code>
      </program>
    </solution>
  </exercise>
  


  <exercise label="itr-ex-100Loopq">
    <statement>
      <p>The following code will loop way too many times. Change one line to
            make the loop only have five iterations.</p>
    </statement>
    <program xml:id="itr-ex-100Loopq_editor" interactive="activecode" language="python">
      <code>
x = 5
while x &lt; 100:
    print(x)
    x = x + 1
        </code>
    </program>
  </exercise>
  <exercise>
    <statement>
      <p>Make 5 changes to the code below to correctly print a count up from -10 to 0, including 0.</p>
    </statement>
    <program interactive="activecode" language="python" xml:id="itr-ex-countq_editor">
      <code>
output = ""
x = -10
while x &lt; 0
    x = x - 1
output = output + str(x) + " "
print(output

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(output,"-10 -9 -8 -7 -6 -5 -4 -3 -2 -1 0 ")
myTests().main()
        </code>
    </program>
    <solution>


    
                Make 5 changes to the code below to correctly print a count up from -10 to 0, including 0.

    
    <program xml:id="itr-ex-counta_editor" language="python"><code>
output = ""
# Start x at -11 so it stays under 0
x = -11

# First line of a loop ends with a colon (:)
while x &lt; 0:

    # Since the iteration variable is negative, increase the count
    x = x + 1

    # Output reassignment is within the loop
    output = output + str(x) + " "

# Close print parentheses
print(output)

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(output,"-10 -9 -8 -7 -6 -5 -4 -3 -2 -1 0 ")
myTests().main()
        </code></program>
            </solution>
  </exercise>
  <exercise label="itr-ex-allValuesq">
    <statement>
      <p>Finish lines 1 and 5 so that the following code correctly prints every integer from -5 to -1, including -1.</p>
    </statement>
    <program xml:id="itr-ex-allValuesq_editor" interactive="activecode" language="python">
      <code>
output =
x = -5
while x &lt; 0:
    output = output + str(x) + " "
    x =
print(output)

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(output,"-5 -4 -3 -2 -1 ")
myTests().main()
        </code>
    </program>
  </exercise>
  <exercise>
    <statement>
      <p>Complete the code on lines 4 and 6 so that it prints the number 6.</p>
    </statement>
    <program interactive="activecode" language="python" xml:id="itr-ex-sixq_editor">
      <code>
x = 3
i = 0
while i &lt; 3:
    x =
    i = i + 1
print()

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(x,6)
myTests().main()
        </code>
    </program>
    <solution>


    
                Complete the code on lines 4 and 6 so that it prints the number 6.

    
    <program xml:id="itr-ex-sixa_editor" language="python"><code>
x = 3
i = 0
while i &lt; 3:
    # Increase x by 1 for each run of the loop
    x = x + 1
    i = i + 1
# Print the x variable
print(x)

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(x,6)
myTests().main()
        </code></program>

            </solution>
  </exercise>

  <exercise label="itr-ex-root_ac" xml:id="itr-ex-root_ac">
    <statement>
      <p>The code below is supposed to print an estimate of the square root of a number (target). You are given an initial guess of 2.
            However, the indention is wrong on 4 lines. Fix these errors to find the estimate.</p>
    </statement>
    <program xml:id="itr-ex-root_ac_editor" interactive="activecode" language="python">
      <code>
target = 6
    guess = 2
guessSquared = guess * guess
while abs(target-guessSquared) &gt; 0.01:
    closer = target / guess
guess = (guess + closer) / 2.0
        guessSquared = guess * guess
    print("Square root of", target, "is", guess)
        </code>
    </program>
  </exercise>
  <exercise label="itr-ex-root_pp" xml:id="itr-ex-root_pp">
    <statement>
      <p>Create code to print an estimate of the square root of a number (target). You are given an initial guess of 2.</p>
    </statement>
    <blocks numbered="left" adaptive="yes" indentation="hide" language="python">
      <block order="6">
        <cline>target = 6</cline>
        <cline>guess = 2</cline>
      </block>
      <block order="2">
        <cline>guessSquared = guess * guess</cline>
      </block>
      <block order="7">
        <cline>while abs(target-guessSquared) &gt; 0.01:</cline>
      </block>
      <block order="3">
        <cline>    closer = target / guess</cline>
      </block>
      <block order="4">
        <cline>    guess = (guess + closer) / 2.0</cline>
      </block>
      <block order="1">
        <cline>    guessSquared = guess * guess</cline>
      </block>
      <block order="5">
        <cline>print("Square root of", target, "is", guess)</cline>
      </block>
    </blocks>
  </exercise>

  <exercise label="iterations_writecode6">
      <select grade="first" questions="itr-ex-root_ac itr-ex-root_pp"/>
  </exercise>

  <exercise>
    <statement>
      <p>This function currently takes a start and stop argument and uses a for loop to find the sum of all the numbers between them (inclusive).
                    Change the for loop to a while loop while still using the same parameters.</p>
    </statement>
    <program interactive="activecode" language="python" xml:id="int-ex-inclusiveq_editor">
      <code>
def sumFunc(start, stop):
    sum = 0
    for num in range(start, stop + 1):
        sum = sum + num
    return sum
print(sumFunc(1,10))

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(sumFunc(1, 10),55,"Tested sumFunc on inputs 1 and 10")
        self.assertEqual(sumFunc(10, 3),0,"Tested sumFunc on inputs 10 and 3")
        self.assertEqual(sumFunc(-5, 0),-15,"Tested sumFunc on inputs 20 and 50")
        self.assertEqual(sumFunc(-3, 12),72,"Tested sumFunc on inputs -3 and 12")
myTests().main()
        </code>
    </program>
    <solution>


    
                This function currently takes a start and stop argument and uses a for loop to find the sum of all the numbers between them (inclusive).
                    Change the for loop to a while loop while still using the same parameters.

    
    <program xml:id="int-ex-inclusivea_editor" language="python"><code>
def sumFunc(start, stop):
    sum = 0
    # Create an iteration variable, initialized to the start argument
    num = start
    # Use while loop until iteration variable is less than
    # or equal to stop argument
    while num &lt;= stop:
        # Add number to sum
        sum = sum + num
        # Increase iteration variable
        num += 1
    # Return the sum
    return sum

print(sumFunc(1,10))

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(sumFunc(1, 10),55,"Tested sumFunc on inputs 1 and 10")
        self.assertEqual(sumFunc(10, 3),0,"Tested sumFunc on inputs 10 and 3")
        self.assertEqual(sumFunc(-5, 0),-15,"Tested sumFunc on inputs 20 and 50")
        self.assertEqual(sumFunc(-3, 12),72,"Tested sumFunc on inputs -3 and 12")
myTests().main()
        </code></program>
            </solution>
  </exercise>
  <exercise label="int-ex-timesq">
    <statement>
      <p>The program below is supposed to print the times tables from 1 to 3, but
            there are 6 errors.  Fix the errors.</p>
    </statement>
    <program xml:id="int-ex-timesq_editor" interactive="activecode" language="python">
      <code>
for x in range(1, 3):
    for y in range(1, 4)
        print(str(x) + " * " str(y) + " = " x * y)
        </code>
    </program>
  </exercise>
  <exercise>
    <statement>
      <p>The following code prints the times tables for 1 to 3 using two for loops.
                    Rewrite the program to use a while and a for loop instead.</p>
    </statement>
    <program interactive="activecode" language="python" xml:id="int-ex-timesWhileq_editor">
      <code>

for x in range(1, 4):
    for y in range(1, 11):
        print(str(x) + " * " + str(y) + " = " + str(x * y))
        </code>
    </program>
    <solution>
      <program xml:id="int-ex-timesWhileA" language="python">
        <code>
# There are a few different ways this can be done
# One is shown here

for x in range(1, 4):
    # Create an iteration variable, starting in the range
    y = 1
    # Use while loop if the iteration variable is less than 11
    while y &lt; 11:
        # Print the string
        print(str(x) + " * " + str(y) + " = " + str(x * y))
        # Increment the iteration variable
        y = y + 1
        </code>
      </program>
    </solution>
  </exercise>

  <exercise label="int-ex-while_ac" xml:id="int-ex-while_ac">
    <statement>
      <p>Rewrite the following code to use a while loop instead of a for loop to calculate the value of ten factorial.</p>
    </statement>
    <program xml:id="int-ex-while_ac_editor" interactive="activecode" language="python">
      <code>
product = 1  # Start out with nothing
numbers = range(1,11)
for number in numbers:
    product = product * number
print(product)

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(product, 3628800)

myTests().main()
        </code>
    </program>
  </exercise>

  <exercise label="int-ex-while_pp" xml:id="int-ex-while_pp">
    <statement>
      <p>Using the blocks below, create code to calculate the value of ten factorial.</p>
    </statement>
    <blocks numbered="left" adaptive="yes" indentation="hide" language="python">
      <block order="5">
        <cline>product = 1</cline>
        <cline>number = 1</cline>
      </block>
      <block order="4">
        <cline>while number &lt; 11:</cline>
      </block>
      <block order="1">
        <choice correct="yes">
          <cline>    product = product * number</cline>
        </choice>
        <choice>
          <cline>    product = number * number </cline>
        </choice>
      </block>
      <block order="6">
        <cline>    number = number + 1</cline>
      </block>
      <block order="2">
        <cline>print(product)</cline>
      </block>
    </blocks>
  </exercise>

  <exercise label="iterations_writecode10">
      <select grade="first" questions="int-ex-while_ac int-ex-while_pp"/>
  </exercise>

</section>