<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>11.13. Write Code Questions &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />
    <link rel="stylesheet" type="text/css" href="../_static/637.css?v=bddd2011b20db76a20ea" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.css?v=5b1dc1bb539ccf7c69f7" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/runtime.bundle.js?v=f3c8dd9e2667ee8dc613"></script>
    <script src="../_static/637.bundle.js?v=1123ad75f9b821cdd636"></script>
    <script src="../_static/runestone.bundle.js?v=ef2d58494168454a1d6a"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="12. Regular Expressions" href="../regex/toctree.html" />
    <link rel="prev" title="11.12. Mixed-Up Code Questions" href="mixedupcode.html" />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'py4e-int';
    eBookConfig.basecourse = 'py4e-int';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar" aria-label="navbar toggle">toggle navbar</span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
          <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          py4e-int
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=writecode'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">11.12. </span>Mixed-Up Code Questions" data-toggle="tooltip">
    <a href="mixedupcode.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">12. </span>Regular Expressions' data-toggle="tooltip" >              
    <a href="../regex/toctree.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  document.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="write-code-questions">
<h1><span class="section-number">11.13. </span>Write Code Questions<a class="headerlink" href="#write-code-questions" title="Permalink to this headline">¶</a></h1>
<ol class="arabic">
<li><div id='tup_writecode1' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<p>Create a function called <code class="docutils literal notranslate"><span class="pre">tup_creation</span></code> that returns a tuple named <code class="docutils literal notranslate"><span class="pre">tup1</span></code> that has three elements: ‘a’, ‘b’, and ‘c’.</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=tup_writecode1q data-question_label="11.13.1">
<div id=tup_writecode1q_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="tup_writecode1q_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def tup_creation():
    # write code here

====
from unittest.gui import TestCaseGui

class MyTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(tup_creation(), ('a', 'b', 'c'), "Checking to make tup_creation() returns a tuple with the correct elements.")

MyTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=tup_writecode1a data-question_label="11.13.2">
<div id=tup_writecode1a_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="tup_writecode1a_editor" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
# Parentheses indicate a tuple
def tup_creation():
    tup1 = ("a","b","c")
    return tup1
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=tup_writecode2 data-question_label="11.13.3">
<div id=tup_writecode2_question class="ac_question col-md-12">
<p>Provided is a list of tuples. Create a function called <code class="docutils literal notranslate"><span class="pre">check_third_element</span></code> that returns a new list
called <code class="docutils literal notranslate"><span class="pre">t_check</span></code> that contains the third element of each tuple.</p>

</div>
<textarea data-lang="python" id="tup_writecode2_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

def check_third_element():
    lst_tups = [('Articuno', 'Moltres', 'Zapdos'), ('Beedrill', 'Metapod', 'Charizard', 'Venasaur', 'Squirtle'), ('Oddish', 'Poliwag', 'Diglett', 'Bellsprout'), ('Ponyta', "Farfetch'd", 'Tauros', 'Dragonite'), ('Hoothoot', 'Chikorita', 'Lanturn', 'Flaaffy', 'Teddiursa', 'Phanpy'), ('Loudred', 'Volbeat', 'Seviper', 'Wailord', 'Sealeo')]
    # write code here

====
from unittest.gui import TestCaseGui

class MyTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(check_third_element(), ['Zapdos', 'Charizard', 'Diglett', 'Tauros', 'Lanturn', 'Seviper'], "Checking that the correct entries made it into t_check.")

MyTests().main()
</textarea>
</div>
</div>
</li>
<li><div id='tup_writecode3' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<p>The dictionary ‘majors’ contains college major codes as keys and college major
names as values. Write a function named <code class="docutils literal notranslate"><span class="pre">find_major()</span></code> that takes one parameter,
<code class="docutils literal notranslate"><span class="pre">major</span> <span class="pre">code</span></code>. If the major code exists in ‘majors’, your function should
return a tuple where the first value is the major code and the second is
the name of the major. If the major code doesn’t exist, return a tuple where
the first value is None and the second is a string containing ‘Error’.
Finally, test your function by printing the name of the major with code 3084.</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=tup_writecode3q data-question_label="11.13.4">
<div id=tup_writecode3q_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="tup_writecode3q_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
majors = {3084: 'Computer Science', 3025: 'Electrical Engineering', 3020: 'Computer Engineering', 3027: 'Cybersecurity', 3068: 'Biometric Systems Engineering'}
def find_major(major_code):
    # write code here

====
from unittest.gui import TestCaseGui

class MyTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(find_major(3084),(3084, 'Computer Science'), "Checking that 'Computer Science' is associated with 3084.")
        self.assertEqual(find_major(0), (None, 'Error'), "Making sure that major code 0 returns Error.")

MyTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=tup_writecode3a data-question_label="11.13.5">
<div id=tup_writecode3a_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="tup_writecode3a_editor" 
      data-timelimit=25000     data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
majors = {3084: 'Computer Science', 3025: 'Electrical Engineering', 3020: 'Computer Engineering', 3027: 'Cybersecurity', 3068: 'Biometric Systems Engineering'}
# Define find_major function with one parameter
def find_major(major_code):
    # if the parameter is in the keys of the dictionary above
    if major_code in majors.keys():
        # return the parameter and the name from the dictionary
        return (major_code, majors[major_code])
    # Create an else statement
    else:
        # Return an error if not found
        return (None, 'Error')
# Test your function!
print(find_major(3084))
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=tup_writecode4 data-question_label="11.13.6">
<div id=tup_writecode4_question class="ac_question col-md-12">
<p>Create a function called <code class="docutils literal notranslate"><span class="pre">interchange_values</span></code> which returns a tuple which the values of <code class="docutils literal notranslate"><span class="pre">t</span></code> interchanged.</p>

</div>
<textarea data-lang="python" id="tup_writecode4_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def interchange_values():
    t = ("LeBron", "James")
    # write code here

====
from unittest.gui import TestCaseGui

class MyTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(interchange_values, ("James", "LeBron"), "Testing interchange_values() returns the correct tuple.")

MyTests().main()
</textarea>
</div>
</div>
</li>
<li><div id='tup_writecode5' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<p>If you remember, the <code class="docutils literal notranslate"><span class="pre">items()</span></code> dictionary method produces a sequence
of tuples. Keeping this in mind, we have provided you a dictionary called
<code class="docutils literal notranslate"><span class="pre">pokemon</span></code>. Write a function called <code class="docutils literal notranslate"><span class="pre">pokemon_transformation</span></code> that appends
the key to the list <code class="docutils literal notranslate"><span class="pre">p_names</span></code> and appends the value to the list <code class="docutils literal notranslate"><span class="pre">p_number</span></code>,
for each key-value pair. Then, zip these two lists together, and return that list.
Do not use the <code class="docutils literal notranslate"><span class="pre">keys()</span></code> or <code class="docutils literal notranslate"><span class="pre">values()</span></code> methods.</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=tup_writecode5q data-question_label="11.13.7">
<div id=tup_writecode5q_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="tup_writecode5q_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def pokemon_transformation():
    pokemon = {'Rattata': 19, 'Machop': 66, 'Seel': 86, 'Volbeat': 86, 'Solrock': 126}
    # write code here

====
from unittest.gui import TestCaseGui

class MyTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(p_names, ['Rattata', 'Machop', 'Seel', 'Volbeat', 'Solrock'], "Testing that p_name has the correct values.")
        self.assertEqual(p_number, [19, 66, 86, 86, 126], "Testing that p_number has the correct values.")

MyTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=tup_writecode5a data-question_label="11.13.8">
<div id=tup_writecode5a_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="tup_writecode5a_editor" 
      data-timelimit=25000     data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def pokemon_transformation():
    pokemon = {'Rattata': 19, 'Machop': 66, 'Seel': 86, 'Volbeat': 86, 'Solrock': 126}
    # Create lists to hold names and numbers for each Pokemon
    p_names = []
    p_number = []
    # Use a for loop to iterate through the key, value pairs from pokemon dictionary
    for key, val in pokemon.items():
        # Add the keys and values to their respective lists
        p_names.append(key)
        p_number.append(val)
    # Zip the two lists
    zipped_list = list(zip(p_names, p_number))
    return zipped_list
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=tup_writecode6 data-question_label="11.13.9">
<div id=tup_writecode6_question class="ac_question col-md-12">
<p>Create a function called <code class="docutils literal notranslate"><span class="pre">my_data</span></code> that returns a tuple that contains one element, the integer 99.</p>

</div>
<textarea data-lang="python" id="tup_writecode6_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def my_data():
    # write code here
====
from unittest.gui import TestCaseGui

class MyTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(my_data(), (99,), "Testing that my_data() returns a tuple that has one element, 99.")

MyTests().main()
</textarea>
</div>
</div>
</li>
<li><div id='tup_writecode7' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<p>Write the function <code class="docutils literal notranslate"><span class="pre">mod_tuples</span></code> which takes a list of tuples, <code class="docutils literal notranslate"><span class="pre">tup_list</span></code> and returns a copy where the
last element in each tuple is modified to be 100. For example, <code class="docutils literal notranslate"><span class="pre">mod_tuples([(3,4),</span> <span class="pre">(20,</span> <span class="pre">-3,</span> <span class="pre">2)])</span></code> returns
[(3,100), (20, -3, 100)]``.</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=tup_writecode7q_v2 data-question_label="11.13.10">
<div id=tup_writecode7q_v2_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="tup_writecode7q_v2_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def mod_tuples(tup_list):

====
from unittest.gui import TestCaseGui

class MyTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(mod_tuples([(3, 4), (20, -3, 2)]),[(3, 100), (20, -3, 100)] "mod_tuples([(3, 4), (20, -3, 2)]")
        self.assertEqual(mod_tuples([(3, ), (20, -3, 2)]),[(100, ), (20, -3, 100)] "mod_tuples([(3, ), (20, -3, 2)]")
        self.assertEqual(mod_tuples([(3, 4), (-3, )]),[(3, 100), (100,)] "mod_tuples([(3, 4), (-3, )]),[(3, 100), (100,)]")
        self.assertEqual(mod_tuples([(20, )]),[(100, )] "mod_tuples([(20, )])")
        self.assertEqual(mod_tuples([(20, 1, 80, 2 )]),[(20, 1, 80, 100)] "mod_tuples([(20, 1, 80, 2 )])")

MyTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=tup_writecode7a data-question_label="11.13.11">
<div id=tup_writecode7a_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="tup_writecode7a_editor" 
      data-timelimit=25000     data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def mod_tuples(tup_list):
    # Access the last element of each list (-1) and replace with 100 in each element of the tuple
    updated_list = [tup[:-1] + (100,) for tup in tup_list]
# return the updated list
return updated_list
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=tup_writecode8 data-question_label="11.13.12">
<div id=tup_writecode8_question class="ac_question col-md-12">
<p>Define a function called <code class="docutils literal notranslate"><span class="pre">info</span></code> with the following required parameters: name, age, birth_year, year_in_college, and hometown. The
function should return a tuple that contains all the inputted information. For example, <code class="docutils literal notranslate"><span class="pre">info('Troy',</span> <span class="pre">24,</span> <span class="pre">1996,</span> <span class="pre">'Sophomore',</span> <span class="pre">'Ann</span> <span class="pre">Arbor')</span></code>
would return <code class="docutils literal notranslate"><span class="pre">('Troy',</span> <span class="pre">24,</span> <span class="pre">1996,</span> <span class="pre">'Sophomore',</span> <span class="pre">'Ann</span> <span class="pre">Arbor')</span></code>.</p>

</div>
<textarea data-lang="python" id="tup_writecode8_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class MyTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(info('Troy', 24, 1996, 'Sophomore', 'Ann Arbor'), ('Troy', 24, 1996, 'Sophomore', 'Ann Arbor'), "Testing that the function correctly outputs the tuple.")
        self.assertEqual(info('Jane', 18, 2002, 'Freshman', 'Manhattan'), ('Jane', 18, 2002, 'Freshman', 'Manhattan'), "Testing that the function correctly outputs the tuple.")

MyTests().main()
</textarea>
</div>
</div>
</li>
<li><div id='tup_writecode9' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<p>Write a function <code class="docutils literal notranslate"><span class="pre">list_link</span></code> that accepts two lists and returns a dictionary
with the first list as the key and the second list as the value. For example,
<code class="docutils literal notranslate"><span class="pre">list_link(['what',</span> <span class="pre">'do',</span> <span class="pre">'you',</span> <span class="pre">'do'],</span> <span class="pre">[1,2,3,4])</span></code> should return <code class="docutils literal notranslate"><span class="pre">{'what':</span> <span class="pre">1,</span> <span class="pre">'do':</span> <span class="pre">4,</span> <span class="pre">'you':</span> <span class="pre">3}</span></code>.</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=tup_writecode9q data-question_label="11.13.13">
<div id=tup_writecode9q_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="tup_writecode9q_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def list_link(lst1, lst2):
    #write your code here

print(list_link(['what', 'do', 'you', 'do'], [1,2,3,4]))

====
from unittest.gui import TestCaseGui

class MyTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(list_link(['what', 'do', 'you', 'do'], [1,2,3,4]), {'what': 1, 'do': 4, 'you': 3}, "Testing that the function was properly defined.")
        self.assertEqual(list_link(['I', 'love', 'coding'], [2,0,-8]), {'I': 2, 'love': 0, 'coding': -8}, "Testing that the function was properly defined.")

MyTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=tup_writecode9a data-question_label="11.13.14">
<div id=tup_writecode9a_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="tup_writecode9a_editor" 
      data-timelimit=25000     data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
# Define function with 2 lists as arguments
def list_link(lst1, lst2):
    # Create dictionary
    diction = {}
    # Create counter variable to count iterations
    counter = 0
    # Create condition for when lists are the same length
    if len(list1) == len(list2):
        # iterate through item in list1
        for i in list1:
            # Add that item to dictionary with its place in list2
            diction[i] = list2[counter]
            # Increment counter
            counter += 1
    # Return the dictionary
    return diction

print(list_link(['what', 'do', 'you', 'do'], [1,2,3,4]))
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=tup_writecode10 data-question_label="11.13.15">
<div id=tup_writecode10_question class="ac_question col-md-12">
<p>Create a function <code class="docutils literal notranslate"><span class="pre">tuplize()</span></code> that accepts two inputs and returns a tuple containing those inputs in order.
For example, <code class="docutils literal notranslate"><span class="pre">tuplize('Stephen',</span> <span class="pre">'Curry')</span></code> would return <code class="docutils literal notranslate"><span class="pre">('Stephen',</span> <span class="pre">'Curry')</span></code>.</p>

</div>
<textarea data-lang="python" id="tup_writecode10_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class MyTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(tuplize('Stephen', 'Curry'), ('Stephen', 'Curry'), "Testing that the function creates the correct tuple.")
        self.assertEqual(tuplize(1, 2), (1, 2), "Testing that the function creates the correct tuple.")

MyTests().main()
</textarea>
</div>
</div>
</li>
</ol>
</div>


  <div id=scprogresscontainer>
    You have attempted <span id=scprogresstotal></span> of <span id=scprogressposs></span> activities on this page <div id="subchapterprogress"></div>
  </div>
  
      
<ul>
  <li id="relations-prev" style="opacity:0.4" onmouseover="mouseOverPrev()" onmouseout="mouseOutPrev()"class="navLink" title='Previous Section - 11.12. Mixed-Up Code Questions' data-toggle="tooltip">
    <a href="mixedupcode.html" aria-label="11.12. Mixed-Up Code Questions">
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">
            <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
        </div>
    </a>
  </li>
</ul>

  
<ul>
  <li id="relations-next" style="opacity:0.4" onmouseover="mouseOverNext()" onmouseout="mouseOutNext()" class="navLink" title='Next Section - 12. Regular Expressions' data-toggle="tooltip" >
    <a href="../regex/toctree.html" aria-label="12. Regular Expressions">
      <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">
          <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
      </div>
    </a>
  </li>
</ul>

<script>
document.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>




<script> 
  document.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>