<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.17.1: http://docutils.sourceforge.net/" />

    <title>5.32. Functions and Lists Write Code Questions &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.cd688d1623c9932e.css?v=630908D8" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.bf73008a1eba9f4e.css?v=7DED055E" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/runtime.439d581f758f2b99.bundle.js?v=2E54442"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.51d9dde4e5162190.bundle.js?v=9CAA2D33"></script>
    <script src="../_static/runestone.42b576791a4ed54b.bundle.js?v=B7293199"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="5.33. Group Work: Functions with Loops" href="funcWithLoops.html" />
    <link rel="prev" title="5.31. Functions and Lists Mixed-Up Code Questions" href="function_list_mixedupcode.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    
    eBookConfig.useRunestoneServices = {{use_services}};
    eBookConfig.host = '';
    eBookConfig.app = eBookConfig.host + '/runestone';
    eBookConfig.course = '{{ course_name }}';
    eBookConfig.basecourse = '{{ base_course }}';
    eBookConfig.isLoggedIn = {{ is_logged_in}};
    eBookConfig.email = '{{ user_email }}';
    eBookConfig.isInstructor = {{ is_instructor }};
    eBookConfig.username = '{{ user_id}}';
    eBookConfig.readings = {{ readings|safe}};
    eBookConfig.activities = {{ activity_info|safe }}
    eBookConfig.downloadsEnabled = {{downloads_enabled}};
    eBookConfig.allow_pairs = {{allow_pairs}}
    eBookConfig.enableCompareMe = {{enable_compare_me}};
    // _`new_server_prefix`: Defined by the server to be the path to the new server.
    eBookConfig.new_server_prefix = '{{ new_server_prefix }}';
    
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->

    
        {% if serve_ad is defined and serve_ad and settings.adsenseid %}
         <script data-ad-client="{{settings.adsenseid}}" async src="https://pagead2.googlesyndication.com/pagead/js/adsbygoogle.js"></script>
        {% endif %}
    



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
            
            {{ course_name }}
            
          
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="browsing_warning"></span></li>

        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li id="inst_peer_link"><a href='/runestone/peer/instructor.html'>Peer Instruction (Instructor)</a></li>            
            <li><a href='/runestone/peer/student.html'>Peer Instruction (Student)</a></li>            
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/ns/books/published/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=function_list_writecode'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
            <script></script>
<li class="dropdown globaltoc-container">
    <a
        href="../index.html"
        class="dropdown-toggle"
        data-toggle="dropdown"
        aria-label="index"
        >This Chapter<b class="caret"></b
    ></a>
    <ul class="dropdown-menu globaltoc">
         
        {% for row in subchapter_list %}
        <li class="toctree-l1">
            <a class="reference internal" href="{{ row['subchap_uri']}}"
                >{{ row['title']}}</a
            >
        </li>
        {% endfor %} 
        
    </ul>
</li>
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">5.31. </span>Functions and Lists Mixed-Up Code Questions" data-toggle="tooltip">
    <a href="function_list_mixedupcode.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">5.33. </span>Group Work: Functions with Loops' data-toggle="tooltip" >              
    <a href="funcWithLoops.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->

    
        {% if settings.num_banners > 0 and settings.show_rs_banner %}
         <div style="width: 100%;">
            <a href="/runestone/default/donate?banner={{banner_num}}">
            <img class="runestone-banner"
                 src="/runestone/static/images/RunestoneBanner{{banner_num}}.png"
                 alt="Please Support Runestone"
                 style="display: block; width: 728px; margin-left: auto; margin-right: auto; margin-top: 10px;">
            </a>
         </div>
        {% endif %}
    


  
  <section id="functions-and-lists-write-code-questions">
<h1><span class="section-number">5.32. </span>Functions and Lists Write Code Questions<a class="headerlink" href="#functions-and-lists-write-code-questions" title="Permalink to this headline">¶</a></h1>
<ol class="arabic">
<li><div id='funct_list_writecode1' data-component="tabbedStuff"  class='alert alert-warning '><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=funct_list_writecode1q data-question_label="5.32.1">
<div id=funct_list_writecode1q_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">average_of_num_list</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">num_list</span></code> and returns the average of all the numbers in num_list.
For example, <code class="docutils literal notranslate"><span class="pre">average_of_num_list([0,</span> <span class="pre">20.8,</span> <span class="pre">80,</span> <span class="pre">5])</span></code> would return <code class="docutils literal notranslate"><span class="pre">26.45</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_list_writecode1q_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def average_of_num_list(num_list):
    # write code here

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(average_of_num_list([0]),0,"average_of_num_list([0])")
        self.assertAlmostEqual(average_of_num_list([0, 20.8, 80, 5]),26.45,2,"average_of_num_list([0, 20.8, 80, 5])")
        self.assertAlmostEqual(average_of_num_list([0, 15, 13, 14, 7]),9.8,1,"average_of_num_list([0, 15, 13, 14, 7])")
        self.assertAlmostEqual(average_of_num_list([0, 100, 5, 10]),28.75,2,"average_of_num_list([0, 100, 5, 10])")
        self.assertAlmostEqual(average_of_num_list([0, 25, 40, -20, -100]),-11.0,1,"average_of_num_list([0, 25, 40, -20, -100])")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=funct_list_writecode1a data-question_label="5.32.2">
<div id=funct_list_writecode1a_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">average_of_num_list</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">num_list</span></code> and returns the average of all the numbers in num_list.
For example, <code class="docutils literal notranslate"><span class="pre">average_of_num_list([0,</span> <span class="pre">20.8,</span> <span class="pre">80,</span> <span class="pre">5])</span></code> would return <code class="docutils literal notranslate"><span class="pre">26.45</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_list_writecode1a_editor" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def average_of_num_list(num_list):
    return sum(num_list) / len(num_list)
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=funct_list_writecode2q data-question_label="5.32.3">
<div id=funct_list_writecode2q_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">names</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">name_list</span></code> and returns an alphabetically sorted <code class="docutils literal notranslate"><span class="pre">name_list</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">names(['Susan',</span> <span class="pre">'Sara',</span> <span class="pre">'Sammy',</span> <span class="pre">'Sarah'])</span></code> would return <code class="docutils literal notranslate"><span class="pre">['Sammy',</span> <span class="pre">'Sara',</span> <span class="pre">'Sarah',</span> <span class="pre">'Susan']</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_list_writecode2q_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def names(name_list):
    # write code here

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(names(['Sally', 'Jimmy', 'Peter', 'Tinkerbell']),['Jimmy', 'Peter', 'Sally', 'Tinkerbell'],"names(['Sally', 'Jimmy', 'Peter', 'Tinkerbell'])")
        self.assertEqual(names(['Timmy']),['Timmy'],"names(['Timmy'])")
        self.assertEqual(names(['Susan', 'Sara', 'Sammy', 'Sarah']),['Sammy', 'Sara', 'Sarah', 'Susan'],"names(['Susan', 'Sara', 'Sammy', 'Sarah'])")

myTests().main()
</textarea>
</div>
</div>
</li>
<li><div id='funct_list_writecode3' data-component="tabbedStuff"  class='alert alert-warning '><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=funct_list_writecode3q data-question_label="5.32.4">
<div id=funct_list_writecode3q_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">remove_min_value</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">num_list</span></code> and returns a <code class="docutils literal notranslate"><span class="pre">num_list</span></code> without the
minimum value from <code class="docutils literal notranslate"><span class="pre">num_list</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">remove_min_value([20,</span> <span class="pre">4,</span> <span class="pre">1203,</span> <span class="pre">7482,</span> <span class="pre">3])</span></code> would return <code class="docutils literal notranslate"><span class="pre">[20,</span> <span class="pre">4,</span> <span class="pre">1203,</span> <span class="pre">7482]</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_list_writecode3q_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def remove_min_value(num_list):
    # write code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(remove_min_value([20, 4, 1203, 7482, 3]),[20, 4, 1203, 7482],"remove_min_value([20, 4, 1203, 7482, 3])")
        self.assertEqual(remove_min_value([3]),[],"remove_min_value([3])")
        self.assertEqual(remove_min_value([3, 0, 100]),[3, 100],"remove_min_value([3, 0, 100])")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=funct_list_writecode3a data-question_label="5.32.5">
<div id=funct_list_writecode3a_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">remove_min_value</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">num_list</span></code> and returns a <code class="docutils literal notranslate"><span class="pre">num_list</span></code> without the
minimum value from <code class="docutils literal notranslate"><span class="pre">num_list</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">remove_min_value([20,</span> <span class="pre">4,</span> <span class="pre">1203,</span> <span class="pre">7482,</span> <span class="pre">3])</span></code> would return <code class="docutils literal notranslate"><span class="pre">[20,</span> <span class="pre">4,</span> <span class="pre">1203,</span> <span class="pre">7482]</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_list_writecode3a_editor" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def remove_min_value(num_list):
    num_list.remove(min(num_list))
    return num_list
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=funct_list_writecode4q data-question_label="5.32.6">
<div id=funct_list_writecode4q_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">range_given_list</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">list_of_nums</span></code> and
returns the range (max value - min value) of the values. Try using the sort method and indexing.
For example <code class="docutils literal notranslate"><span class="pre">range_given_list([20,</span> <span class="pre">100,</span> <span class="pre">2000,</span> <span class="pre">15,</span> <span class="pre">3,</span> <span class="pre">12])</span></code> would return <code class="docutils literal notranslate"><span class="pre">1997</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_list_writecode4q_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def range_given_list(list_of_nums):
    # write code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(range_given_list([20, 100, 2000, 15, 3, 12]),1997,"range_given_list([20, 100, 2000, 15, 3, 12])")
        self.assertEqual(range_given_list([20, 100, 2000, 15, 3, 0]),2000,"range_given_list([20, 100, 2000, 15, 3, 0])")
        self.assertEqual(range_given_list([20, 100, 2000, 15, 3, -12]),2012,"range_given_list([20, 100, 2000, 15, 3, -12])")
        self.assertEqual(range_given_list([-20, -100, -2000, -15, -3, -12]),1997,"range_given_list([-20, -100, -2000, -15, -3, -12])")
        self.assertAlmostEqual(range_given_list([20.7, 100.3, 2000.5, 2000.4, 15.7, 3.6, -12.9]),2013.4,1,"range_given_list([20.7, 100.3, 2000.5, 2000.4, 15.7, 3.6, -12.9])")

myTests().main()
</textarea>
</div>
</div>
</li>
<li><div id='funct_list_writecode5' data-component="tabbedStuff"  class='alert alert-warning '><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=funct_list_writecode5q data-question_label="5.32.7">
<div id=funct_list_writecode5q_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">remove_indices_after_first_max_value</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">num_list</span></code> and
returns a <code class="docutils literal notranslate"><span class="pre">new_num_list</span></code> with values up to the max value of the list. For example, <code class="docutils literal notranslate"><span class="pre">remove_indices_after_first_max_value([200,</span> <span class="pre">10,</span> <span class="pre">5,</span> <span class="pre">200])</span></code>
would return <code class="docutils literal notranslate"><span class="pre">[5,</span> <span class="pre">10,</span> <span class="pre">5,</span> <span class="pre">200]</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_list_writecode5q_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def remove_indices_after_first_max_value(num_list):
    # write code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(remove_indices_after_first_max_value([200, 10, 5, 200]),[200],"remove_indices_after_first_max_value([200, 10, 5, 200])")
        self.assertEqual(remove_indices_after_first_max_value([5, 13, 0, -201, 200, 10, 5, 200]),[5, 13, 0, -201, 200],"remove_indices_after_first_max_value([5, 13, 0, -201, 200, 10, 5, 200])")
        self.assertEqual(remove_indices_after_first_max_value([5.6, 7.9, 11.3, 10, 200.1, 10, 5, 200]),[5.6, 7.9, 11.3, 10, 200.1],"remove_indices_after_first_max_value([5.6, 7.9, 11.3, 10, 200.1, 10, 5, 200])")
        self.assertEqual(remove_indices_after_first_max_value([100, 10, 5, 200]),[100, 10, 5, 200],"remove_indices_after_first_max_value([100, 10, 5, 200])")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=funct_list_writecode5a data-question_label="5.32.8">
<div id=funct_list_writecode5a_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">remove_indices_after_first_max_value</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">num_list</span></code> and
returns a <code class="docutils literal notranslate"><span class="pre">new_num_list</span></code> with values up to the max value of the list. For example, <code class="docutils literal notranslate"><span class="pre">remove_indices_after_first_max_value([200,</span> <span class="pre">10,</span> <span class="pre">5,</span> <span class="pre">200])</span></code>
would return <code class="docutils literal notranslate"><span class="pre">[5,</span> <span class="pre">10,</span> <span class="pre">5,</span> <span class="pre">200]</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_list_writecode5a_editor" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def remove_indices_after_first_max_value(num_list):
    index_value = num_list.index(max(num_list)) + 1
    new_num_list = num_list[:index_value]
    return new_num_list
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=funct_list_writecode6q data-question_label="5.32.9">
<div id=funct_list_writecode6q_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">transform_and_combine</span></code> that takes in two parameters, <code class="docutils literal notranslate"><span class="pre">list_one</span></code>, which must have at least one element, and <code class="docutils literal notranslate"><span class="pre">list_two</span></code>.
Remove the last element from <code class="docutils literal notranslate"><span class="pre">list_one</span></code>, then reverse the list. Sort <code class="docutils literal notranslate"><span class="pre">list_two</span></code>, then extend <code class="docutils literal notranslate"><span class="pre">list_two</span></code> with <code class="docutils literal notranslate"><span class="pre">list_one</span></code>, and return <code class="docutils literal notranslate"><span class="pre">list_two</span></code>.
Hint: Use list methods (e.g., pop, sort, append, reverse, and extend). For example, <code class="docutils literal notranslate"><span class="pre">transform_and_combine([5,</span> <span class="pre">20,</span> <span class="pre">3,</span> <span class="pre">15,</span> <span class="pre">200,</span> <span class="pre">0,</span> <span class="pre">17],</span> <span class="pre">['Hello',</span> <span class="pre">'Bye',</span> <span class="pre">'How</span> <span class="pre">are</span> <span class="pre">you?'])</span></code>
would return <code class="docutils literal notranslate"><span class="pre">['Bye',</span> <span class="pre">'Hello',</span> <span class="pre">'How</span> <span class="pre">are</span> <span class="pre">you?',</span> <span class="pre">0,</span> <span class="pre">200,</span> <span class="pre">15,</span> <span class="pre">3,</span> <span class="pre">20,</span> <span class="pre">5]</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_list_writecode6q_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def transform_and_combine(list_one, list_two):
    # write code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(transform_and_combine([5, 20, 3, 15, 200, 0, 17], ['Hello', 'Bye', 'How are you?']),['Bye', 'Hello', 'How are you?', 0, 200, 15, 3, 20, 5],"transform_and_combine([5, 20, 3, 15, 200, 0, 17], ['Hello', 'Bye', 'How are you?'])")
        self.assertEqual(transform_and_combine(['Susan', 'Sara', 'Sammy', 'Sarah'], ['Hello', 'Bye', 'How are you?']),['Bye', 'Hello', 'How are you?', 'Sammy', 'Sara', 'Susan'],"transform_and_combine(['Susan', 'Sara', 'Sammy', 'Sarah'], ['Hello', 'Bye', 'How are you?'])")
        self.assertEqual(transform_and_combine([5, 13, 0, -201, 200, 10, 5, 200], [200, 10, 5, 200]),[5, 10, 200, 200, 5, 10, 200, -201, 0, 13, 5],"transform_and_combine([5, 13, 0, -201, 200, 10, 5, 200], [200, 10, 5, 200])")
        self.assertEqual(transform_and_combine([1], [4]),[4],"transform_and_combine([1], [4]")
        self.assertEqual(transform_and_combine([1], []),[],"transform_and_combine([1], []")



myTests().main()
</textarea>
</div>
</div>
</li>
</ol>
</section>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 5.31. Functions and Lists Mixed-Up Code Questions' data-toggle="tooltip">
  <a href="function_list_mixedupcode.html" aria-label="Previous - 5.31. Functions and Lists Mixed-Up Code Questions">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 5.33. Group Work: Functions with Loops' data-toggle="tooltip" >
    <a href="funcWithLoops.html" aria-label="Next - 5.33. Group Work: Functions with Loops">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>



  
    {% if settings.google_ga %}
      <script>
        var _gaq = _gaq || [];
        _gaq.push(['_setAccount', '{{ settings.google_ga }}']);
        _gaq.push(['_trackPageview']);

        (function() {
          var ga = document.createElement('script'); ga.async = true;
          ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
          var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
        })();
      </script>
    {% endif %}
  

  
    <script type="application/json" class="js-hypothesis-config">
    {
      "showHighlights": false
    }
    </script>

    <script async src="https://hypothes.is/embed.js"></script>
    <!--provide space to prevent hypothes.is sidebar from obscuring content-->
    <style>@media screen and (max-width:767px) { body { padding-right: 21px; }}</style>

  


<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>