<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.17.1: http://docutils.sourceforge.net/" />

    <title>Functions and Lists Mixed-Up Code Questions &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.cd688d1623c9932e.css?v=630908D8" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.bf73008a1eba9f4e.css?v=7DED055E" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/runtime.439d581f758f2b99.bundle.js?v=2E54442"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.51d9dde4e5162190.bundle.js?v=9CAA2D33"></script>
    <script src="../_static/runestone.42b576791a4ed54b.bundle.js?v=B7293199"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    
    eBookConfig.useRunestoneServices = {{use_services}};
    eBookConfig.host = '';
    eBookConfig.app = eBookConfig.host + '/runestone';
    eBookConfig.course = '{{ course_name }}';
    eBookConfig.basecourse = '{{ base_course }}';
    eBookConfig.isLoggedIn = {{ is_logged_in}};
    eBookConfig.email = '{{ user_email }}';
    eBookConfig.isInstructor = {{ is_instructor }};
    eBookConfig.username = '{{ user_id}}';
    eBookConfig.readings = {{ readings|safe}};
    eBookConfig.activities = {{ activity_info|safe }}
    eBookConfig.downloadsEnabled = {{downloads_enabled}};
    eBookConfig.allow_pairs = {{allow_pairs}}
    eBookConfig.enableCompareMe = {{enable_compare_me}};
    // _`new_server_prefix`: Defined by the server to be the path to the new server.
    eBookConfig.new_server_prefix = '{{ new_server_prefix }}';
    
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->

    
        {% if serve_ad is defined and serve_ad and settings.adsenseid %}
         <script data-ad-client="{{settings.adsenseid}}" async src="https://pagead2.googlesyndication.com/pagead/js/adsbygoogle.js"></script>
        {% endif %}
    



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
            
            {{ course_name }}
            
          
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="browsing_warning"></span></li>

        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li id="inst_peer_link"><a href='/runestone/peer/instructor.html'>Peer Instruction (Instructor)</a></li>            
            <li><a href='/runestone/peer/student.html'>Peer Instruction (Student)</a></li>            
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/ns/books/published/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=function_list_mixedupcode_orig'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
            <script></script>
<li class="dropdown globaltoc-container">
    <a
        href="../index.html"
        class="dropdown-toggle"
        data-toggle="dropdown"
        aria-label="index"
        >This Chapter<b class="caret"></b
    ></a>
    <ul class="dropdown-menu globaltoc">
         
        {% for row in subchapter_list %}
        <li class="toctree-l1">
            <a class="reference internal" href="{{ row['subchap_uri']}}"
                >{{ row['title']}}</a
            >
        </li>
        {% endfor %} 
        
    </ul>
</li>
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->

    
        {% if settings.num_banners > 0 and settings.show_rs_banner %}
         <div style="width: 100%;">
            <a href="/runestone/default/donate?banner={{banner_num}}">
            <img class="runestone-banner"
                 src="/runestone/static/images/RunestoneBanner{{banner_num}}.png"
                 alt="Please Support Runestone"
                 style="display: block; width: 728px; margin-left: auto; margin-right: auto; margin-top: 10px;">
            </a>
         </div>
        {% endif %}
    


  
  <section id="functions-and-lists-mixed-up-code-questions">
<h1>Functions and Lists Mixed-Up Code Questions<a class="headerlink" href="#functions-and-lists-mixed-up-code-questions" title="Permalink to this headline">¶</a></h1>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="functions_lists_mixedupcode_1" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">second_half</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">lst</span></code> as a parameter and returns the second half of <code class="docutils literal notranslate"><span class="pre">lst</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">second_half([1])</span></code> should return <code class="docutils literal notranslate"><span class="pre">[1]</span></code>, <code class="docutils literal notranslate"><span class="pre">second_half([1,</span> <span class="pre">2,</span> <span class="pre">3])</span></code> should return <code class="docutils literal notranslate"><span class="pre">[2,</span> <span class="pre">3]</span></code>, and <code class="docutils literal notranslate"><span class="pre">second_half([1,</span> <span class="pre">2,</span> <span class="pre">3,</span> <span class="pre">4])</span></code>
should return <code class="docutils literal notranslate"><span class="pre">[3,</span> <span class="pre">4]</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="1"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def second_half(lst):
---
    return lst[(len(lst) // 2):]
---
    return lst[(len(lst) / 2):] #paired
---
    return lst[(len(lst) / 2):-1] #paired
---
    return lst[(len(lst) // 2):(len(lst) - 1)] #paired
---
    return lst[(len(lst) % 2):len(lst)] #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=functions_lists_mixedupcode_1_ac data-question_label="2">
<div id=functions_lists_mixedupcode_1_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">second_half</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">lst</span></code> as a parameter and returns the second half of <code class="docutils literal notranslate"><span class="pre">lst</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">second_half([1])</span></code> should return <code class="docutils literal notranslate"><span class="pre">[1]</span></code>, <code class="docutils literal notranslate"><span class="pre">second_half([1,</span> <span class="pre">2,</span> <span class="pre">3])</span></code> should return <code class="docutils literal notranslate"><span class="pre">[2,</span> <span class="pre">3]</span></code>, and <code class="docutils literal notranslate"><span class="pre">second_half([1,</span> <span class="pre">2,</span> <span class="pre">3,</span> <span class="pre">4])</span></code>
should return <code class="docutils literal notranslate"><span class="pre">[3,</span> <span class="pre">4]</span></code>.</p>

</div>
<textarea data-lang="python" id="functions_lists_mixedupcode_1_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(second_half([1, 2, 3]),[2, 3],"second_half([1, 2, 3])")
        self.assertEqual(second_half([1]),[1],"second_half([1])")
        self.assertEqual(second_half([1, 2, 3, 4]),[3, 4],"second_half([1, 2, 3, 4])")
        self.assertEqual(second_half(['hello', 'many', False, True, 2, 3, 'happy']),[True, 2, 3, 'happy'],"second_half(['hello', 'many', False, True, 2, 3, 'happy'])")


myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="functions_lists_mixedupcode_2" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">sum_of_split_num</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">str_num_with_hyphens</span></code> as a parameter and returns the sum of the integers that
are split up by hyphens. For example, <code class="docutils literal notranslate"><span class="pre">sum_of_split_num('978-876-4908')</span></code> should return <code class="docutils literal notranslate"><span class="pre">6762</span></code> because 978 + 876 + 4908 = 6762.
(Note: Integers are 0 or greater.)</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="3"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def sum_of_split_num(str_num_with_hyphens):
---
    num_without_hyphens = str_num_with_hyphens.split('-')
---
    num_without_hyphens = str_num_with_hyphens.split() #paired
---
    map_object_to_ints = map(int, num_without_hyphens)
---
    lst_of_ints = list(map_object_to_ints)
---
    return sum(lst_of_ints)
---
    return sum(lst_of_ints) / 2 #paired
---
    return lst_of_ints.sum() #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=functions_lists_mixedupcode_2_ac data-question_label="4">
<div id=functions_lists_mixedupcode_2_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">sum_of_split_num</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">str_num_with_hyphens</span></code> as a parameter and returns the sum of the integers that
are split up by hyphens. For example, <code class="docutils literal notranslate"><span class="pre">sum_of_split_num('978-876-4908')</span></code> should return <code class="docutils literal notranslate"><span class="pre">6762</span></code> because 978 + 876 + 4908 = 6762. (Note: Integers
are 0 or greater.)</p>

</div>
<textarea data-lang="python" id="functions_lists_mixedupcode_2_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sum_of_split_num('978-876-4908'),6762,"sum_of_split_num('978-876-4908')")
        self.assertEqual(sum_of_split_num('000-05746-28736'),34482,"sum_of_split_num('000-05746-28736')")


myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="functions_lists_mixedupcode_3" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">sort_descending</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">lst_of_nums</span></code> as a parameter and returns <code class="docutils literal notranslate"><span class="pre">lst_of_nums</span></code> in descending order.
For example, <code class="docutils literal notranslate"><span class="pre">sort_descending([100,</span> <span class="pre">2000,</span> <span class="pre">-50])</span></code> should return <code class="docutils literal notranslate"><span class="pre">[2000,</span> <span class="pre">100,</span> <span class="pre">-50]</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="5"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def sort_descending(lst_of_nums):
---
    lst_of_nums.sort(reverse=True)
---
    lst_of_nums.sort(reverse=False) #paired
---
    lst_of_nums.sort() #paired
---
    lst_of_nums = lst_of_nums.sort(reverse=True) #paired
---
    return lst_of_nums
---
    return sorted(lst_of_nums, reverse=False) #paired
---
    return sorted(lst_of_nums) #paired
---
    return lst_of_nums.sort(reverse=True) #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=functions_lists_mixedupcode_3_ac data-question_label="6">
<div id=functions_lists_mixedupcode_3_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">sort_descending</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">lst_of_nums</span></code> as a parameter and returns <code class="docutils literal notranslate"><span class="pre">lst_of_nums</span></code> in descending order.
For example, <code class="docutils literal notranslate"><span class="pre">sort_descending([100,</span> <span class="pre">2000,</span> <span class="pre">-50])</span></code> should return <code class="docutils literal notranslate"><span class="pre">[2000,</span> <span class="pre">100,</span> <span class="pre">-50]</span></code>.</p>

</div>
<textarea data-lang="python" id="functions_lists_mixedupcode_3_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sort_descending([100, 2000, -50]),[2000, 100, -50],"sort_descending([100, 2000, -50])")
        self.assertEqual(sort_descending([83, 0, -2000]),[83, 0, -2000],"sort_descending([83, 0, -2000])")
        self.assertEqual(sort_descending([-83]),[-83],"sort_descending([-83])")


myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="functions_lists_mixedupcode_4" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">extend_two_lists_sort</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">lst1</span></code> and <code class="docutils literal notranslate"><span class="pre">lst2</span></code> as parameters, sorts both lists, and extends <code class="docutils literal notranslate"><span class="pre">lst1</span></code> with <code class="docutils literal notranslate"><span class="pre">lst2</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">extend_two_lists_sort(['string',</span> <span class="pre">'hello'],</span> <span class="pre">['my',</span> <span class="pre">'goodness'])</span></code> should return <code class="docutils literal notranslate"><span class="pre">['hello',</span> <span class="pre">'string',</span> <span class="pre">'goodness',</span> <span class="pre">'my']</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="7"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def extend_two_lists_sort(lst1, lst2):
---
    lst1.sort()
---
    sorted(lst1) #paired
---
    lst2.sort()
---
    sorted(lst2) #paired
---
    lst1.extend(lst2)
---
    lst1.append(lst2) #paired
---
    return lst1
---
    return lst1.extend(lst2) #paired
---
    return lst1.extend(lst2).sort() #paired
---
    return sorted(lst1.extend(lst2)) #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=functions_lists_mixedupcode_4_ac data-question_label="8">
<div id=functions_lists_mixedupcode_4_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">extend_two_lists_sort</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">lst1</span></code> and <code class="docutils literal notranslate"><span class="pre">lst2</span></code> as parameters, sorts both lists, and extends <code class="docutils literal notranslate"><span class="pre">lst1</span></code> with <code class="docutils literal notranslate"><span class="pre">lst2</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">extend_two_lists_sort(['string',</span> <span class="pre">'hello'],</span> <span class="pre">['my',</span> <span class="pre">'goodness'])</span></code> should return <code class="docutils literal notranslate"><span class="pre">['hello',</span> <span class="pre">'string',</span> <span class="pre">'goodness',</span> <span class="pre">'my']</span></code>.</p>

</div>
<textarea data-lang="python" id="functions_lists_mixedupcode_4_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(extend_two_lists_sort(['string', 'hello'], ['my', 'goodness']),['hello', 'string', 'goodness', 'my'],"extend_two_lists_sort(['string', 'hello'], ['my', 'goodness'])")
        self.assertEqual(extend_two_lists_sort([1, -5, 0], ['my', 'hello']),[-5, 0, 1, 'hello', 'my'],"extend_two_lists_sort([1, -5, 0], ['my', 'hello'])")
        self.assertEqual(extend_two_lists_sort([1, -5, 0], [8000, -23, -1, 0]),[-5, 0, 1, -23, -1, 0, 8000],"extend_two_lists_sort([1, -5, 0], [8000, -23, -1, 0])")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="functions_lists_mixedupcode_5" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">sort_and_get_median_num</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">lst_of_nums</span></code> as a parameter, sorts <code class="docutils literal notranslate"><span class="pre">lst_of_nums</span></code>, and returns the median
of <code class="docutils literal notranslate"><span class="pre">lst_of_nums</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">sort_and_get_median_num([200,</span> <span class="pre">-5,</span> <span class="pre">0,</span> <span class="pre">75,</span> <span class="pre">80,</span> <span class="pre">60])</span></code> should return <code class="docutils literal notranslate"><span class="pre">67.5</span></code>, and
<code class="docutils literal notranslate"><span class="pre">sort_and_get_median_num([200,</span> <span class="pre">-5,</span> <span class="pre">0,</span> <span class="pre">75,</span> <span class="pre">80])</span></code> should return <code class="docutils literal notranslate"><span class="pre">75</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="9"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def sort_and_get_median_num(lst_of_nums):
---
    lst_of_nums.sort()
---
    if len(lst_of_nums) % 2 == 1:
---
        return lst_of_nums[len(lst_of_nums) // 2]
---
        return lst_of_nums[len(lst_of_nums) / 2] #paired
---
    else:
---
        sum_of_middle_two_nums = sum(lst_of_nums[len(lst_of_nums) // 2 - 1: len(lst_of_nums) // 2 + 1])
---
        sum_of_middle_two_nums = sum(lst_of_nums[len(lst_of_nums) // 2: len(lst_of_nums) // 2 + 2]) #paired
---
        sum_of_middle_two_nums = sum(lst_of_nums[len(lst_of_nums) // 2 - 2: len(lst_of_nums) // 2]) #paired
---
        sum_of_middle_two_nums = sum(lst_of_nums[len(lst_of_nums) / 2 - 1: len(lst_of_nums) / 2 + 1]) #paired
---
        mean_of_two_nums = sum_of_middle_two_nums / 2
---
        mean_of_two_nums = sum_of_middle_two_nums // 2 #paired
---
        return mean_of_two_nums
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=functions_lists_mixedupcode_5_ac data-question_label="10">
<div id=functions_lists_mixedupcode_5_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">sort_and_get_median_num</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">lst_of_nums</span></code> as a parameter, sorts <code class="docutils literal notranslate"><span class="pre">lst_of_nums</span></code>, and returns the median of <code class="docutils literal notranslate"><span class="pre">lst_of_nums</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">sort_and_get_median_num([200,</span> <span class="pre">-5,</span> <span class="pre">0,</span> <span class="pre">75,</span> <span class="pre">80,</span> <span class="pre">60])</span></code> should return <code class="docutils literal notranslate"><span class="pre">67.5</span></code>, and <code class="docutils literal notranslate"><span class="pre">sort_and_get_median_num([200,</span> <span class="pre">-5,</span> <span class="pre">0,</span> <span class="pre">75,</span> <span class="pre">80])</span></code> should return <code class="docutils literal notranslate"><span class="pre">75</span></code>.</p>

</div>
<textarea data-lang="python" id="functions_lists_mixedupcode_5_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sort_and_get_median_num([200, -5, 0, 75, 80]),75,"sort_and_get_median_num([200, -5, 0, 75, 80])")
        self.assertEqual(sort_and_get_median_num([200, -5, 0, 75, 80, 60]),67.5,"sort_and_get_median_num([200, -5, 0, 75, 80, 60])")
        self.assertEqual(sort_and_get_median_num([200]),200,"sort_and_get_median_num([200])")
        self.assertEqual(sort_and_get_median_num([200, -5]),97.5,"sort_and_get_median_num([200, -5])")


myTests().main()
</textarea>
</div>
</div>
</section>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>



  
    {% if settings.google_ga %}
      <script>
        var _gaq = _gaq || [];
        _gaq.push(['_setAccount', '{{ settings.google_ga }}']);
        _gaq.push(['_trackPageview']);

        (function() {
          var ga = document.createElement('script'); ga.async = true;
          ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
          var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
        })();
      </script>
    {% endif %}
  

  
    <script type="application/json" class="js-hypothesis-config">
    {
      "showHighlights": false
    }
    </script>

    <script async src="https://hypothes.is/embed.js"></script>
    <!--provide space to prevent hypothes.is sidebar from obscuring content-->
    <style>@media screen and (max-width:767px) { body { padding-right: 21px; }}</style>

  


<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>