<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.17.1: http://docutils.sourceforge.net/" />

    <title>Function and String Mixed-Up Code Questions &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.cd688d1623c9932e.css?v=630908D8" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.bf73008a1eba9f4e.css?v=7DED055E" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/runtime.439d581f758f2b99.bundle.js?v=2E54442"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.51d9dde4e5162190.bundle.js?v=9CAA2D33"></script>
    <script src="../_static/runestone.42b576791a4ed54b.bundle.js?v=B7293199"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    
    eBookConfig.useRunestoneServices = {{use_services}};
    eBookConfig.host = '';
    eBookConfig.app = eBookConfig.host + '/runestone';
    eBookConfig.course = '{{ course_name }}';
    eBookConfig.basecourse = '{{ base_course }}';
    eBookConfig.isLoggedIn = {{ is_logged_in}};
    eBookConfig.email = '{{ user_email }}';
    eBookConfig.isInstructor = {{ is_instructor }};
    eBookConfig.username = '{{ user_id}}';
    eBookConfig.readings = {{ readings|safe}};
    eBookConfig.activities = {{ activity_info|safe }}
    eBookConfig.downloadsEnabled = {{downloads_enabled}};
    eBookConfig.allow_pairs = {{allow_pairs}}
    eBookConfig.enableCompareMe = {{enable_compare_me}};
    // _`new_server_prefix`: Defined by the server to be the path to the new server.
    eBookConfig.new_server_prefix = '{{ new_server_prefix }}';
    
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->

    
        {% if serve_ad is defined and serve_ad and settings.adsenseid %}
         <script data-ad-client="{{settings.adsenseid}}" async src="https://pagead2.googlesyndication.com/pagead/js/adsbygoogle.js"></script>
        {% endif %}
    



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
            
            {{ course_name }}
            
          
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="browsing_warning"></span></li>

        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li id="inst_peer_link"><a href='/runestone/peer/instructor.html'>Peer Instruction (Instructor)</a></li>            
            <li><a href='/runestone/peer/student.html'>Peer Instruction (Student)</a></li>            
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/ns/books/published/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=function_string_mixedUpCode_orig'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
            <script></script>
<li class="dropdown globaltoc-container">
    <a
        href="../index.html"
        class="dropdown-toggle"
        data-toggle="dropdown"
        aria-label="index"
        >This Chapter<b class="caret"></b
    ></a>
    <ul class="dropdown-menu globaltoc">
         
        {% for row in subchapter_list %}
        <li class="toctree-l1">
            <a class="reference internal" href="{{ row['subchap_uri']}}"
                >{{ row['title']}}</a
            >
        </li>
        {% endfor %} 
        
    </ul>
</li>
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->

    
        {% if settings.num_banners > 0 and settings.show_rs_banner %}
         <div style="width: 100%;">
            <a href="/runestone/default/donate?banner={{banner_num}}">
            <img class="runestone-banner"
                 src="/runestone/static/images/RunestoneBanner{{banner_num}}.png"
                 alt="Please Support Runestone"
                 style="display: block; width: 728px; margin-left: auto; margin-right: auto; margin-top: 10px;">
            </a>
         </div>
        {% endif %}
    


  
  <section id="function-and-string-mixed-up-code-questions">
<h1>Function and String Mixed-Up Code Questions<a class="headerlink" href="#function-and-string-mixed-up-code-questions" title="Permalink to this headline">¶</a></h1>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Function_String_MixedUpCode_1_v2" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">joinStrings(str1,</span> <span class="pre">str2)</span></code> that takes two different strings
as parameters, <code class="docutils literal notranslate"><span class="pre">str1</span></code> and <code class="docutils literal notranslate"><span class="pre">str2</span></code>, and returns them as one string with <code class="docutils literal notranslate"><span class="pre">str1</span></code> before <code class="docutils literal notranslate"><span class="pre">str2</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">joinStrings('he','llo')</span></code> returns <code class="docutils literal notranslate"><span class="pre">hello</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="1"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def joinStrings(str1, str2):
---
def joinStrings(): #paired
---
    joined = str1 + str2
---
    joined = str1.join(str2) #paired
---
    return joined
---
    Return joined #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Function_String_MixedUpCode_1_v2_ac data-question_label="2">
<div id=Function_String_MixedUpCode_1_v2_ac_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">joinStrings(str1,</span> <span class="pre">str2)</span></code> that takes two different strings
as parameters, <code class="docutils literal notranslate"><span class="pre">str1</span></code> and <code class="docutils literal notranslate"><span class="pre">str2</span></code>, and returns them as one string with <code class="docutils literal notranslate"><span class="pre">str1</span></code> before <code class="docutils literal notranslate"><span class="pre">str2</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">joinStrings('he','llo')</span></code> returns <code class="docutils literal notranslate"><span class="pre">hello</span></code>.</p>

</div>
<textarea data-lang="python" id="Function_String_MixedUpCode_1_v2_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(joinStrings('he', 'llo'),"hello","joinStrings('he', 'llo')")
        self.assertEqual(joinStrings('bye', 'bye'),"byebye","joinStrings('bye', 'bye')")
        self.assertEqual(joinStrings('Walk', 'Away'),"WalkAway","joinStrings('Walk', 'Away')")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Function_String_MixedUpCode_2_v2" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">rect_per(w,</span> <span class="pre">h)</span></code> that takes the width, <code class="docutils literal notranslate"><span class="pre">w</span></code>, and the height, <code class="docutils literal notranslate"><span class="pre">h</span></code>, as parameters. Calculate the perimeter, which is the width plus the height times two,
and return a string in the form - <code class="docutils literal notranslate"><span class="pre">Perimeter</span> <span class="pre">of</span> <span class="pre">rectangle</span> <span class="pre">with</span> <span class="pre">width</span> <span class="pre">of</span> <span class="pre">(w)</span> <span class="pre">and</span> <span class="pre">height</span> <span class="pre">of</span> <span class="pre">(h)</span> <span class="pre">is</span> <span class="pre">(perimeter)</span></code>.  For example, <code class="docutils literal notranslate"><span class="pre">rect_per(5,10)</span></code>
returns <code class="docutils literal notranslate"><span class="pre">Perimeter</span> <span class="pre">of</span> <span class="pre">rectangle</span> <span class="pre">with</span> <span class="pre">width</span> <span class="pre">of</span> <span class="pre">5</span> <span class="pre">and</span> <span class="pre">height</span> <span class="pre">of</span> <span class="pre">10</span> <span class="pre">is</span> <span class="pre">30</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="3"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def rect_per(w, h):
---
def rect_per(w, h) #paired
---
    perimeter = (w + h) * 2
---
    perimeter = w + h * 2 #paired
---
    answer = "Perimeter of rectangle with width of " + str(w) + " and height of " + str(h) + " is " + str(perimeter)
---
    answer = "Perimeter of rectangle with width of " + str(w) + " and height of " + str(h) + " is " + perimeter #paired
---
    return answer
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Function_String_MixedUpCode_2_v2_ac data-question_label="4">
<div id=Function_String_MixedUpCode_2_v2_ac_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">rect_per(w,</span> <span class="pre">h)</span></code> that takes the width, <code class="docutils literal notranslate"><span class="pre">w</span></code>, and the height, <code class="docutils literal notranslate"><span class="pre">h</span></code>, as parameters. Calculate the perimeter, which is the width plus the height times two,
and return a string in the form - <code class="docutils literal notranslate"><span class="pre">Perimeter</span> <span class="pre">of</span> <span class="pre">rectangle</span> <span class="pre">with</span> <span class="pre">width</span> <span class="pre">of</span> <span class="pre">(w)</span> <span class="pre">and</span> <span class="pre">height</span> <span class="pre">of</span> <span class="pre">(h)</span> <span class="pre">is</span> <span class="pre">(perimeter)</span></code>.  For example, <code class="docutils literal notranslate"><span class="pre">rect_per(5,10)</span></code>
returns <code class="docutils literal notranslate"><span class="pre">Perimeter</span> <span class="pre">of</span> <span class="pre">rectangle</span> <span class="pre">with</span> <span class="pre">width</span> <span class="pre">of</span> <span class="pre">5</span> <span class="pre">and</span> <span class="pre">height</span> <span class="pre">of</span> <span class="pre">10</span> <span class="pre">is</span> <span class="pre">30</span></code>.</p>

</div>
<textarea data-lang="python" id="Function_String_MixedUpCode_2_v2_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(rect_per(5, 10),"Perimeter of rectangle with width of 5 and height of 10 is 30","rect_per(5, 10)")
        self.assertEqual(rect_per(2, 3),"Perimeter of rectangle with width of 2 and height of 3 is 10","rect_per(2, 3)")
        self.assertEqual(rect_per(2.5, 3.5),"Perimeter of rectangle with width of 2.5 and height of 3.5 is 12.0","rect_per(2.5, 3.5)")
        self.assertEqual(rect_per(1.3, 2.6),"Perimeter of rectangle with width of 1.3 and height of 2.6 is 7.8","rect_per(1.3, 2.6)")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Function_String_MixedUpCode_3" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">cube(n)</span></code> that takes a number, <code class="docutils literal notranslate"><span class="pre">n</span></code>, and cubes it and returns a string in the form
<code class="docutils literal notranslate"><span class="pre">Cube</span> <span class="pre">of</span> <span class="pre">(n)</span> <span class="pre">is</span> <span class="pre">(n</span> <span class="pre">cubed)</span></code>.  For example, <code class="docutils literal notranslate"><span class="pre">cube(4)</span></code> returns <code class="docutils literal notranslate"><span class="pre">Cube</span> <span class="pre">of</span> <span class="pre">4</span> <span class="pre">is</span> <span class="pre">64</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="5"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def cube(n):
---
    cube_num = n ** 3
---
    cube_num = n * 3 #paired
---
    answer = "Cube of " + str(n) + " is " + str(cube_num)
---
    return answer
---
    return cube_num #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Function_String_MixedUpCode_3_ac data-question_label="6">
<div id=Function_String_MixedUpCode_3_ac_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">cube(n)</span></code> that takes a number, <code class="docutils literal notranslate"><span class="pre">n</span></code>, and cubes it and returns a string in the form
<code class="docutils literal notranslate"><span class="pre">Cube</span> <span class="pre">of</span> <span class="pre">(n)</span> <span class="pre">is</span> <span class="pre">(n</span> <span class="pre">cubed)</span></code>.  For example, <code class="docutils literal notranslate"><span class="pre">cube(4)</span></code> returns <code class="docutils literal notranslate"><span class="pre">Cube</span> <span class="pre">of</span> <span class="pre">4</span> <span class="pre">is</span> <span class="pre">64</span></code>.</p>

</div>
<textarea data-lang="python" id="Function_String_MixedUpCode_3_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(cube(3),"Cube of 3 is 27","cube(3)")
        self.assertEqual(cube(4),"Cube of 4 is 64","cube(4)")
        self.assertEqual(cube(5),"Cube of 5 is 125","cube(5)")
        self.assertEqual(cube(1),"Cube of 1 is 1","cube(1)")
        self.assertEqual(cube(1.5),"Cube of 1.5 is 3.375","cube(1.5)")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Function_String_MixedUpCode_4" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">atlas(c_name)</span></code> that takes a country name, <code class="docutils literal notranslate"><span class="pre">c_name</span></code>,
as a parameter and returns a string in the form of
<code class="docutils literal notranslate"><span class="pre">The</span> <span class="pre">last</span> <span class="pre">letter</span> <span class="pre">of</span> <span class="pre">(c_name)</span> <span class="pre">is</span> <span class="pre">(last)</span></code>.  For example, <code class="docutils literal notranslate"><span class="pre">atlas('Germany')</span></code> returns
<code class="docutils literal notranslate"><span class="pre">The</span> <span class="pre">last</span> <span class="pre">letter</span> <span class="pre">of</span> <span class="pre">Germany</span> <span class="pre">is</span> <span class="pre">y</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="7"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def atlas(c_name):
---
    last = c_name[-1]
---
    last = c_name[:-1] #paired
---
    return("The last letter of " + c_name + " is " + last)
---
    return("The last letter of" + c_name + "is" + last) #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Function_String_MixedUpCode_4_ac data-question_label="8">
<div id=Function_String_MixedUpCode_4_ac_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">atlas(c_name)</span></code> that takes a country name, <code class="docutils literal notranslate"><span class="pre">c_name</span></code>,
as a parameter and returns a string in the form of
<code class="docutils literal notranslate"><span class="pre">The</span> <span class="pre">last</span> <span class="pre">letter</span> <span class="pre">of</span> <span class="pre">(c_name)</span> <span class="pre">is</span> <span class="pre">(last)</span></code>.  For example, <code class="docutils literal notranslate"><span class="pre">atlas('Germany')</span></code> returns
<code class="docutils literal notranslate"><span class="pre">The</span> <span class="pre">last</span> <span class="pre">letter</span> <span class="pre">of</span> <span class="pre">Germany</span> <span class="pre">is</span> <span class="pre">y</span></code>.</p>

</div>
<textarea data-lang="python" id="Function_String_MixedUpCode_4_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(atlas('Germany'),"The last letter of Germany is y","atlas('Germany')")
        self.assertEqual(atlas('USA'),"The last letter of USA is A","atlas('USA')")
        self.assertEqual(atlas('India'),"The last letter of India is a","atlas('India')")
        self.assertEqual(atlas('China'),"The last letter of China is a","atlas('China')")
        self.assertEqual(atlas('Ghana'),"The last letter of Ghana is a","atlas('Ghana')")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Function_String_MixedUpCode_5" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">dateYear(date)</span></code> that takes a string, <code class="docutils literal notranslate"><span class="pre">date</span></code>, in the format MM/DD/YYYY and returns the year.
For example, <code class="docutils literal notranslate"><span class="pre">dateYear(11/07/2000)</span></code> should return <code class="docutils literal notranslate"><span class="pre">2000</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="9"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def dateYear(date):
---
    year = date[6:]
---
    year = date[7:] #paired
---
    return year
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Function_String_MixedUpCode_5_ac data-question_label="10">
<div id=Function_String_MixedUpCode_5_ac_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">dateYear(date)</span></code> that takes a string, <code class="docutils literal notranslate"><span class="pre">date</span></code>, in the format MM/DD/YYYY and returns the year.
For example, <code class="docutils literal notranslate"><span class="pre">dateYear(11/07/2000)</span></code> should return <code class="docutils literal notranslate"><span class="pre">2000</span></code>.</p>

</div>
<textarea data-lang="python" id="Function_String_MixedUpCode_5_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(dateYear('03/07/1961'),"1961","dateYear('03/07/1961')")
        self.assertEqual(dateYear('12/24/1945'),"1945","dateYear('12/24/1945')")
        self.assertEqual(dateYear('11/07/2000'),"2000","dateYear('11/07/2000')")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Function_String_MixedUpCode_6" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">capitalize(str)</span></code>, that takes a string, <code class="docutils literal notranslate"><span class="pre">str</span></code>, and returns the string with the first letter capitalized.
For example, <code class="docutils literal notranslate"><span class="pre">capitalize('america')</span></code> would return <code class="docutils literal notranslate"><span class="pre">America</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="11"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def capitalize(str):
---
Def Capitalize(str): #paired
---
    answer = str[0].upper() + str[1:]
---
    answer = str.upper() + str[1:] #paired
---
    return answer
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Function_String_MixedUpCode_6_ac data-question_label="12">
<div id=Function_String_MixedUpCode_6_ac_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">capitalize(str)</span></code>, that takes a string, <code class="docutils literal notranslate"><span class="pre">str</span></code>, and returns the string with the first letter capitalized.
For example, <code class="docutils literal notranslate"><span class="pre">capitalize('america')</span></code> would return <code class="docutils literal notranslate"><span class="pre">America</span></code>.</p>

</div>
<textarea data-lang="python" id="Function_String_MixedUpCode_6_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(capitalize('america'),"America","capitalize('america')")
        self.assertEqual(capitalize('walk'),"Walk","capitalize('walk')")
        self.assertEqual(capitalize('Bend'),"Bend","capitalize('Bend')")
        self.assertEqual(capitalize('flower'),"Flower","capitalize('flower')")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Function_String_MixedUpCode_7" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">mod_password(password)</span></code>, that takes a string, <code class="docutils literal notranslate"><span class="pre">password</span></code>, and returns a new string that replaces all <code class="docutils literal notranslate"><span class="pre">'s'</span></code> in the string with <code class="docutils literal notranslate"><span class="pre">'$'</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">mod_password('suspense')</span></code> returns <code class="docutils literal notranslate"><span class="pre">$u$pen$e</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="13"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def mod_password(password):
---
    new = password.replace('s', '$')
---
    new = password.remove('s') + '$' #paired
---
    return new
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Function_String_MixedUpCode_7_ac data-question_label="14">
<div id=Function_String_MixedUpCode_7_ac_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">mod_password(password)</span></code>, that takes a string, <code class="docutils literal notranslate"><span class="pre">password</span></code>, and returns a new string that replaces all <code class="docutils literal notranslate"><span class="pre">'s'</span></code> in the string with <code class="docutils literal notranslate"><span class="pre">'$'</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">mod_password('suspense')</span></code> returns <code class="docutils literal notranslate"><span class="pre">$u$pen$e</span></code>.</p>

</div>
<textarea data-lang="python" id="Function_String_MixedUpCode_7_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(mod_password('suspense'),"$u$pen$e","mod_password('suspense')")
        self.assertEqual(mod_password('sue'),"$ue","mod_password('sue')")
        self.assertEqual(mod_password('sass'),"$a$$","mod_password('sass')")
        self.assertEqual(mod_password('away'),"away","mod_password('away')")
        self.assertEqual(mod_password('apples'),"apple$","mod_password('apples')")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Function_String_MixedUpCode_8" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">first_last(str)</span></code>, that takes a string, <code class="docutils literal notranslate"><span class="pre">str</span></code>, and returns a new string with the first two characters of the word
followed by the last two characters. Assume <code class="docutils literal notranslate"><span class="pre">str</span></code> is four characters or more.
For example, <code class="docutils literal notranslate"><span class="pre">first_last('wander')</span></code> returns <code class="docutils literal notranslate"><span class="pre">waer</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="15"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def first_last(str):
---
    out = str[:2] + str[-2:]
---
    out = str[-2:] + str[:2] #paired
---
    return out
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Function_String_MixedUpCode_8_ac data-question_label="16">
<div id=Function_String_MixedUpCode_8_ac_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">first_last(str)</span></code>, that takes a string, <code class="docutils literal notranslate"><span class="pre">str</span></code>, and returns a new string with the first two characters of the word
followed by the last two characters. Assume <code class="docutils literal notranslate"><span class="pre">str</span></code> is four characters or more.
For example, <code class="docutils literal notranslate"><span class="pre">first_last('wander')</span></code> returns <code class="docutils literal notranslate"><span class="pre">waer</span></code>.</p>

</div>
<textarea data-lang="python" id="Function_String_MixedUpCode_8_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(first_last('wander'),"waer","first_last('wander')")
        self.assertEqual(first_last('hope'),"hope","first_last('hope')")
        self.assertEqual(first_last('through'),"thgh","first_last('through')")

myTests().main()
</textarea>
</div>
</div>
</section>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>



  
    {% if settings.google_ga %}
      <script>
        var _gaq = _gaq || [];
        _gaq.push(['_setAccount', '{{ settings.google_ga }}']);
        _gaq.push(['_trackPageview']);

        (function() {
          var ga = document.createElement('script'); ga.async = true;
          ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
          var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
        })();
      </script>
    {% endif %}
  

  
    <script type="application/json" class="js-hypothesis-config">
    {
      "showHighlights": false
    }
    </script>

    <script async src="https://hypothes.is/embed.js"></script>
    <!--provide space to prevent hypothes.is sidebar from obscuring content-->
    <style>@media screen and (max-width:767px) { body { padding-right: 21px; }}</style>

  


<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>