<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.17.1: http://docutils.sourceforge.net/" />

    <title>5.28. Functions and Conditionals Write Code Questions &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.cd688d1623c9932e.css?v=630908D8" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.bf73008a1eba9f4e.css?v=7DED055E" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/runtime.439d581f758f2b99.bundle.js?v=2E54442"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.51d9dde4e5162190.bundle.js?v=9CAA2D33"></script>
    <script src="../_static/runestone.42b576791a4ed54b.bundle.js?v=B7293199"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="5.29. Group Work: Functions and Lists" href="funcWithLists.html" />
    <link rel="prev" title="5.27. Functions and Conditionals Mixed-Up Code Questions" href="function_conditional_mixedupcode.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    
    eBookConfig.useRunestoneServices = {{use_services}};
    eBookConfig.host = '';
    eBookConfig.app = eBookConfig.host + '/runestone';
    eBookConfig.course = '{{ course_name }}';
    eBookConfig.basecourse = '{{ base_course }}';
    eBookConfig.isLoggedIn = {{ is_logged_in}};
    eBookConfig.email = '{{ user_email }}';
    eBookConfig.isInstructor = {{ is_instructor }};
    eBookConfig.username = '{{ user_id}}';
    eBookConfig.readings = {{ readings|safe}};
    eBookConfig.activities = {{ activity_info|safe }}
    eBookConfig.downloadsEnabled = {{downloads_enabled}};
    eBookConfig.allow_pairs = {{allow_pairs}}
    eBookConfig.enableCompareMe = {{enable_compare_me}};
    // _`new_server_prefix`: Defined by the server to be the path to the new server.
    eBookConfig.new_server_prefix = '{{ new_server_prefix }}';
    
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->

    
        {% if serve_ad is defined and serve_ad and settings.adsenseid %}
         <script data-ad-client="{{settings.adsenseid}}" async src="https://pagead2.googlesyndication.com/pagead/js/adsbygoogle.js"></script>
        {% endif %}
    



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
            
            {{ course_name }}
            
          
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="browsing_warning"></span></li>

        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li id="inst_peer_link"><a href='/runestone/peer/instructor.html'>Peer Instruction (Instructor)</a></li>            
            <li><a href='/runestone/peer/student.html'>Peer Instruction (Student)</a></li>            
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/ns/books/published/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=function_conditional_writecode'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
            <script></script>
<li class="dropdown globaltoc-container">
    <a
        href="../index.html"
        class="dropdown-toggle"
        data-toggle="dropdown"
        aria-label="index"
        >This Chapter<b class="caret"></b
    ></a>
    <ul class="dropdown-menu globaltoc">
         
        {% for row in subchapter_list %}
        <li class="toctree-l1">
            <a class="reference internal" href="{{ row['subchap_uri']}}"
                >{{ row['title']}}</a
            >
        </li>
        {% endfor %} 
        
    </ul>
</li>
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">5.27. </span>Functions and Conditionals Mixed-Up Code Questions" data-toggle="tooltip">
    <a href="function_conditional_mixedupcode.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">5.29. </span>Group Work: Functions and Lists' data-toggle="tooltip" >              
    <a href="funcWithLists.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->

    
        {% if settings.num_banners > 0 and settings.show_rs_banner %}
         <div style="width: 100%;">
            <a href="/runestone/default/donate?banner={{banner_num}}">
            <img class="runestone-banner"
                 src="/runestone/static/images/RunestoneBanner{{banner_num}}.png"
                 alt="Please Support Runestone"
                 style="display: block; width: 728px; margin-left: auto; margin-right: auto; margin-top: 10px;">
            </a>
         </div>
        {% endif %}
    


  
  <section id="functions-and-conditionals-write-code-questions">
<h1><span class="section-number">5.28. </span>Functions and Conditionals Write Code Questions<a class="headerlink" href="#functions-and-conditionals-write-code-questions" title="Permalink to this headline">¶</a></h1>
<ol class="arabic">
<li><div id='funct_conditional_writecode1' data-component="tabbedStuff"  class='alert alert-warning '><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=funct_conditional_writecode1q data-question_label="5.28.1">
<div id=funct_conditional_writecode1q_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">cropped_photo</span></code> that takes the <code class="docutils literal notranslate"><span class="pre">height</span></code> and <code class="docutils literal notranslate"><span class="pre">width</span></code> of a photo.
If the <code class="docutils literal notranslate"><span class="pre">height</span></code> and <code class="docutils literal notranslate"><span class="pre">width</span></code> are equal, then the function returns <code class="docutils literal notranslate"><span class="pre">&quot;Cropped</span> <span class="pre">Photo&quot;</span></code>. Otherwise,
it returns <code class="docutils literal notranslate"><span class="pre">&quot;The</span> <span class="pre">photo</span> <span class="pre">needs</span> <span class="pre">to</span> <span class="pre">be</span> <span class="pre">cropped.&quot;</span></code> Assume <code class="docutils literal notranslate"><span class="pre">height</span></code> and <code class="docutils literal notranslate"><span class="pre">width</span></code> are positive. For example,
<code class="docutils literal notranslate"><span class="pre">cropped_photo(23.2,</span> <span class="pre">23.2)</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;Cropped</span> <span class="pre">Photo&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_conditional_writecode1q_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def cropped_photo(height, width):
    # write code here

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(cropped_photo(30, 80),"The photo needs to be cropped.","cropped_photo(30, 80)")
        self.assertEqual(cropped_photo(50, 50),"Cropped Photo","cropped_photo(50, 50)")
        self.assertEqual(cropped_photo(50.5, 50),"The photo needs to be cropped.","cropped_photo(50.5, 50)")
        self.assertEqual(cropped_photo(23.2, 23.2),"Cropped Photo","cropped_photo(23.2, 23.2)")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=funct_conditional_writecode1a data-question_label="5.28.2">
<div id=funct_conditional_writecode1a_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">cropped_photo</span></code> that takes the <code class="docutils literal notranslate"><span class="pre">height</span></code> and <code class="docutils literal notranslate"><span class="pre">width</span></code> of a photo.
If the <code class="docutils literal notranslate"><span class="pre">height</span></code> and <code class="docutils literal notranslate"><span class="pre">width</span></code> are equal, then the function returns <code class="docutils literal notranslate"><span class="pre">&quot;Cropped</span> <span class="pre">Photo&quot;</span></code>. Otherwise,
it returns <code class="docutils literal notranslate"><span class="pre">&quot;The</span> <span class="pre">photo</span> <span class="pre">needs</span> <span class="pre">to</span> <span class="pre">be</span> <span class="pre">cropped.&quot;</span></code> Assume <code class="docutils literal notranslate"><span class="pre">height</span></code> and <code class="docutils literal notranslate"><span class="pre">width</span></code> are positive. For example,
<code class="docutils literal notranslate"><span class="pre">cropped_photo(23.2,</span> <span class="pre">23.2)</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;Cropped</span> <span class="pre">Photo&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_conditional_writecode1a_editor" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def cropped_photo(height, width):
    if height == width:
        return "Cropped Photo"
    else:
        return "The photo needs to be cropped."
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=funct_conditional_writecode2q data-question_label="5.28.3">
<div id=funct_conditional_writecode2q_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">bonus</span></code> that takes two parameters, <code class="docutils literal notranslate"><span class="pre">years</span></code> and <code class="docutils literal notranslate"><span class="pre">quality</span></code>.
If <code class="docutils literal notranslate"><span class="pre">years</span></code> is more than 5, then the person gets a bonus.
The person also gets a bonus if their <code class="docutils literal notranslate"><span class="pre">quality</span></code> of work was more than 90.
They can also get a bonus if they worked for at least 3 <code class="docutils literal notranslate"><span class="pre">years</span></code> and had <code class="docutils literal notranslate"><span class="pre">quality</span></code> more than 80.
If the person is eligible for a bonus, return <code class="docutils literal notranslate"><span class="pre">&quot;Eligible</span> <span class="pre">for</span> <span class="pre">Bonus&quot;</span></code>. Otherwise, return <code class="docutils literal notranslate"><span class="pre">&quot;Ineligible</span> <span class="pre">for</span> <span class="pre">Bonus&quot;</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">bonus(3,</span> <span class="pre">81)</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;Eligible</span> <span class="pre">for</span> <span class="pre">Bonus&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_conditional_writecode2q_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def bonus(years, quality):
    # write code here

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(bonus(7, 40),"Eligible for Bonus","bonus(7, 40)")
        self.assertEqual(bonus(5, 40),"Ineligible for Bonus","bonus(5, 40)")
        self.assertEqual(bonus(2, 98),"Eligible for Bonus","bonus(2, 98)")
        self.assertEqual(bonus(2, 90),"Ineligible for Bonus","bonus(2, 90)")
        self.assertEqual(bonus(3, 81),"Eligible for Bonus","bonus(3, 81)")
        self.assertEqual(bonus(4, 81),"Eligible for Bonus","bonus(4, 81)")
        self.assertEqual(bonus(3, 67),"Ineligible for Bonus","bonus(3, 67)")
        self.assertEqual(bonus(3, 80),"Ineligible for Bonus","bonus(3, 80)")
        self.assertEqual(bonus(2, 88),"Ineligible for Bonus","bonus(2, 88)")

myTests().main()
</textarea>
</div>
</div>
</li>
<li><div id='funct_conditional_writecode3' data-component="tabbedStuff"  class='alert alert-warning '><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=funct_conditional_writecode3q data-question_label="5.28.4">
<div id=funct_conditional_writecode3q_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">move_elevator</span></code> that takes two parameters, <code class="docutils literal notranslate"><span class="pre">current_floor</span></code> and <code class="docutils literal notranslate"><span class="pre">next_floor</span></code>. If the elevator moves
to a floor above, then it should return <code class="docutils literal notranslate"><span class="pre">&quot;Up&quot;</span></code>. Otherwise, it should return <code class="docutils literal notranslate"><span class="pre">&quot;Down&quot;</span></code>. Also, if <code class="docutils literal notranslate"><span class="pre">next_floor</span></code> is the same as <code class="docutils literal notranslate"><span class="pre">current_floor</span></code>, it should return 0.
Assume there are no 0 or negative floors. For example, <code class="docutils literal notranslate"><span class="pre">move_elevator(5,</span> <span class="pre">8)</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;Up&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_conditional_writecode3q_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def move_elevator(current_floor, next_floor):
    # write code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(move_elevator(5, 8),"Up","move_elevator(5, 8)")
        self.assertEqual(move_elevator(2, 2),0,"move_elevator(2, 2)")
        self.assertEqual(move_elevator(9, 3),"Down","move_elevator(9, 3)")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=funct_conditional_writecode3a data-question_label="5.28.5">
<div id=funct_conditional_writecode3a_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">move_elevator</span></code> that takes two parameters, <code class="docutils literal notranslate"><span class="pre">current_floor</span></code> and <code class="docutils literal notranslate"><span class="pre">next_floor</span></code>. If the elevator moves
to a floor above, then it should return <code class="docutils literal notranslate"><span class="pre">&quot;Up&quot;</span></code>. Otherwise, it should return <code class="docutils literal notranslate"><span class="pre">&quot;Down&quot;</span></code>. Also, if <code class="docutils literal notranslate"><span class="pre">next_floor</span></code> is the same as <code class="docutils literal notranslate"><span class="pre">current_floor</span></code>, it should return 0.
Assume there are no 0 or negative floors. For example, <code class="docutils literal notranslate"><span class="pre">move_elevator(5,</span> <span class="pre">8)</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;Up&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_conditional_writecode3a_editor" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def move_elevator(current_floor, next_floor):
    if next_floor < current_floor:
        return "Down"
    elif next_floor > current_floor:
        return "Up"
    else:
        return 0
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=funct_conditional_writecode4q data-question_label="5.28.6">
<div id=funct_conditional_writecode4q_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">lunch_break</span></code> that takes a parameter <code class="docutils literal notranslate"><span class="pre">class_standing</span></code> and determines how long a lunch break will be based on
the student’s class standing. If a person is a ‘Freshman’, return <code class="docutils literal notranslate"><span class="pre">40</span></code>. If a person is a ‘Sophomore’, return <code class="docutils literal notranslate"><span class="pre">30</span></code>.
If a person is a ‘Junior’, return <code class="docutils literal notranslate"><span class="pre">20</span></code>. If a person is a ‘Senior’, return <code class="docutils literal notranslate"><span class="pre">15</span></code>. If a person is none of these, return <code class="docutils literal notranslate"><span class="pre">0</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">lunch_break('Alumni')</span></code> should return <code class="docutils literal notranslate"><span class="pre">0</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_conditional_writecode4q_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def lunch_break(class_standing):
    # write code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(lunch_break('Freshman'),40,"lunch_break('Freshman')")
        self.assertEqual(lunch_break('Sophomore'),30,"lunch_break('Sophomore')")
        self.assertEqual(lunch_break('Junior'),20,"lunch_break('Junior')")
        self.assertEqual(lunch_break('Senior'),15,"lunch_break('Senior')")
        self.assertEqual(lunch_break('Alumni'),0,"lunch_break('Alumni')")

myTests().main()
</textarea>
</div>
</div>
</li>
<li><div id='funct_conditional_writecode5' data-component="tabbedStuff"  class='alert alert-warning '><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=funct_conditional_writecode5q data-question_label="5.28.7">
<div id=funct_conditional_writecode5q_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">pay_rent</span></code> that takes in one parameter,
<code class="docutils literal notranslate"><span class="pre">units</span></code>. If <code class="docutils literal notranslate"><span class="pre">units</span></code> is less than 200, return <code class="docutils literal notranslate"><span class="pre">&quot;$100&quot;</span></code>.
If <code class="docutils literal notranslate"><span class="pre">units</span></code> is less than 500 and greater than or equal to 200, return <code class="docutils literal notranslate"><span class="pre">&quot;$200&quot;</span></code>. Otherwise, return <code class="docutils literal notranslate"><span class="pre">&quot;$300&quot;</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">pay_rent(175)</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;$100&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_conditional_writecode5q_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def pay_rent(units):
    # write code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(pay_rent(175), "$100","pay_rent(175)")
        self.assertEqual(pay_rent(199), "$100","pay_rent(199)")
        self.assertEqual(pay_rent(200), "$200","pay_rent(200)")
        self.assertEqual(pay_rent(201), "$200","pay_rent(201)")
        self.assertEqual(pay_rent(373), "$200","pay_rent(373)")
        self.assertEqual(pay_rent(499), "$200","pay_rent(499)")
        self.assertEqual(pay_rent(500), "$300","pay_rent(500)")
        self.assertEqual(pay_rent(501), "$300","pay_rent(501)")
        self.assertEqual(pay_rent(1007), "$300","pay_rent(1007)")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=funct_conditional_writecode5a data-question_label="5.28.8">
<div id=funct_conditional_writecode5a_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">pay_rent</span></code> that takes in one parameter,
<code class="docutils literal notranslate"><span class="pre">units</span></code>. If <code class="docutils literal notranslate"><span class="pre">units</span></code> is less than 200, return <code class="docutils literal notranslate"><span class="pre">&quot;$100&quot;</span></code>.
If <code class="docutils literal notranslate"><span class="pre">units</span></code> is less than 500 and greater than or equal to 200, return <code class="docutils literal notranslate"><span class="pre">&quot;$200&quot;</span></code>. Otherwise, return <code class="docutils literal notranslate"><span class="pre">&quot;$300&quot;</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">pay_rent(175)</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;$100&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_conditional_writecode5a_editor" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def pay_rent(units):
    if units < 200:
        return "$100"
    elif units < 500:
        return "$200"
    else:
        return "$300"
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=funct_conditional_writecode6q data-question_label="5.28.9">
<div id=funct_conditional_writecode6q_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">play_weather</span></code> that takes the string <code class="docutils literal notranslate"><span class="pre">weather</span></code> as a parameter and
returns if one should play or not. If the <code class="docutils literal notranslate"><span class="pre">weather</span></code> is ‘sunny’ or ‘cloudy’, then the function
should return <code class="docutils literal notranslate"><span class="pre">&quot;The</span> <span class="pre">person</span> <span class="pre">should</span> <span class="pre">play&quot;</span></code>. If the <code class="docutils literal notranslate"><span class="pre">weather</span></code> is ‘rainy’, ‘windy’ or ‘snowy’,
then the function should return <code class="docutils literal notranslate"><span class="pre">&quot;The</span> <span class="pre">person</span> <span class="pre">should</span> <span class="pre">not</span> <span class="pre">play&quot;</span></code>. If the <code class="docutils literal notranslate"><span class="pre">weather</span></code> is none of these, the
function should return <code class="docutils literal notranslate"><span class="pre">&quot;It</span> <span class="pre">is</span> <span class="pre">not</span> <span class="pre">safe</span> <span class="pre">to</span> <span class="pre">go</span> <span class="pre">outside&quot;</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">play_weather('sunny')</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;The</span> <span class="pre">person</span> <span class="pre">should</span> <span class="pre">play&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_conditional_writecode6q_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def play_weather(weather):
    # write code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(play_weather('sunny'), "The person should play", "play_weather('sunny')")
        self.assertEqual(play_weather('cloudy'), "The person should play", "play_weather('cloudy')")
        self.assertEqual(play_weather('rainy'), "The person should not play", "play_weather('rainy')")
        self.assertEqual(play_weather('windy'), "The person should not play", "play_weather('windy')")
        self.assertEqual(play_weather('snowy'), "The person should not play", "play_weather('snowy')")
        self.assertEqual(play_weather('stormy'), "It is not safe to go outside", "play_weather('stormy')")

myTests().main()
</textarea>
</div>
</div>
</li>
</ol>
</section>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 5.27. Functions and Conditionals Mixed-Up Code Questions' data-toggle="tooltip">
  <a href="function_conditional_mixedupcode.html" aria-label="Previous - 5.27. Functions and Conditionals Mixed-Up Code Questions">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 5.29. Group Work: Functions and Lists' data-toggle="tooltip" >
    <a href="funcWithLists.html" aria-label="Next - 5.29. Group Work: Functions and Lists">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>



  
    {% if settings.google_ga %}
      <script>
        var _gaq = _gaq || [];
        _gaq.push(['_setAccount', '{{ settings.google_ga }}']);
        _gaq.push(['_trackPageview']);

        (function() {
          var ga = document.createElement('script'); ga.async = true;
          ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
          var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
        })();
      </script>
    {% endif %}
  

  
    <script type="application/json" class="js-hypothesis-config">
    {
      "showHighlights": false
    }
    </script>

    <script async src="https://hypothes.is/embed.js"></script>
    <!--provide space to prevent hypothes.is sidebar from obscuring content-->
    <style>@media screen and (max-width:767px) { body { padding-right: 21px; }}</style>

  


<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>