<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.17.1: http://docutils.sourceforge.net/" />

    <title>Functions and Conditionals Mixed-Up Code Questions &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.cd688d1623c9932e.css?v=630908D8" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.bf73008a1eba9f4e.css?v=7DED055E" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/runtime.439d581f758f2b99.bundle.js?v=2E54442"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.51d9dde4e5162190.bundle.js?v=9CAA2D33"></script>
    <script src="../_static/runestone.42b576791a4ed54b.bundle.js?v=B7293199"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    
    eBookConfig.useRunestoneServices = {{use_services}};
    eBookConfig.host = '';
    eBookConfig.app = eBookConfig.host + '/runestone';
    eBookConfig.course = '{{ course_name }}';
    eBookConfig.basecourse = '{{ base_course }}';
    eBookConfig.isLoggedIn = {{ is_logged_in}};
    eBookConfig.email = '{{ user_email }}';
    eBookConfig.isInstructor = {{ is_instructor }};
    eBookConfig.username = '{{ user_id}}';
    eBookConfig.readings = {{ readings|safe}};
    eBookConfig.activities = {{ activity_info|safe }}
    eBookConfig.downloadsEnabled = {{downloads_enabled}};
    eBookConfig.allow_pairs = {{allow_pairs}}
    eBookConfig.enableCompareMe = {{enable_compare_me}};
    // _`new_server_prefix`: Defined by the server to be the path to the new server.
    eBookConfig.new_server_prefix = '{{ new_server_prefix }}';
    
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->

    
        {% if serve_ad is defined and serve_ad and settings.adsenseid %}
         <script data-ad-client="{{settings.adsenseid}}" async src="https://pagead2.googlesyndication.com/pagead/js/adsbygoogle.js"></script>
        {% endif %}
    



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
            
            {{ course_name }}
            
          
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="browsing_warning"></span></li>

        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li id="inst_peer_link"><a href='/runestone/peer/instructor.html'>Peer Instruction (Instructor)</a></li>            
            <li><a href='/runestone/peer/student.html'>Peer Instruction (Student)</a></li>            
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/ns/books/published/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=function_conditional_mixedupcode_orig'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
            <script></script>
<li class="dropdown globaltoc-container">
    <a
        href="../index.html"
        class="dropdown-toggle"
        data-toggle="dropdown"
        aria-label="index"
        >This Chapter<b class="caret"></b
    ></a>
    <ul class="dropdown-menu globaltoc">
         
        {% for row in subchapter_list %}
        <li class="toctree-l1">
            <a class="reference internal" href="{{ row['subchap_uri']}}"
                >{{ row['title']}}</a
            >
        </li>
        {% endfor %} 
        
    </ul>
</li>
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->

    
        {% if settings.num_banners > 0 and settings.show_rs_banner %}
         <div style="width: 100%;">
            <a href="/runestone/default/donate?banner={{banner_num}}">
            <img class="runestone-banner"
                 src="/runestone/static/images/RunestoneBanner{{banner_num}}.png"
                 alt="Please Support Runestone"
                 style="display: block; width: 728px; margin-left: auto; margin-right: auto; margin-top: 10px;">
            </a>
         </div>
        {% endif %}
    


  
  <section id="functions-and-conditionals-mixed-up-code-questions">
<h1>Functions and Conditionals Mixed-Up Code Questions<a class="headerlink" href="#functions-and-conditionals-mixed-up-code-questions" title="Permalink to this headline">¶</a></h1>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="funct_conditional_mixedupcode_1" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">get_grade</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">percentage</span></code> and returns a string variable called <code class="docutils literal notranslate"><span class="pre">letter_grade</span></code>.
If the <code class="docutils literal notranslate"><span class="pre">percentage</span></code> is greater than 100, return <code class="docutils literal notranslate"><span class="pre">&quot;You</span> <span class="pre">can't</span> <span class="pre">get</span> <span class="pre">over</span> <span class="pre">100.&quot;</span></code>. Otherwise, it should return <code class="docutils literal notranslate"><span class="pre">'A'</span></code> if the score is greater than or equal to 90,
<code class="docutils literal notranslate"><span class="pre">'B'</span></code> if greater than or equal to 80, <code class="docutils literal notranslate"><span class="pre">'C'</span></code> if greater than or equal to 70, <code class="docutils literal notranslate"><span class="pre">'D'</span></code> if greater than or equal to 60, and otherwise <code class="docutils literal notranslate"><span class="pre">'E'</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">get_grade(90)</span></code> should return <code class="docutils literal notranslate"><span class="pre">'A'</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="1"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def get_grade(percentage):
---
def get_grade(percentage) #paired
---
    if percentage > 100:
---
        return "You can't get over 100."
---
        return 'You can't get over 100.' #paired
---
    elif percentage <= 100 and percentage >= 90:
---
        letter_grade = 'A'
---
        letter_grade == 'A' #paired
---
    elif percentage >= 80:
---
        letter_grade = 'B'
---
    elif percentage >= 70:
---
        letter_grade = 'C'
---
    elif percentage >= 60:
---
        letter_grade = 'D'
---
    else:
        letter_grade = 'E'
---
    return letter_grade
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=funct_conditional_mixedupcode_1_ac data-question_label="2">
<div id=funct_conditional_mixedupcode_1_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">get_grade</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">percentage</span></code> and returns a string variable called <code class="docutils literal notranslate"><span class="pre">letter_grade</span></code>.
If the <code class="docutils literal notranslate"><span class="pre">percentage</span></code> is greater than 100, return <code class="docutils literal notranslate"><span class="pre">&quot;You</span> <span class="pre">can't</span> <span class="pre">get</span> <span class="pre">over</span> <span class="pre">100.&quot;</span></code>. Otherwise, it should return <code class="docutils literal notranslate"><span class="pre">'A'</span></code> if the score is greater than or equal to 90,
<code class="docutils literal notranslate"><span class="pre">'B'</span></code> if greater than or equal to 80, <code class="docutils literal notranslate"><span class="pre">'C'</span></code> if greater than or equal to 70, <code class="docutils literal notranslate"><span class="pre">'D'</span></code> if greater than or equal to 60, and otherwise <code class="docutils literal notranslate"><span class="pre">'E'</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">get_grade(90)</span></code> should return <code class="docutils literal notranslate"><span class="pre">'A'</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_conditional_mixedupcode_1_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def get_grade(percentage):
    # write code here

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(get_grade(101),"You can't get over 100.","get_grade(101)")
        self.assertEqual(get_grade(100),"A","get_grade(100)")
        self.assertEqual(get_grade(95),"A","get_grade(95)")
        self.assertEqual(get_grade(90),"A","get_grade(90)")
        self.assertEqual(get_grade(89),"B","get_grade(89)")
        self.assertEqual(get_grade(80),"B","get_grade(80)")
        self.assertEqual(get_grade(79),"C","get_grade(79)")
        self.assertEqual(get_grade(70),"C","get_grade(70)")
        self.assertEqual(get_grade(69),"D","get_grade(69)")
        self.assertEqual(get_grade(60),"D","get_grade(60)")
        self.assertEqual(get_grade(59),"E","get_grade(59)")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="funct_conditional_mixedupcode_2" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">even_or_odd</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">num</span></code>. If <code class="docutils literal notranslate"><span class="pre">num</span></code> is odd, return <code class="docutils literal notranslate"><span class="pre">&quot;This</span> <span class="pre">is</span> <span class="pre">odd.&quot;</span></code>, and if <code class="docutils literal notranslate"><span class="pre">num</span></code> is even, return <code class="docutils literal notranslate"><span class="pre">&quot;This</span> <span class="pre">is</span> <span class="pre">even.&quot;</span> <span class="pre">``.</span>
<span class="pre">For</span> <span class="pre">example,</span> <span class="pre">``even_or_odd(-65)</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;This</span> <span class="pre">is</span> <span class="pre">odd.&quot;</span></code>. Note: For this function, <code class="docutils literal notranslate"><span class="pre">num</span></code> is solely an integer.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="3"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def even_or_odd(num):
---
Def even_or_odd(num): #paired
---
    if num % 2 == 0:
---
    if num % 2 = 1: #paired
---
    if num % 2 = 0: #paired
---
        return "This is even."
---
    else:
---
        return "This is odd."
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=funct_conditional_mixedupcode_2_ac data-question_label="4">
<div id=funct_conditional_mixedupcode_2_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">even_or_odd</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">num</span></code>. If <code class="docutils literal notranslate"><span class="pre">num</span></code> is odd, return <code class="docutils literal notranslate"><span class="pre">&quot;This</span> <span class="pre">is</span> <span class="pre">odd.&quot;</span></code>, and if <code class="docutils literal notranslate"><span class="pre">num</span></code> is even, return <code class="docutils literal notranslate"><span class="pre">&quot;This</span> <span class="pre">is</span> <span class="pre">even.&quot;</span> <span class="pre">``.</span>
<span class="pre">For</span> <span class="pre">example,</span> <span class="pre">``even_or_odd(-65)</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;This</span> <span class="pre">is</span> <span class="pre">odd.&quot;</span></code>. Note: For this function, <code class="docutils literal notranslate"><span class="pre">num</span></code> is solely an integer.</p>

</div>
<textarea data-lang="python" id="funct_conditional_mixedupcode_2_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def even_or_odd(num):
    # write code here

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(even_or_odd(100),"This is even.","even_or_odd(100)")
        self.assertEqual(even_or_odd(58),"This is even.","even_or_odd(58)")
        self.assertEqual(even_or_odd(1002),"This is even.","even_or_odd(1002)")
        self.assertEqual(even_or_odd(199994),"This is even.","even_or_odd(199994)")
        self.assertEqual(even_or_odd(0),"This is even.","even_or_odd(0)")
        self.assertEqual(even_or_odd(6),"This is even.","even_or_odd(6)")
        self.assertEqual(even_or_odd(79),"This is odd.","even_or_odd(79)")
        self.assertEqual(even_or_odd(-65),"This is odd.","even_or_odd(-65)")
        self.assertEqual(even_or_odd(-43.0),"This is odd.","even_or_odd(-43.0)")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="funct_conditional_mixedupcode_3" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">which_quadrant</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">x</span></code> and <code class="docutils literal notranslate"><span class="pre">y</span></code> and returns its corresponding quadrants in increasing order
(“Quadrant 1”, “Quadrant 2”, etc.). Refer to <a class="reference external" href="https://en.wikipedia.org/wiki/Quadrant_(plane_geometry">https://en.wikipedia.org/wiki/Quadrant_(plane_geometry</a>) for quadrants on a coordinate
plane if needed. If coordinate is on the x or y axis, return <code class="docutils literal notranslate"><span class="pre">&quot;It's</span> <span class="pre">not</span> <span class="pre">in</span> <span class="pre">a</span> <span class="pre">quadrant.&quot;</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">which_quadrant(1,</span> <span class="pre">100)</span></code>
should return <code class="docutils literal notranslate"><span class="pre">&quot;Quadrant</span> <span class="pre">1&quot;</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="5"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def which_quadrant(x, y):
---
    if (x > 0) and (y > 0):
---
    if (x >= 0) and (y >= 0): #paired
---
        return "Quadrant 1"
---
    elif (x < 0) and (y > 0):
---
    elif (x <= 0) and (y >= 0): #paired
---
        return "Quadrant 2"
---
    elif (x < 0) and (y < 0):
---
    elif (x <= 0) and (y <= 0): #paired
---
        return "Quadrant 3"
---
    elif (x > 0) and (y < 0):
---
    elif (x >= 0) and (y <= 0): #paired
---
        return "Quadrant 4"
---
    else:
---
        return "It's not in a quadrant."
---
        return 'It's not in a quadrant.' #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=funct_conditional_mixedupcode_3_ac data-question_label="6">
<div id=funct_conditional_mixedupcode_3_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">which_quadrant</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">x</span></code> and <code class="docutils literal notranslate"><span class="pre">y</span></code> and returns its corresponding quadrants
(“Quadrant 1”, “Quadrant 2”, etc.). Refer to <a class="reference external" href="https://en.wikipedia.org/wiki/Quadrant_(plane_geometry">https://en.wikipedia.org/wiki/Quadrant_(plane_geometry</a>) for quadrants on a coordinate
plane if needed. If coordinate is on the x or y axis, return <code class="docutils literal notranslate"><span class="pre">&quot;It's</span> <span class="pre">not</span> <span class="pre">in</span> <span class="pre">a</span> <span class="pre">quadrant.&quot;</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">which_quadrant(1,</span> <span class="pre">100)</span></code>
should return <code class="docutils literal notranslate"><span class="pre">&quot;Quadrant</span> <span class="pre">1&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_conditional_mixedupcode_3_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def which_quadrant(x, y):
    # write code here

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(which_quadrant(0, 0),"It's not in a quadrant.","which_quadrant(0, 0)")
        self.assertEqual(which_quadrant(0, 15),"It's not in a quadrant.","which_quadrant(0, 15)")
        self.assertEqual(which_quadrant(-200, 0),"It's not in a quadrant.","which_quadrant(-200, 0)")
        self.assertEqual(which_quadrant(11, 0),"It's not in a quadrant.","which_quadrant(11, 0)")
        self.assertEqual(which_quadrant(0, 10),"It's not in a quadrant.","which_quadrant(0, 10)")
        self.assertEqual(which_quadrant(1, 100),"Quadrant 1","which_quadrant(1, 100)")
        self.assertEqual(which_quadrant(-5, 20),"Quadrant 2","which_quadrant(-5, 20)")
        self.assertEqual(which_quadrant(-100, -85),"Quadrant 3","which_quadrant(-100, -85)")
        self.assertEqual(which_quadrant(2020.4, -4.5),"Quadrant 4","which_quadrant(2020.4, -4.5)")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="funct_conditional_mixedupcode_4" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">axis_or_origin</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">x</span></code> and <code class="docutils literal notranslate"><span class="pre">y</span></code>. If both coordinates aren’t 0,
return <code class="docutils literal notranslate"><span class="pre">&quot;This</span> <span class="pre">coordinate</span> <span class="pre">is</span> <span class="pre">not</span> <span class="pre">on</span> <span class="pre">an</span> <span class="pre">axis.</span> <span class="pre">It's</span> <span class="pre">in</span> <span class="pre">a</span> <span class="pre">quadrant.&quot;</span></code>. If both coordinates are 0, return <code class="docutils literal notranslate"><span class="pre">&quot;This</span> <span class="pre">coordinate</span> <span class="pre">is</span> <span class="pre">the</span> <span class="pre">origin.&quot;</span></code>.
Lastly, if the coordinate isn’t in a quadrant and isn’t the origin, determine if the coordinate is on the x or y axis by returning
<code class="docutils literal notranslate"><span class="pre">&quot;This</span> <span class="pre">coordinate</span> <span class="pre">is</span> <span class="pre">on</span> <span class="pre">the</span> <span class="pre">(y-axis</span> <span class="pre">or</span> <span class="pre">x-axis).&quot;</span></code>. Refer to <a class="reference external" href="https://www.math.net/x-and-y-axis">https://www.math.net/x-and-y-axis</a> for information on axes if needed.
For example, <code class="docutils literal notranslate"><span class="pre">axis_or_origin(500,</span> <span class="pre">0)</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;This</span> <span class="pre">coordinate</span> <span class="pre">is</span> <span class="pre">on</span> <span class="pre">the</span> <span class="pre">x-axis.&quot;</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="7"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def axis_or_origin(x, y):
---
    if (x != 0) and (y != 0):
---
    if (x != 0) && (y != 0): #paired
---
    if (x != 0) or (y != 0): #paired
---
        return "This coordinate is not on an axis. It's in a quadrant."
---
    elif (x == 0) and (y == 0):
---
        return "This coordinate is the origin."
---
    elif x == 0:
---
    if x == 0: #paired
---
    elif x = 0: #paired
---
        return "This coordinate is on the y-axis."
---
    else:
---
        return "This coordinate is on the x-axis."
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=funct_conditional_mixedupcode_4_ac data-question_label="8">
<div id=funct_conditional_mixedupcode_4_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">axis_or_origin</span></code> that takes in parameters <code class="docutils literal notranslate"><span class="pre">x</span></code> and <code class="docutils literal notranslate"><span class="pre">y</span></code>. If both coordinates aren’t 0,
return <code class="docutils literal notranslate"><span class="pre">&quot;This</span> <span class="pre">coordinate</span> <span class="pre">is</span> <span class="pre">not</span> <span class="pre">on</span> <span class="pre">an</span> <span class="pre">axis.</span> <span class="pre">It's</span> <span class="pre">in</span> <span class="pre">a</span> <span class="pre">quadrant.&quot;</span></code>. If both coordinates are 0, return <code class="docutils literal notranslate"><span class="pre">&quot;This</span> <span class="pre">coordinate</span> <span class="pre">is</span> <span class="pre">the</span> <span class="pre">origin.&quot;</span></code>.
Lastly, if the coordinate isn’t in a quadrant and isn’t the origin, determine if the coordinate is on the x or y axis by returning
<code class="docutils literal notranslate"><span class="pre">&quot;This</span> <span class="pre">coordinate</span> <span class="pre">is</span> <span class="pre">on</span> <span class="pre">the</span> <span class="pre">(y-axis</span> <span class="pre">or</span> <span class="pre">x-axis).&quot;</span></code>. Refer to <a class="reference external" href="https://www.math.net/x-and-y-axis">https://www.math.net/x-and-y-axis</a> for information on axes if needed.
For example, <code class="docutils literal notranslate"><span class="pre">axis_or_origin(500,</span> <span class="pre">0)</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;This</span> <span class="pre">coordinate</span> <span class="pre">is</span> <span class="pre">on</span> <span class="pre">the</span> <span class="pre">x-axis.&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_conditional_mixedupcode_4_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def axis_or_origin(x, y):
    # write code here

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(axis_or_origin(-5, 20),"This coordinate is not on an axis. It's in a quadrant.","axis_or_origin(-5, 20)")
        self.assertEqual(axis_or_origin(5, 20.5),"This coordinate is not on an axis. It's in a quadrant.","axis_or_origin(5, 20.5)")
        self.assertEqual(axis_or_origin(-400, -101),"This coordinate is not on an axis. It's in a quadrant.","axis_or_origin(-400, -101)")
        self.assertEqual(axis_or_origin(10004, -40.3),"This coordinate is not on an axis. It's in a quadrant.","axis_or_origin(10004, -40.3)")
        self.assertEqual(axis_or_origin(0, 0),"This coordinate is the origin.","axis_or_origin(0, 0)")
        self.assertEqual(axis_or_origin(0, -200),"This coordinate is on the y-axis.","axis_or_origin(0, -200)")
        self.assertEqual(axis_or_origin(500, 0),"This coordinate is on the x-axis.","axis_or_origin(500, 0)")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="funct_conditional_mixedupcode_5" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">num_to_day</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">num</span></code>. If <code class="docutils literal notranslate"><span class="pre">num</span></code> is between 1-5 inclusive,
create a nested conditional statement that returns <code class="docutils literal notranslate"><span class="pre">&quot;This</span> <span class="pre">is</span> <span class="pre">a</span> <span class="pre">weekday,</span> <span class="pre">and</span> <span class="pre">it's</span> <span class="pre">(insert</span> <span class="pre">day</span> <span class="pre">here).&quot;</span></code>, 1 being Monday through 5 being Friday.
If <code class="docutils literal notranslate"><span class="pre">num</span></code> is 6 or 7, create a nested conditional statement that returns <code class="docutils literal notranslate"><span class="pre">&quot;It</span> <span class="pre">is</span> <span class="pre">the</span> <span class="pre">weekend,</span> <span class="pre">and</span> <span class="pre">it's</span> <span class="pre">(insert</span> <span class="pre">day</span> <span class="pre">here).&quot;</span></code>,
6 being Saturday and 7 being Sunday. If none of the other conditions have been met, return <code class="docutils literal notranslate"><span class="pre">&quot;The</span> <span class="pre">number</span> <span class="pre">does</span> <span class="pre">not</span> <span class="pre">correspond</span> <span class="pre">to</span> <span class="pre">a</span> <span class="pre">day</span> <span class="pre">of</span> <span class="pre">the</span> <span class="pre">week.&quot;</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">num_to_day(1)</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;This</span> <span class="pre">is</span> <span class="pre">a</span> <span class="pre">weekday,</span> <span class="pre">and</span> <span class="pre">it's</span> <span class="pre">Monday.&quot;</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="9"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def num_to_day(num):
---
    if (num >= 1) and (num <= 5):
---
    if (num > 1) and (num < 5): #paired
---
        if num == 1:
---
            return "This is a weekday, and it's Monday."
---
        elif num == 2:
---
            return "This is a weekday, and it's Tuesday."
---
        elif num == 3:
---
            return "This is a weekday, and it's Wednesday."
---
        elif num == 4:
---
            return "This is a weekday, and it's Thursday."
---
        else:
---
            return "This is a weekday, and it's Friday."
---
    elif (num > 5) and (num <= 7):
---
    elif num <= 7: #paired
---
    elif (num > 5) or (num <= 7): #paired
---
    elif (num >= 5) or (num <= 7): #paired
---
        if num == 6:
---
            return "It is the weekend, and it's Saturday."
---
        else:
---
            return "It is the weekend, and it's Sunday."
---
    else:
---
        return "The number does not correspond to a day of the week."
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=funct_conditional_mixedupcode_5_ac data-question_label="10">
<div id=funct_conditional_mixedupcode_5_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">num_to_day</span></code> that takes in a parameter <code class="docutils literal notranslate"><span class="pre">num</span></code>. If <code class="docutils literal notranslate"><span class="pre">num</span></code> is between 1-5 inclusive,
create a nested conditional statement that returns <code class="docutils literal notranslate"><span class="pre">&quot;This</span> <span class="pre">is</span> <span class="pre">a</span> <span class="pre">weekday,</span> <span class="pre">and</span> <span class="pre">it's</span> <span class="pre">(insert</span> <span class="pre">day</span> <span class="pre">here).&quot;</span></code>, 1 being Monday through 5 being Friday.
If <code class="docutils literal notranslate"><span class="pre">num</span></code> is 6 or 7, create a nested conditional statement that returns <code class="docutils literal notranslate"><span class="pre">&quot;It</span> <span class="pre">is</span> <span class="pre">the</span> <span class="pre">weekend,</span> <span class="pre">and</span> <span class="pre">it's</span> <span class="pre">(insert</span> <span class="pre">day</span> <span class="pre">here).&quot;</span></code>,
6 being Saturday and 7 being Sunday. If none of the other conditions have been met, return <code class="docutils literal notranslate"><span class="pre">&quot;The</span> <span class="pre">number</span> <span class="pre">does</span> <span class="pre">not</span> <span class="pre">correspond</span> <span class="pre">to</span> <span class="pre">a</span> <span class="pre">day</span> <span class="pre">of</span> <span class="pre">the</span> <span class="pre">week.&quot;</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">num_to_day(1)</span></code> should return <code class="docutils literal notranslate"><span class="pre">&quot;This</span> <span class="pre">is</span> <span class="pre">a</span> <span class="pre">weekday,</span> <span class="pre">and</span> <span class="pre">it's</span> <span class="pre">Monday.&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="funct_conditional_mixedupcode_5_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def num_to_day(num):
    # write code here

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(num_to_day(-1),"The number does not correspond to a day of the week.","num_to_day(-1)")
        self.assertEqual(num_to_day(0),"The number does not correspond to a day of the week.","num_to_day(0)")
        self.assertEqual(num_to_day(1),"This is a weekday, and it's Monday.","num_to_day(1)")
        self.assertEqual(num_to_day(2),"This is a weekday, and it's Tuesday.","num_to_day(2)")
        self.assertEqual(num_to_day(3),"This is a weekday, and it's Wednesday.","num_to_day(3)")
        self.assertEqual(num_to_day(4),"This is a weekday, and it's Thursday.","num_to_day(4)")
        self.assertEqual(num_to_day(5),"This is a weekday, and it's Friday.","num_to_day(5)")
        self.assertEqual(num_to_day(6),"It is the weekend, and it's Saturday.","num_to_day(6)")
        self.assertEqual(num_to_day(7),"It is the weekend, and it's Sunday.","num_to_day(7)")
        self.assertEqual(num_to_day(8),"The number does not correspond to a day of the week.","num_to_day(8)")
        self.assertEqual(num_to_day(-100),"The number does not correspond to a day of the week.","num_to_day(-100)")
        self.assertEqual(num_to_day(1001),"The number does not correspond to a day of the week.","num_to_day(1001)")

myTests().main()
</textarea>
</div>
</div>
</section>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>



  
    {% if settings.google_ga %}
      <script>
        var _gaq = _gaq || [];
        _gaq.push(['_setAccount', '{{ settings.google_ga }}']);
        _gaq.push(['_trackPageview']);

        (function() {
          var ga = document.createElement('script'); ga.async = true;
          ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
          var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
        })();
      </script>
    {% endif %}
  

  
    <script type="application/json" class="js-hypothesis-config">
    {
      "showHighlights": false
    }
    </script>

    <script async src="https://hypothes.is/embed.js"></script>
    <!--provide space to prevent hypothes.is sidebar from obscuring content-->
    <style>@media screen and (max-width:767px) { body { padding-right: 21px; }}</style>

  


<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>