<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.17.1: http://docutils.sourceforge.net/" />

    <title>Functions Mixed-Up Code Questions &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.cd688d1623c9932e.css?v=630908D8" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.bf73008a1eba9f4e.css?v=7DED055E" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/runtime.439d581f758f2b99.bundle.js?v=2E54442"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.51d9dde4e5162190.bundle.js?v=9CAA2D33"></script>
    <script src="../_static/runestone.42b576791a4ed54b.bundle.js?v=B7293199"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    
    eBookConfig.useRunestoneServices = {{use_services}};
    eBookConfig.host = '';
    eBookConfig.app = eBookConfig.host + '/runestone';
    eBookConfig.course = '{{ course_name }}';
    eBookConfig.basecourse = '{{ base_course }}';
    eBookConfig.isLoggedIn = {{ is_logged_in}};
    eBookConfig.email = '{{ user_email }}';
    eBookConfig.isInstructor = {{ is_instructor }};
    eBookConfig.username = '{{ user_id}}';
    eBookConfig.readings = {{ readings|safe}};
    eBookConfig.activities = {{ activity_info|safe }}
    eBookConfig.downloadsEnabled = {{downloads_enabled}};
    eBookConfig.allow_pairs = {{allow_pairs}}
    eBookConfig.enableCompareMe = {{enable_compare_me}};
    // _`new_server_prefix`: Defined by the server to be the path to the new server.
    eBookConfig.new_server_prefix = '{{ new_server_prefix }}';
    
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->

    
        {% if serve_ad is defined and serve_ad and settings.adsenseid %}
         <script data-ad-client="{{settings.adsenseid}}" async src="https://pagead2.googlesyndication.com/pagead/js/adsbygoogle.js"></script>
        {% endif %}
    



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
            
            {{ course_name }}
            
          
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="browsing_warning"></span></li>

        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li id="inst_peer_link"><a href='/runestone/peer/instructor.html'>Peer Instruction (Instructor)</a></li>            
            <li><a href='/runestone/peer/student.html'>Peer Instruction (Student)</a></li>            
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/ns/books/published/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=functions_mixedupcode_orig'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
            <script></script>
<li class="dropdown globaltoc-container">
    <a
        href="../index.html"
        class="dropdown-toggle"
        data-toggle="dropdown"
        aria-label="index"
        >This Chapter<b class="caret"></b
    ></a>
    <ul class="dropdown-menu globaltoc">
         
        {% for row in subchapter_list %}
        <li class="toctree-l1">
            <a class="reference internal" href="{{ row['subchap_uri']}}"
                >{{ row['title']}}</a
            >
        </li>
        {% endfor %} 
        
    </ul>
</li>
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->

    
        {% if settings.num_banners > 0 and settings.show_rs_banner %}
         <div style="width: 100%;">
            <a href="/runestone/default/donate?banner={{banner_num}}">
            <img class="runestone-banner"
                 src="/runestone/static/images/RunestoneBanner{{banner_num}}.png"
                 alt="Please Support Runestone"
                 style="display: block; width: 728px; margin-left: auto; margin-right: auto; margin-top: 10px;">
            </a>
         </div>
        {% endif %}
    


  
  <section id="functions-mixed-up-code-questions">
<h1>Functions Mixed-Up Code Questions<a class="headerlink" href="#functions-mixed-up-code-questions" title="Permalink to this headline">¶</a></h1>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Functions_MixedUpCode_1P_v2" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">exp</span></code> that takes in two parameters, <code class="docutils literal notranslate"><span class="pre">num1</span></code> and <code class="docutils literal notranslate"><span class="pre">num2</span></code>, and returns <code class="docutils literal notranslate"><span class="pre">num1</span></code> raised to the
power of <code class="docutils literal notranslate"><span class="pre">num2</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">exp(2,3)</span></code> should return <code class="docutils literal notranslate"><span class="pre">8</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="1"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def exp(num1, num2):
---
def exp(num1, num2) #paired
---
    value = num1 ** num2
---
    value = num1 * num2 #paired
---
    return value
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Functions_MixedUpCode_1P_v2_ac data-question_label="2">
<div id=Functions_MixedUpCode_1P_v2_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">exp</span></code> that takes in two parameters, <code class="docutils literal notranslate"><span class="pre">num1</span></code> and <code class="docutils literal notranslate"><span class="pre">num2</span></code>, and returns <code class="docutils literal notranslate"><span class="pre">num1</span></code> raised to the
power of <code class="docutils literal notranslate"><span class="pre">num2</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">exp(2,3)</span></code> should return <code class="docutils literal notranslate"><span class="pre">8</span></code>.</p>

</div>
<textarea data-lang="python" id="Functions_MixedUpCode_1P_v2_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertAlmostEqual(exp(2,3),8,0, "exp(2,3)")
        self.assertAlmostEqual(exp(1,1),1,0,"exp(1,1)")
        self.assertAlmostEqual(exp(2,-2),0.25,2,"exp(2,-2)")
        self.assertAlmostEqual(exp(0,0),1,0, "exp(0,0)")
        self.assertAlmostEqual(exp(1.2,4),2.0736,2,"exp(1.2,4)")
        self.assertAlmostEqual(exp(2,4.1),17.148375400580687,2,"exp(2,4.1)")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Functions_MixedUpCode_2" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">poly_perimeter</span></code> that takes in two parameters, <code class="docutils literal notranslate"><span class="pre">len_side</span></code> and <code class="docutils literal notranslate"><span class="pre">num_sides</span></code>, and
returns the perimeter of the polygon.  The perimeter of a polygon is the length of each side times the number of sides.  For example, <code class="docutils literal notranslate"><span class="pre">poly_perimeter(4,5)</span></code>
should return <code class="docutils literal notranslate"><span class="pre">20</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="3"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def poly_perimeter(len_side, num_sides):
---
Def poly_perimeter(len_side, num_sides): #paired
---
    per = len_side * num_sides
---
    per = len_side ** num_sides #paired
---
    return per
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Functions_MixedUpCode_2_ac data-question_label="4">
<div id=Functions_MixedUpCode_2_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">poly_perimeter</span></code> that takes in two parameters, <code class="docutils literal notranslate"><span class="pre">len_side</span></code> and <code class="docutils literal notranslate"><span class="pre">num_sides</span></code>, and
returns the perimeter of the polygon.  The perimeter of a polygon is the length of each side times the number of sides.  For example, <code class="docutils literal notranslate"><span class="pre">poly_perimeter(4,5)</span></code>
should return <code class="docutils literal notranslate"><span class="pre">20</span></code>.</p>

</div>
<textarea data-lang="python" id="Functions_MixedUpCode_2_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertAlmostEqual(poly_perimeter(4,5),20,0,"poly_perimeter(4,5)")
        self.assertAlmostEqual(poly_perimeter(2,3),6,0,"poly_perimeter(2,3)")
        self.assertAlmostEqual(poly_perimeter(0.1,6),0.6000000000000001,2,"poly_perimeter(0.1,6)")
        self.assertAlmostEqual(poly_perimeter(2,9999),19998,0,"poly_perimeter(2,9999)")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Functions_MixedUpCode_3" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">sphere_volume</span></code> that takes in one parameter, <code class="docutils literal notranslate"><span class="pre">radius</span></code>, and
returns the volume of a sphere with the given <code class="docutils literal notranslate"><span class="pre">radius</span></code>.  The volume of a sphere is 4 / 3 times PI (in which 3.14 should be used) times the <code class="docutils literal notranslate"><span class="pre">radius</span></code> cubed.
For example, <code class="docutils literal notranslate"><span class="pre">sphere_volume(3)</span></code> should return <code class="docutils literal notranslate"><span class="pre">113.04</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="5"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def sphere_volume(radius):
---
    volume = (4/3) * (3.14) * radius ** 3
---
    volume = (4/3) * (3.14) * (radius * 3) #paired
---
    return volume
---
    ret volume #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Functions_MixedUpCode_3_ac data-question_label="6">
<div id=Functions_MixedUpCode_3_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">sphere_volume</span></code> that takes in one parameter, <code class="docutils literal notranslate"><span class="pre">radius</span></code>, and
returns the volume of a sphere with the given <code class="docutils literal notranslate"><span class="pre">radius</span></code>.  The volume of a sphere is 4 / 3 times PI (in which 3.14 should be used) times the <code class="docutils literal notranslate"><span class="pre">radius</span></code> cubed.
For example, <code class="docutils literal notranslate"><span class="pre">sphere_volume(3)</span></code> should return <code class="docutils literal notranslate"><span class="pre">113.04</span></code>.</p>

</div>
<textarea data-lang="python" id="Functions_MixedUpCode_3_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertAlmostEqual(sphere_volume(3),113.04,2,"sphere_volume(3)")
        self.assertAlmostEqual(sphere_volume(0.1),0.004186666666666668,5,"sphere_volume(0.1)")
        self.assertAlmostEqual(sphere_volume(5),523.3333333333334,2,"sphere_volume(5)")
        self.assertAlmostEqual(sphere_volume(10),4186.666666666667,2,"sphere_volume(10)")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Functions_MixedUpCode_4" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">circle_circumference</span></code> that takes in one parameter, <code class="docutils literal notranslate"><span class="pre">radius</span></code>, and returns the
circumference of a circle with the given <code class="docutils literal notranslate"><span class="pre">radius</span></code>.  The circumference of a circle is 2 times PI (in which 3.14 should be used) times the <code class="docutils literal notranslate"><span class="pre">radius</span></code>. For example,
<code class="docutils literal notranslate"><span class="pre">circle_circumference(3)</span></code> should return <code class="docutils literal notranslate"><span class="pre">18.84</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="7"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def circle_circumference(radius):
---
def circle_circumference(radius: #paired
---
    value = 2 * 3.14 * radius
---
    value = 2 ** 3.14 * radius #paired
---
    return value
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Functions_MixedUpCode_4_ac data-question_label="8">
<div id=Functions_MixedUpCode_4_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">circle_circumference</span></code> that takes in one parameter, <code class="docutils literal notranslate"><span class="pre">radius</span></code>, and returns the
circumference of a circle with the given <code class="docutils literal notranslate"><span class="pre">radius</span></code>.  The circumference of a circle is 2 times PI (in which 3.14 should be used) times the <code class="docutils literal notranslate"><span class="pre">radius</span></code>. For example,
<code class="docutils literal notranslate"><span class="pre">circle_circumference(3)</span></code> should return <code class="docutils literal notranslate"><span class="pre">18.84</span></code>.</p>

</div>
<textarea data-lang="python" id="Functions_MixedUpCode_4_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertAlmostEqual(circle_circumference(3),18.84,2,"circle_circumference(3)")
        self.assertAlmostEqual(circle_circumference(0.1),0.6280000000000001,3,"circle_circumference(0.1)")
        self.assertAlmostEqual(circle_circumference(5),31.400000000000002,2,"circle_circumference(5)")
        self.assertAlmostEqual(circle_circumference(10),62.800000000000004,2,"circle_circumference(10)")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Functions_MixedUpCode_5" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">slope</span></code> that takes in four parameters, <code class="docutils literal notranslate"><span class="pre">x1</span></code>, <code class="docutils literal notranslate"><span class="pre">x2</span></code>, <code class="docutils literal notranslate"><span class="pre">y1</span></code>, and <code class="docutils literal notranslate"><span class="pre">y2</span></code>, and
returns the slope of a line formed by the two points.  The slope is <code class="docutils literal notranslate"><span class="pre">y2</span></code> minus <code class="docutils literal notranslate"><span class="pre">y1</span></code> divided by <code class="docutils literal notranslate"><span class="pre">x2</span></code> minus <code class="docutils literal notranslate"><span class="pre">x1</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">slope(4,</span> <span class="pre">3,</span> <span class="pre">7,</span> <span class="pre">8)</span></code> should return <code class="docutils literal notranslate"><span class="pre">-1</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="9"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def slope(x1, x2, y1, y2):
---
def slope(x1, x2, y1, y2) #paired
---
    value = (y2 - y1) / (x2 - x1)
---
    value = (x2 - x1) / (y2 - y1) #paired
---
    return value
---
    return Value #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Functions_MixedUpCode_5_ac data-question_label="10">
<div id=Functions_MixedUpCode_5_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">slope</span></code> that takes in four parameters, <code class="docutils literal notranslate"><span class="pre">x1</span></code>, <code class="docutils literal notranslate"><span class="pre">x2</span></code>, <code class="docutils literal notranslate"><span class="pre">y1</span></code>, and <code class="docutils literal notranslate"><span class="pre">y2</span></code>, and
returns the slope of a line formed by the two points.  The slope is <code class="docutils literal notranslate"><span class="pre">y2</span></code> minus <code class="docutils literal notranslate"><span class="pre">y1</span></code> divided by <code class="docutils literal notranslate"><span class="pre">x2</span></code> minus <code class="docutils literal notranslate"><span class="pre">x1</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">slope(4,</span> <span class="pre">3,</span> <span class="pre">7,</span> <span class="pre">8)</span></code> should return <code class="docutils literal notranslate"><span class="pre">-1</span></code>.</p>

</div>
<textarea data-lang="python" id="Functions_MixedUpCode_5_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertAlmostEqual(slope(4,3,7,8), -1.0, 1, "slope(4,3,7,8)")
        self.assertAlmostEqual(slope(1,2,0,0), 0.0, 1, "slope(1,2,0,0)")
        self.assertAlmostEqual(slope(0,1,1,1), 0.0, 1,"slope(0,1,1,1)")
        self.assertAlmostEqual(slope(-1,2,3,4.4), 0.4666666666666668, 2, "slope(-1,2,3,4.4)")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Functions_MixedUpCode_6" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">get_hypotenuse</span></code> that takes in two parameters, <code class="docutils literal notranslate"><span class="pre">angle</span></code> and <code class="docutils literal notranslate"><span class="pre">opposite_side</span></code>, and
returns the hypotenuse of a right triangle. <code class="docutils literal notranslate"><span class="pre">angle</span></code> contains the angle facing the opposite side in degrees.
The sine function takes in a value in radians. The sine of an angle is equal to the opposite side over the hypotenuse.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="11"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        import math
---
def get_hypotenuse(angle, opposite_side):
---
def get_hypotenuse(angle, opposite_side) #paired
---
    sine = math.sin(math.radians(angle))
---
    sine = math.sin(angle) #paired
---
    hypotenuse = opposite_side / sine
---
    hypotenuse = sine / opposite_side #paired
---
    return hypotenuse
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Functions_MixedUpCode_6_ac data-question_label="12">
<div id=Functions_MixedUpCode_6_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">get_hypotenuse</span></code> that takes in two parameters, <code class="docutils literal notranslate"><span class="pre">angle</span></code> and <code class="docutils literal notranslate"><span class="pre">opposite_side</span></code>, and
returns the hypotenuse of a right triangle. <code class="docutils literal notranslate"><span class="pre">angle</span></code> contains the angle facing the opposite side in degrees.
The sine function takes in a value in radians. The sine of an angle is equal to the opposite side over the hypotenuse.</p>

</div>
<textarea data-lang="python" id="Functions_MixedUpCode_6_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
import math


====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertAlmostEqual(get_hypotenuse(45, 1), 1.41421356237, 5, "get_hypotenuse(45, 1)")
        self.assertAlmostEqual(get_hypotenuse(16, 7), 25.3991291727, 1, "get_hypotenuse(16, 7)")
        self.assertAlmostEqual(get_hypotenuse(34, 14), 25.035765379, 2,"get_hypotenuse(34, 14)")
        self.assertAlmostEqual(get_hypotenuse(32, 15.9), 30, 2,"get_hypotenuse(32, 15.9)")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Functions_MixedUpCode_7" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">convert_to_miles</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">feet</span></code> as a parameter and
returns the value in miles. <code class="docutils literal notranslate"><span class="pre">feet</span></code> contains the length of an object in feet. 1 mile is equal to 5280 feet.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="13"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def convert_to_miles(feet):
---
def convert_to_miles(feet) #paired
---
    miles = feet / 5280
---
    miles = 5280 / feet #paired
---
    return miles
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Functions_MixedUpCode_7_ac data-question_label="14">
<div id=Functions_MixedUpCode_7_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">convert_to_miles</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">feet</span></code> as a parameter and
returns the value in miles. <code class="docutils literal notranslate"><span class="pre">feet</span></code> contains the length of an object in feet. 1 mile is equal to 5280 feet.</p>

</div>
<textarea data-lang="python" id="Functions_MixedUpCode_7_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">


====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertAlmostEqual(convert_to_miles(32451), 6.1460227, 5, "convert_to_miles(32451)")
        self.assertAlmostEqual(convert_to_miles(4), 0.000757576, 9, "convert_to_miles(4)")
        self.assertAlmostEqual(convert_to_miles(5280), 1, 1,"convert_to_miles(5280)")
        self.assertAlmostEqual(convert_to_miles(12475), 2.3626894, 4,"convert_to_miles(12475)")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Functions_MixedUpCode_8" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">even_odd</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">integer</span></code> as a parameter and
returns 1 if the integer is even, 2 if the integer is odd, or 3 if neither is true.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="15"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def even_odd(integer):
---
def even_odd(integer) #paired
---
    if integer == 0:
        return 3
---
    if integer = 0: #paired
        return 3
---
    elif integer % 2 == 0:
        return 1
---
    else:
        return 2
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Functions_MixedUpCode_8_ac data-question_label="16">
<div id=Functions_MixedUpCode_8_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">even_odd</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">integer</span></code> as a parameter and
returns 1 if the integer is even, 2 if the integer is odd, or 3 if neither is true.</p>

</div>
<textarea data-lang="python" id="Functions_MixedUpCode_8_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">


====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(even_odd(0), 3, "even_odd(0)")
        self.assertEqual(even_odd(32), 1, "even_odd(32)")
        self.assertEqual(even_odd(5), 2, "even_odd(5)")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Functions_MixedUpCode_9" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">larger_string</span></code> that takes in two parameters, <code class="docutils literal notranslate"><span class="pre">str1</span></code> and <code class="docutils literal notranslate"><span class="pre">str2</span></code>, and
returns the string that has the larger length. If the lengths are equal, return <code class="docutils literal notranslate"><span class="pre">str1</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="17"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def larger_string(str1, str2):
---
def larger_string(str1, str2) #paired
---
    if len(str1) >= len(str2):
        return str1
---
    if len(str1) > len(str2): #paired
        return str1
---
    else:
        return str2
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Functions_MixedUpCode_9_ac data-question_label="18">
<div id=Functions_MixedUpCode_9_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">larger_string</span></code> that takes in two parameters, <code class="docutils literal notranslate"><span class="pre">str1</span></code> and <code class="docutils literal notranslate"><span class="pre">str2</span></code>, and
returns the string that has the larger length. If the lengths are equal, return <code class="docutils literal notranslate"><span class="pre">str1</span></code>.</p>

</div>
<textarea data-lang="python" id="Functions_MixedUpCode_9_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">


====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(larger_string("Three", "Four"), "Three", "larger_string('Three', 'Four')")
        self.assertEqual(larger_string("That", "Knew"), "That", "larger_string('That', 'Knew')")
        self.assertEqual(larger_string("Cat", "People"), "People", "larger_string('Cat', 'People')")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="Functions_MixedUpCode_10" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">find_distance</span></code> that takes in four parameters, <code class="docutils literal notranslate"><span class="pre">x1</span></code>, <code class="docutils literal notranslate"><span class="pre">y1</span></code>, <code class="docutils literal notranslate"><span class="pre">x2</span></code>, and <code class="docutils literal notranslate"><span class="pre">y2</span></code>, and
returns the distance between these two points.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="19"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        import math
---
def find_distance(x1, y1, x2, y2):
---
def find_distance(x1, y1, x2, y2) #paired
---
    sqr_dist = (x1 - x2)**2 + (y1 - y2)**2
---
    sqr_dist = (x1 - y1)**2 + (x2 - y2)**2 #paired
---
    return math.sqrt(sqr_dist)
---
    return sqrt(sqr_distance) #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=Functions_MixedUpCode_10_ac data-question_label="20">
<div id=Functions_MixedUpCode_10_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">find_distance</span></code> that takes in four parameters, <code class="docutils literal notranslate"><span class="pre">x1</span></code>, <code class="docutils literal notranslate"><span class="pre">y1</span></code>, <code class="docutils literal notranslate"><span class="pre">x2</span></code>, and <code class="docutils literal notranslate"><span class="pre">y2</span></code>, and
returns the distance between these two points.</p>

</div>
<textarea data-lang="python" id="Functions_MixedUpCode_10_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
import math

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertAlmostEqual(find_distance(1, 2, 21, 14), 23.3238, 4, "find_distance(1, 2, 21, 14)")
        self.assertAlmostEqual(find_distance(5, 4, 15, 345), 341.1466, 4, "find_distance(5, 4, 15, 345)")
        self.assertAlmostEqual(find_distance(51, 42, 75, 12), 38.41875, 3, "find_distance(51, 42, 75, 12)")
        self.assertAlmostEqual(find_distance(1, 2, 3, 4), 2.828427, 3, "find_distance(1, 2, 3, 4)")


myTests().main()
</textarea>
</div>
</div>
</section>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>



  
    {% if settings.google_ga %}
      <script>
        var _gaq = _gaq || [];
        _gaq.push(['_setAccount', '{{ settings.google_ga }}']);
        _gaq.push(['_trackPageview']);

        (function() {
          var ga = document.createElement('script'); ga.async = true;
          ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
          var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
        })();
      </script>
    {% endif %}
  

  
    <script type="application/json" class="js-hypothesis-config">
    {
      "showHighlights": false
    }
    </script>

    <script async src="https://hypothes.is/embed.js"></script>
    <!--provide space to prevent hypothes.is sidebar from obscuring content-->
    <style>@media screen and (max-width:767px) { body { padding-right: 21px; }}</style>

  


<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>