<?xml version="1.0"?>
<section xml:id="lists_mixed-up-code-questions">
  <title>Mixed-up Code Questions</title>
  <exercise label="list_MUC_password_creator" numbered="yes" adaptive="yes" indentation="show" language="python">
    <statement>
      <p>Create a function <c>password_maker(word_list)</c> that takes a list of words, <c>word_list</c> and returns a string with the first and last character of
            each word in a list. For example, <c>password_maker("show", "me", "the", "money")</c> would return <c>swmetemy</c>.</p>
    </statement>
    <blocks>
      <block order="2">
        <cline>def password_maker(word_list):</cline>
      </block>
      <block order="7">
        <cline>    password = ""</cline>
      </block>
      <block order="5">
        <cline>    for word in word_list:</cline>
      </block>
      <block order="6">
        <cline>        first_char = word[0]</cline>
      </block>
      <block order="8">
        <cline>        last_char = word[-1]</cline>
      </block>
      <block order="9">
        <cline>        password += first_char</cline>
      </block>
      <block order="1">
        <cline>        password += last_char</cline>
      </block>
      <block order="3">
        <choice correct="yes">
          <cline>    return password</cline>
        </choice>
        <choice>
          <cline>    return word_list </cline>
        </choice>
      </block>
    </blocks>
  </exercise>
  <exercise label="list_MUC_password_creator_ac">
    <statement>
      <p>Write a function <c>password_maker(word_list)</c> that takes a a list of words, <c>word_list</c> and returns a string with the first and last character of each
            word in a list. For example, <c>password_maker(["show", "me", "the", "money"])</c> would return <c>swmetemy</c>.</p>
    </statement>
    <program xml:id="list_MUC_password_creator_ac_editor" interactive="activecode" language="python">
      <input>

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):

        self.assertEqual(password_maker(["show", "me", "the", "money"]), "swmetemy", 'password_maker("show", "me", "the", "money")')
        self.assertEqual(password_maker(["Bye", "Now"]), "BeNw", 'password_maker(["Bye", "Now"])')
        self.assertEqual(password_maker(["Beautiful"]), "Bl", 'password_maker(["Beautiful"])')
        self.assertEqual(password_maker(["You", "complete", "me"]), "Yuceme", 'password_maker(["You", "complete", "me"])')

myTests().main()
        </input>
    </program>
  </exercise>
  <exercise label="list_MUC_rainfall_v2" numbered="yes" adaptive="yes" indentation="show" language="python">
    <statement>
      <p>Create a function <c>avg_rainfall(day_list)</c> that takes a list with the amount of rain for each day.
            There are some problems with the equipment, so ignore days that have a negative number for the amount of rain. Calculate the average rainfall as the sum of the non-negative values divided by the number of non-negative values. If at least one day has a non-negative rainfall return <c>Average: (average)</c> and otherwise return <c>No rain</c>.  For example,  <c>avg_rainfall([30,-2, 0])</c> should return <c>Average: 15</c> and <c>avg_rainfall([-3, -2])</c> should return <c>No rain</c>.</p>
    </statement>
    <blocks>
      <block order="4">
        <cline>def avg_rainfall(day_list):</cline>
      </block>
      <block order="3">
        <cline>    sum = 0</cline>
        <cline>    count = 0</cline>
      </block>
      <block order="12">
        <cline>    for day in day_list:</cline>
      </block>
      <block order="1">
        <choice correct="yes">
          <cline>        if day &gt;= 0:</cline>
        </choice>
        <choice>
          <cline>        if day &gt; 0: </cline>
        </choice>
      </block>
      <block order="11">
        <cline>            sum = sum + day</cline>
        <cline>            count = count + 1</cline>
      </block>
      <block order="7">
        <cline>    if count &gt; 0:</cline>
      </block>
      <block order="10">
        <cline>        avg = sum / count</cline>
      </block>
      <block order="8">
        <choice correct="yes">
          <cline>        return "Average: " + str(avg)</cline>
        </choice>
        <choice>
          <cline>        return "Average: " + avg </cline>
        </choice>
      </block>
      <block order="5">
        <cline>    else:</cline>
      </block>
      <block order="6">
        <cline>        return "No rain"</cline>
      </block>
    </blocks>
  </exercise>
  <exercise label="list_MUC_rainfall_v2_ac">
    <statement>
      <p>Write a function <c>avg_rainfall(day_list)</c> that takes a list with the amount of rain for each day.
            There are some problems with the equipment, so ignore days that have a negative number for the amount of rain. Calculate the average rainfall as the sum of the non-negative values divided by the number of non-negative values. If at least one day has a non-negative rainfall return <c>Average: (average)</c> and otherwise return <c>No rain</c>.  For example,  <c>avg_rainfall([30,-2, 0])</c> should return <c>Average: 15</c> and <c>avg_rainfall([-3, -2])</c> should return <c>No rain</c>.</p>
    </statement>
    <program xml:id="list_MUC_rainfall_v2_ac_editor" interactive="activecode" language="python">
      <input>

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(avg_rainfall([30, -2, 0]), "Average: 15.0", "rainfall([30, -2, 0])")
        self.assertEqual(avg_rainfall([-2, -3]), "No rain", "rainfall([-2, -3])")
        self.assertEqual(avg_rainfall([4, 0, -3, 20]), "Average: 8.0", "avg_rainfall([4, 0, -3, 20])")
        self.assertEqual(avg_rainfall([20, 4]), "Average: 12.0", "rainfall([20, 4])")

myTests().main()
        </input>
    </program>
  </exercise>
  <exercise label="list_MUC_mix_items" numbered="yes" adaptive="yes" indentation="show" language="python">
    <statement>
      <p>Create a function <c>mix</c> that takes a list <c>list1</c> and returns a new list with the first and last items from the original list swapped.  For example, <c>mix([3, 2, 1, 4]</c>
            should return <c>[4, 2, 1, 3]</c>.</p>
    </statement>
    <blocks>
      <block order="2">
        <cline>def mix(list1):</cline>
      </block>
      <block order="7">
        <cline>    first = list1[0]</cline>
        <cline>    last = list1[-1]</cline>
      </block>
      <block order="4">
        <cline>    out = [last]</cline>
      </block>
      <block order="3">
        <choice correct="yes">
          <cline>    out.extend(list1[1:-1])</cline>
        </choice>
        <choice>
          <cline>    out.extend(list1[1:-2]) </cline>
        </choice>
      </block>
      <block order="8">
        <choice correct="yes">
          <cline>    out.append(first)</cline>
        </choice>
        <choice>
          <cline>    out.extend(first) </cline>
        </choice>
      </block>
      <block order="6">
        <cline>    return out</cline>
      </block>
    </blocks>
  </exercise>
  <exercise label="list_MUC_mix_items_ac">
    <statement>
      <p>Write a function <c>mix</c> that takes a list <c>list1</c> and returns a new list with the first and last items from the original list swapped.  For example, <c>mix([3, 2, 1, 4])</c>
            should return <c>[4, 2, 1, 3]</c>.</p>
    </statement>
    <program xml:id="list_MUC_mix_items_ac_editor" interactive="activecode" language="python">
      <input>

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(mix([3, 2, 1, 4]), [4, 2, 1, 3], "mix([3, 2, 1, 4])")
        self.assertEqual(mix(['a', 'b', 'c']), ['c', 'b', 'a'], "mix(['a', 'b', 'c'])")
        self.assertEqual(mix(['a', 'b']), ['b', 'a'], "mix(['a', 'b'])")
        self.assertEqual(mix([2, 2, 1, 1]), [1, 2, 1, 2], "mix([2, 2, 1, 1])")


myTests().main()
        </input>
    </program>
  </exercise>
  <exercise label="list_MUC_discount" numbered="yes" adaptive="yes" indentation="show" language="python">
    <statement>
      <p>Create a function <c>discount(price_list)</c> that takes a list of prices and returns a new list with the original price reduced by half.
            Limit each price to two digits after the decimal by rounding to the closest value.  Hint, you can use the <c>round</c> function to do this.</p>
    </statement>
    <blocks>
      <block order="3">
        <cline>def discount(price_list):</cline>
      </block>
      <block order="4">
        <cline>    discounts = []</cline>
      </block>
      <block order="6">
        <cline>    for price in price_list:</cline>
      </block>
      <block order="5">
        <choice correct="yes">
          <cline>        new_price = round(price / 2, 2)</cline>
        </choice>
        <choice>
          <cline>        new_price = price / 2 </cline>
        </choice>
      </block>
      <block order="8">
        <choice correct="yes">
          <cline>        discounts.append(new_price)</cline>
        </choice>
        <choice>
          <cline>        discounts.extend(new_price) </cline>
        </choice>
      </block>
      <block order="1">
        <cline>    return discounts</cline>
      </block>
    </blocks>
  </exercise>
  <exercise label="list_MUC_discount_ac">
    <statement>
      <p>Write a function <c>discount(price_list)</c> that takes a list of prices and returns a new list with the original price reduced by half.
            Limit each price to two digits after the decimal by rounding to the closest value.  Hint, you can use the <c>round</c> function to do this.</p>
    </statement>
    <program xml:id="list_MUC_discount_ac_editor" interactive="activecode" language="python">
      <input>

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(discount([20.00, 15.60, 22.99]), [10.00, 7.80, 11.50], "discount([20.00, 15.60, 22.99])")
        self.assertEqual(discount([14.50]), [7.25], "discount([14.50])")
        self.assertEqual(discount([12.22]), [6.11], "discount([12.22])")
        self.assertEqual(discount([12.31]), [6.16], "discount([12.31])")

myTests().main()
        </input>
    </program>
  </exercise>
  <exercise label="list_MUC_append_if_a" numbered="yes" adaptive="yes" indentation="show" language="python">
    <statement>
      <p>Create a function <c>filter_a(word_list)</c> that takes a list of words, <c>word_list</c> and returns a new list with only the words in <c>word_list</c> that start with the letter <c>a</c>.
            For example, <c>filter_a(['a', 'bye', 'above'])</c> should return <c>['a', 'above']</c>.</p>
    </statement>
    <blocks>
      <block order="6">
        <cline>def filter_a(word_list):</cline>
      </block>
      <block order="4">
        <cline>    res = []</cline>
      </block>
      <block order="7">
        <cline>    for word in word_list:</cline>
      </block>
      <block order="5">
        <choice correct="yes">
          <cline>        if word[0] == 'a':</cline>
        </choice>
        <choice>
          <cline>        if word[0] = 'a': </cline>
        </choice>
      </block>
      <block order="3">
        <choice correct="yes">
          <cline>            res.append(word)</cline>
        </choice>
        <choice>
          <cline>            word_list.append(res) </cline>
        </choice>
      </block>
      <block order="2">
        <cline>    return res</cline>
      </block>
    </blocks>
  </exercise>
  <exercise label="list_MUC_append_if_a_ac">
    <statement>
      <p>Write a function <c>filter_a(word_list)</c> that takes a list of words, <c>word_list</c> and returns a new list with only the words in <c>word_list</c> that start with the letter <c>a</c>.
            For example, <c>filter_a(['a', 'bye', 'above'])</c> should return <c>['a', 'above']</c>.</p>
    </statement>
    <program xml:id="list_MUC_append_if_a_ac_editor" interactive="activecode" language="python">
      <input>

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(filter_a(['a', 'bye', 'above']), ['a', 'above'], "filter_a(['a', 'bye', 'above'])")
        self.assertEqual(filter_a(['bye']), [], "filter_a(['bye'])")
        self.assertEqual(filter_a(['bye', 'a', 'around']), ['a', 'around'], "filter_a(['bye', 'a', 'around'])")
        self.assertEqual(filter_a(['bye', 'above']), ['above'], "filter_a(['bye', 'above'])")



myTests().main()
        </input>
    </program>
  </exercise>
  <exercise label="list_MUC_add_vocab" numbered="yes" adaptive="yes" indentation="show" language="python">
    <statement>
      <p>Create a function <c>vocab_list</c> that takes two lists: <c>terms</c> and <c>vocab</c>. It should loop through the list of terms and add each item to
            <c>vocab</c> if it is not already in that list. It should return the list <c>vocab</c>.</p>
    </statement>
    <blocks>
      <block order="2">
        <cline>def vocab_list(terms, vocab):</cline>
      </block>
      <block order="3">
        <cline>    for word in terms:</cline>
      </block>
      <block order="9">
        <choice correct="yes">
          <cline>    if response not in wordList:</cline>
        </choice>
        <choice>
          <cline>    if response in wordList: </cline>
        </choice>
      </block>
      <block order="1">
        <cline>        vocabulary.append(response)</cline>
      </block>
      <block order="5">
        <choice correct="yes">
          <cline>        if word not in vocab:</cline>
        </choice>
        <choice>
          <cline>        if word !in vocab: </cline>
        </choice>
      </block>
      <block order="6">
        <choice correct="yes">
          <cline>            vocab.append(word)</cline>
        </choice>
        <choice>
          <cline>            word.append(vocab) </cline>
        </choice>
      </block>
      <block order="7">
        <cline>    return vocab</cline>
      </block>
    </blocks>
  </exercise>
  <exercise label="list_MUC_add_vocab_ac">
    <statement>
      <p>Write a function <c>vocab_list</c> that takes two lists: <c>terms</c> and <c>vocab</c>. It should loop through the list of terms and add each item to
            <c>vocab</c> if it is not already in that list. It should return the list <c>vocab</c>.</p>
    </statement>
    <program xml:id="list_MUC_add_vocab_ac_editor" interactive="activecode" language="python">
      <input>

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(vocab_list(['hi', 'bye'], ['hi', 'bye']), ['hi', 'bye'], "vocab_list(['hi', 'bye'], ['hi', 'bye'])")
        self.assertEqual(vocab_list(['hi', 'bye', 'run'], ['hi', 'bye']), ['hi', 'bye', 'run'], "vocab_list(['hi', 'bye', 'run'], ['hi', 'bye'])")
        self.assertEqual(vocab_list(['hi', 'bye'], []), ['hi', 'bye'], "vocab_list(['hi', 'bye'], [])")
        self.assertEqual(vocab_list(['an', 'anniversary'], ['anniversary']), ['anniversary', 'an'] , "vocab_list(['an', 'anniversary'], ['anniversary'])")


myTests().main()
        </input>
    </program>
  </exercise>
</section>
