<?xml version="1.0"?>
<section xml:id="lists_write-code-questions">
  <title>Write Code Questions</title>

  <activity xml:id="list_writeMyListq" list="list_writeMyListq">
    <statement>
      <p>Write a function <c>add_to_new_list</c> that takes in a list of strings, <c>lst</c>, as a parameter and creates a new list with the length
                    of <c>lst</c> and the first element of <c>lst</c> three times. For example, <c>add_to_new_list(["1","2","3"])</c> would return <c>[3, '111']</c>.</p>
    </statement>
    <program interactive="activecode" language="python">
      <code>
def add_to_new_list(lst):
    # write code here

print(add_to_new_list(["1","2","3"]))
=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(add_to_new_list(['1','2','3']), [3, '111'], "add_to_new_list(['1','2','3'])")
        self.assertEqual(add_to_new_list(['0','0','0','0']), [4, '000'], "add_to_new_list(['0','0','0','0'])")
        self.assertEqual(add_to_new_list(['10.2','0.0','100','-2']), [4, '10.210.210.2'], "add_to_new_list(['10.2','0.0','100','-2'])")


myTests().main()
        </code>
    </program>
    <solution>
        <p>
                Write a function add_to_new_list that takes in a list of strings, lst, as a parameter and creates a new list with the length
                    of lst and the first element of lst three times. For example, add_to_new_list(["1","2","3"]) would return [3, '111']
        </p>
    <program xml:id="list_writeMyListq_sol" interactive="activecode" language="python"><code>
def add_to_new_list(lst):
    new_list = []
    new_list.append(len(lst))
    new_list.append(lst[0] * 3)
    return new_list

print(add_to_new_list(["1","2","3"]))
=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(add_to_new_list(['1','2','3']), [3, '111'], "add_to_new_list(['1','2','3'])")
        self.assertEqual(add_to_new_list(['0','0','0','0']), [4, '000'], "add_to_new_list(['0','0','0','0'])")
        self.assertEqual(add_to_new_list(['10.2','0.0','100','-2']), [4, '10.210.210.2'], "add_to_new_list(['10.2','0.0','100','-2'])")


myTests().main()
        </code></program>
            </solution>
  </activity>
  
<activity label="list_writeItemsq_v2_sq"><select grade="first" questions="list_writeItemsq_v2_ac list_writeItemsq_v2_pp"/></activity>

  <activity xml:id="list_writeAvgq" label="list_writeAvgq">
    <statement>
      <p>Write a function <c>average</c> that takes in a list of integers, <c>aList</c>, as a parameter and returns the average of
                    all of the integers, rounded to one decimal place. For example, <c>average([99, 100, 74, 63, 100, 100])</c> would return <c>89.3</c>.</p>
    </statement>
    <program interactive="activecode" language="python">
      <code>
def average(aList):
    # write code here

print(average([99, 100, 74, 63, 100, 100]))

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertAlmostEqual(average([99, 100, 74, 63, 100, 100]), 89.3, 1, "average([99, 100, 74, 63, 100, 100])")
        self.assertAlmostEqual(average([0, 2, -3, 1.2, 2000]), 400.0, 1, "average([0, 2, -3, 1.2, 2000])")
        self.assertAlmostEqual(average([-2]), -2.0, 1, "average([-2])")


myTests().main()
        </code>
    </program>
    <solution>
                <p>Write a function average that takes in a list of integers, aList, as a parameter and returns the average of
                    all of the integers, rounded to one decimal place. For example, average([99, 100, 74, 63, 100, 100]) would return 89.33.</p>
    <program xml:id="list_writeAvgq_sol" interactive="activecode" language="python"><code>
def average(aList):
    sum = 0
    for num in aList:
        sum += num
    avg = round(sum / len(aList),2)
    return avg

print(average([99, 100, 74, 63, 100, 100]))
=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertAlmostEqual(average([99, 100, 74, 63, 100, 100]), 89.3, 1, "average([99, 100, 74, 63, 100, 100])")
        self.assertAlmostEqual(average([0, 2, -3, 1.2, 2000]), 400.0, 1, "average([0, 2, -3, 1.2, 2000])")
        self.assertAlmostEqual(average([-2]), -2.0, 1, "average([-2])")


myTests().main()
        </code></program>
            </solution>
  </activity>

  <activity label="list_write23_sq"><select grade="first" questions="list_write23_ac list_write23_pp"/></activity>

  <activity xml:id="list_capitalize_q" label="list_capitalize_q">
    <statement>
      <p>Write a function <c>capitalize</c> that takes in a list of lists of strings, <c>lst</c>, and makes the first letter of each element capitalized and adds
                    it to a new list and returns that new list. For example, <c>capitalize([["hi"],["hello", "hey"]])</c> would return <c>['Hi', 'Hello', 'Hey']</c>.</p>
    </statement>
    <program interactive="activecode" language="python">
      <code>
def capitalize(lst):
    # write code here

print(capitalize([["hi"],["hello", "hey"]]))
=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(capitalize([['hi'],['hello', 'hey']]), ['Hi', 'Hello', 'Hey'], "capitalize([['hi'],['hello', 'hey']])")
        self.assertEqual(capitalize([['HI'],['HELLO', 'HEY']]), ['Hi', 'Hello', 'Hey'], "capitalize([['HI'],['HELLO', 'HEY']])")
        self.assertEqual(capitalize([['go', 'blue'],['python', 'IS', 'The', 'Best']]), ['Go', 'Blue', 'Python', 'Is', 'The', 'Best'], "capitalize([['go', 'blue'],['python', 'IS', 'The', 'Best']])")

myTests().main()
        </code>
    </program>
    <solution>
            <p>Write a function capitalize that takes in a list of lists of strings, lst, and makes the first letter of each element capitalized and adds
                    it to a new list and returns that new list. For example, capitalize([["hi"],["hello", "hey"]]) would return ['Hi', 'Hello', 'Hey'].
                </p>
    <program xml:id="list_capitalize_q_sol" interactive="activecode" language="python"><code>
def capitalize(lst):
    new_list = []
    for i in lst:
        for j in i:
            new_list.append(j.capitalize())
    return new_list

print(capitalize([["hi"],["hello", "hey"]]))
=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(capitalize([['hi'],['hello', 'hey']]), ['Hi', 'Hello', 'Hey'], "capitalize([['hi'],['hello', 'hey']])")
        self.assertEqual(capitalize([['HI'],['HELLO', 'HEY']]), ['Hi', 'Hello', 'Hey'], "capitalize([['HI'],['HELLO', 'HEY']])")
        self.assertEqual(capitalize([['go', 'blue'],['python', 'IS', 'The', 'Best']]), ['Go', 'Blue', 'Python', 'Is', 'The', 'Best'], "capitalize([['go', 'blue'],['python', 'IS', 'The', 'Best']])")

myTests().main()
        </code></program>
            </solution>
  </activity>

  <activity label="list_write5_sq"><select grade="first" questions="list_write5_ac list_write5_pp"/></activity>

  <activity xml:id="list_writeChopq"  label="list_writeChopq">
    <statement>
      <p>Write a function <c>chop</c> that takes a list, <c>lst</c>, and modifies it, removing the first and last elements.
                    For example, <c>chop([1,2,3,4,5])</c> should return <c>[2,3,4]</c>.</p>
    </statement>
    <program interactive="activecode" language="python" >
      <code>
def chop(lst):
    # write code here

print(chop([1,2,3,4,5]))
=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(chop([1,2,3,4,5]),[2,3,4],"chop([1,2,3,4,5])")
        self.assertEqual(chop([1,3,5,7,9,10]),[3,5,7,9],"chop([1,3,5,7,9,10])")
        self.assertEqual(chop([2,9]),[],"chop([2,9])")

myTests().main()
        </code>
    </program>
    <solution>
            <p>Write a function chop that takes a list, lst, and modifies it, removing the first and last elements.
                    For example, chop([1,2,3,4,5] should return [2,3,4].</p>
    <program xml:id="list_writeChopq_sol" interactive="activecode" language="python"><code>
def chop(lst):
    lst.pop(0)
    lst.pop(-1)
    return(lst)

print(chop([1,2,3,4,5]))
=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(chop([1,2,3,4,5]),[2,3,4],"chop([1,2,3,4,5])")
        self.assertEqual(chop([1,3,5,7,9,10]),[3,5,7,9],"chop([1,3,5,7,9,10])")
        self.assertEqual(chop([2,9]),[],"chop([2,9])")

myTests().main()
        </code></program>
            </solution>
  </activity>

  <activity label="list_writeReverse_sq"><select grade="first" questions="list_writeReverse_ac list_writeReverse_pp"/></activity>

  <activity xml:id="list_writeSumq" label="list_writeSumq">
    <statement>
      <p>Write a function <c>sumUntilEven</c> that takes in one parameter, <c>lst</c>, and returns the sum of all the
                    elements in the <c>lst</c> up to but not including the first even number. For example, <c>sumUntilEven([1,2,3,4,5])</c>
                    should return <c>1</c> and <c>sumUntilEven([1,3,5,7,9])</c> should return <c>25</c>.</p>
    </statement>
    <program interactive="activecode" language="python" >
      <code>
def sumUntilEven(lst):
    # write code here

print(sumUntilEven([1,2,3,4,5]))

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sumUntilEven([1,2,3,4,5]),1,"sumUntilEven([1,2,3,4,5])")
        self.assertEqual(sumUntilEven([1,3,5,7,9]),25,"sumUntilEven([1,3,5,7,9])")
        self.assertEqual(sumUntilEven([2,4,6,7,9]),0,"sumUntilEven([2,4,6,7,9])")

myTests().main()
        </code>
    </program>
    <solution>
        <p>Write a function called sumUntilEven that takes in one parameter, lst, and returns the sum of all the
                    elements in the lst up to but not including the first even number. For example, sumUntilEven([1,2,3,4,5]
                    should return 1 and sumUntilEven([1,3,5,7,9] should return 25.
        </p>
    <program xml:id="list_writeSumq_sol" interactive="activecode" language="python"><code>
def sumUntilEven(lst):
    total = 0
    element = 0
    while element &lt; len(lst) and lst[element] % 2 != 0:
        total = total + lst[element]
        element += 1
    return total

print(sumUntilEven([1,2,3,4,5]))
====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sumUntilEven([1,2,3,4,5]),1,"sumUntilEven([1,2,3,4,5])")
        self.assertEqual(sumUntilEven([1,3,5,7,9]),25,"sumUntilEven([1,3,5,7,9])")
        self.assertEqual(sumUntilEven([2,4,6,7,9]),0,"sumUntilEven([2,4,6,7,9])")

myTests().main()
        </code></program>
            </solution>
  </activity>

<activity label="list_sortByLen_sq"><select grade="first" questions="list_sortByLen_ac list_sortByLen_pp"/></activity>

  <activity xml:id="list_loop_two_lists" label="list_loop_two_lists">
    <statement>
      <p>Write a function <c>combine(names, ages)</c> that takes in two lists, <c>names</c> and <c>ages</c> and returns a list of strings in the format <c>"Name: name, age: age"</c>.  For example,
            <c>combine(["Claire", "Jennifer"],[23, 19])</c> would return <c>["Name: Claire, age: 23", "Name: Jennfier, age: 19"]</c>.</p>
    </statement>
    <program interactive="activecode" language="python">
      <code>
def combine(names, ages):

print(combine(["Claire", "Jennifer"],[23, 19]))

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
          self.assertEqual(combine(["Mark", "Mary", "Sue"], [59, 60, 61]), ["Name: Mark, age: 59", "Name: Mary, age: 60", "Name: Sue, age: 61"], 'combine(["Mark", "Mary", "Sue"])')
          self.assertEqual(combine(["Iman"], [28]), ["Name: Iman, age: 28"], 'combine(["Iman"], [28])')
          self.assertEqual(combine(["Pedro", "Mez"], [24, 26]), ["Name: Pedro, age: 24", "Name: Mez, age: 26"], 'combine(["Pedro", "Mez"], [24, 26])')
          self.assertEqual(combine(["Claire", "Jennifer"], [23, 19]), ["Name: Claire, age: 23", "Name: Jennifer, age: 19"], 'combine(["Claire", "Jennifer"], [23, 19])')


myTests().main()
        </code>
    </program>
  </activity>
</section>
