<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.17.1: http://docutils.sourceforge.net/" />

    <title>Mixed-Up Code Questions &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.cd688d1623c9932e.css?v=630908D8" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.bf73008a1eba9f4e.css?v=7DED055E" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/runtime.439d581f758f2b99.bundle.js?v=2E54442"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.51d9dde4e5162190.bundle.js?v=9CAA2D33"></script>
    <script src="../_static/runestone.42b576791a4ed54b.bundle.js?v=B7293199"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    
    eBookConfig.useRunestoneServices = {{use_services}};
    eBookConfig.host = '';
    eBookConfig.app = eBookConfig.host + '/runestone';
    eBookConfig.course = '{{ course_name }}';
    eBookConfig.basecourse = '{{ base_course }}';
    eBookConfig.isLoggedIn = {{ is_logged_in}};
    eBookConfig.email = '{{ user_email }}';
    eBookConfig.isInstructor = {{ is_instructor }};
    eBookConfig.username = '{{ user_id}}';
    eBookConfig.readings = {{ readings|safe}};
    eBookConfig.activities = {{ activity_info|safe }}
    eBookConfig.downloadsEnabled = {{downloads_enabled}};
    eBookConfig.allow_pairs = {{allow_pairs}}
    eBookConfig.enableCompareMe = {{enable_compare_me}};
    // _`new_server_prefix`: Defined by the server to be the path to the new server.
    eBookConfig.new_server_prefix = '{{ new_server_prefix }}';
    
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->

    
        {% if serve_ad is defined and serve_ad and settings.adsenseid %}
         <script data-ad-client="{{settings.adsenseid}}" async src="https://pagead2.googlesyndication.com/pagead/js/adsbygoogle.js"></script>
        {% endif %}
    



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
            
            {{ course_name }}
            
          
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="browsing_warning"></span></li>

        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li id="inst_peer_link"><a href='/runestone/peer/instructor.html'>Peer Instruction (Instructor)</a></li>            
            <li><a href='/runestone/peer/student.html'>Peer Instruction (Student)</a></li>            
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/ns/books/published/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=mixedupcode_orig'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
            <script></script>
<li class="dropdown globaltoc-container">
    <a
        href="../index.html"
        class="dropdown-toggle"
        data-toggle="dropdown"
        aria-label="index"
        >This Chapter<b class="caret"></b
    ></a>
    <ul class="dropdown-menu globaltoc">
         
        {% for row in subchapter_list %}
        <li class="toctree-l1">
            <a class="reference internal" href="{{ row['subchap_uri']}}"
                >{{ row['title']}}</a
            >
        </li>
        {% endfor %} 
        
    </ul>
</li>
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->

    
        {% if settings.num_banners > 0 and settings.show_rs_banner %}
         <div style="width: 100%;">
            <a href="/runestone/default/donate?banner={{banner_num}}">
            <img class="runestone-banner"
                 src="/runestone/static/images/RunestoneBanner{{banner_num}}.png"
                 alt="Please Support Runestone"
                 style="display: block; width: 728px; margin-left: auto; margin-right: auto; margin-top: 10px;">
            </a>
         </div>
        {% endif %}
    


  
  <section id="mixed-up-code-questions">
<h1>Mixed-Up Code Questions<a class="headerlink" href="#mixed-up-code-questions" title="Permalink to this headline">¶</a></h1>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="dictionaries_mixedupcode_question10_1" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">letter_dict</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">string</span></code> as a parameter and returns a dictionary containing the amount of times letters appear in the given <code class="docutils literal notranslate"><span class="pre">string</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">letter_dict('It</span> <span class="pre">is')</span></code> should return <code class="docutils literal notranslate"><span class="pre">{'i':</span> <span class="pre">2,</span> <span class="pre">'t':</span> <span class="pre">1,</span> <span class="pre">'s':</span> <span class="pre">1}</span></code>. (Note: Pretend the <code class="docutils literal notranslate"><span class="pre">string</span></code> doesn’t have any punctuation besides spaces.)</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="1"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def letter_dict(string):
---
    d = dict()
---
    d = dictionary() #paired
---
    for char in string:
---
        if char == ' ':
---
        if char = ' ': #paired
---
            continue
---
            break #paired
---
        if char.lower() not in d:
---
        if char in d: #paired
---
            d[char.lower()] = 1
---
        else:
---
            d[char.lower()] += 1
---
    return d
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=dictionaries_mixedupcode_question10_1_ac data-question_label="2">
<div id=dictionaries_mixedupcode_question10_1_ac_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">letter_dict</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">string</span></code> as a parameter and returns a dictionary containing the amount of times letters appear in the given <code class="docutils literal notranslate"><span class="pre">string</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">letter_dict('It</span> <span class="pre">is')</span></code> should return <code class="docutils literal notranslate"><span class="pre">{'i':</span> <span class="pre">2,</span> <span class="pre">'t':</span> <span class="pre">1,</span> <span class="pre">'s':</span> <span class="pre">1}</span></code>. (Note: Pretend the <code class="docutils literal notranslate"><span class="pre">string</span></code> doesn’t have any punctuation besides spaces.)</p>

</div>
<textarea data-lang="python" id="dictionaries_mixedupcode_question10_1_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(letter_dict('It is'),{'i': 2, 't': 1, 's': 1},"letter_dict('It is')")
        self.assertEqual(letter_dict('It isSsSsS'),{'i': 2, 't': 1, 's': 6},"letter_dict('It isSsSsS')")
        self.assertEqual(letter_dict('HELLO'),{'h': 1, 'e': 1, 'l': 2, 'o': 1},"letter_dict('HELLO')")
        self.assertEqual(letter_dict('goodbye'),{'g': 1, 'o': 2, 'd': 1, 'b': 1, 'y': 1, 'e': 1},"letter_dict('goodbye')")
        self.assertEqual(letter_dict(' Nice to meet      YOU   '),{'n': 1, 'i': 1, 'c': 1, 'e': 3, 't': 2, 'o': 2, 'm': 1, 'y': 1, 'u': 1},"letter_dict(' Nice to meet      YOU   ')")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="dictionaries_mixedupcode_question10_2" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">word_dict</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">string</span></code> as a parameter and returns a dictionary containing the amount of times a word appears in the given <code class="docutils literal notranslate"><span class="pre">string</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">word_dict('HELLO</span> <span class="pre">heLlo</span> <span class="pre">there</span>&#160; <span class="pre">')</span></code> should return <code class="docutils literal notranslate"><span class="pre">{'hello':</span> <span class="pre">2,</span> <span class="pre">'there':</span> <span class="pre">1}</span></code>. (Note: Pretend the <code class="docutils literal notranslate"><span class="pre">string</span></code> doesn’t have any punctuation besides spaces.)</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="3"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def word_dict(string):
---
    dictionary = {}
---
    for word in string.split():
---
    for word.lower() in string.split(): #paired
---
        dictionary[word.lower()] = dictionary.get(word.lower(), 0) + 1
---
        dictionary[word.lower()] = dictionary.get(word.lower(), 0) #paired
---
    return dictionary
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=dictionaries_mixedupcode_question10_2_ac data-question_label="4">
<div id=dictionaries_mixedupcode_question10_2_ac_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">word_dict</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">string</span></code> as a parameter and returns a dictionary containing the amount of times a word appears in the given <code class="docutils literal notranslate"><span class="pre">string</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">word_dict('HELLO</span> <span class="pre">heLlo</span> <span class="pre">there</span>&#160; <span class="pre">')</span></code> should return <code class="docutils literal notranslate"><span class="pre">{'hello':</span> <span class="pre">2,</span> <span class="pre">'there':</span> <span class="pre">1}</span></code>. (Note: Pretend the <code class="docutils literal notranslate"><span class="pre">string</span></code> doesn’t have any punctuation besides spaces.)</p>

</div>
<textarea data-lang="python" id="dictionaries_mixedupcode_question10_2_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(word_dict('HELLO heLlo there  '),{'hello': 2, 'there': 1},"word_dict('HELLO heLlo there  ')")
        self.assertEqual(word_dict(' Hello HELLO he helLO hello HE'),{'hello': 4, 'he': 2},"word_dict(' Hello HELLO he helLO hello HE')")
        self.assertEqual(word_dict('There is a rainbow'),{'there': 1, 'is': 1, 'a': 1, 'rainbow': 1},"word_dict('There is a rainbow')")
        self.assertEqual(word_dict('That rainbow is more beautiful than any other rainbow I have seen before'),{'that': 1, 'rainbow': 2, 'is': 1, 'more': 1, 'beautiful': 1, 'than': 1, 'any': 1, 'other': 1, 'i': 1, 'have': 1, 'seen': 1, 'before': 1},"word_dict('That rainbow is more beautiful than any other rainbow I have seen before')")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="dictionaries_mixedupcode_question10_3" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function <code class="docutils literal notranslate"><span class="pre">a_counter</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">string</span></code> as a parameter, creates a <code class="docutils literal notranslate"><span class="pre">dictionary</span></code>, and returns the amount of a’s in the given <code class="docutils literal notranslate"><span class="pre">string</span></code> using the <code class="docutils literal notranslate"><span class="pre">dictionary</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">a_counter('Alice</span> <span class="pre">and</span> <span class="pre">Apples')</span></code> should return <code class="docutils literal notranslate"><span class="pre">3</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="5"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def a_counter(string):
---
    dictionary = {}
---
    for char in string:
---
    for char in string.split(): #paired
---
        dictionary[char.lower()] = dictionary.get(char.lower(), 0) + 1
---
        dictionary[char.lower()] = dictionary.get(char.lower(), 0) #paired
---
    return dictionary.get('a', 0)
---
    return dictionary['a'] #paired
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=dictionaries_mixedupcode_question10_3_ac data-question_label="6">
<div id=dictionaries_mixedupcode_question10_3_ac_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">a_counter</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">string</span></code> as a parameter, creates a <code class="docutils literal notranslate"><span class="pre">dictionary</span></code>, and returns the amount of a’s in the given <code class="docutils literal notranslate"><span class="pre">string</span></code> using the <code class="docutils literal notranslate"><span class="pre">dictionary</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">a_counter('Alice</span> <span class="pre">and</span> <span class="pre">Apples')</span></code> should return <code class="docutils literal notranslate"><span class="pre">3</span></code>.</p>

</div>
<textarea data-lang="python" id="dictionaries_mixedupcode_question10_3_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(a_counter('Alice and Apples'),3,"a_counter('Alice and Apples')")
        self.assertEqual(a_counter('Australia is awesome!'),4,"a_counter('Australia is awesome!')")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="dictionaries_mixedupcode_question10_4" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">sales_tax</span></code> that takes in a <code class="docutils literal notranslate"><span class="pre">list_of_costs</span></code> (floating point values) as a parameter and
returns a dictionary with the original costs as keys and the final costs after the sales tax of 6 percent as values.
The final costs should be rounded to the hundredths place. For example, <code class="docutils literal notranslate"><span class="pre">sales_tax([1.25,</span> <span class="pre">8,</span> <span class="pre">5])</span></code> returns <code class="docutils literal notranslate"><span class="pre">{1.25:</span> <span class="pre">1.33,</span> <span class="pre">8:</span> <span class="pre">8.48,</span> <span class="pre">5:</span> <span class="pre">5.3}</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="7"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def sales_tax(list_of_costs):
---
    final_dictionary = {}
---
    final_dictionary = [] #paired
---
    for cost in list_of_costs:
---
        final_dictionary[cost] = round(cost * 1.06, 2)
---
        final_dictionary[cost] = round(cost * 0.06, 2) #paired
---
    return final_dictionary
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=dictionaries_mixedupcode_question10_4_ac data-question_label="8">
<div id=dictionaries_mixedupcode_question10_4_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">sales_tax</span></code> that takes in a <code class="docutils literal notranslate"><span class="pre">list_of_costs</span></code> (floating point values) as a parameter and
returns a dictionary with the original costs as keys and the final costs after the sales tax of 6 percent as values.
The final costs should be rounded to the hundredths place. For example, <code class="docutils literal notranslate"><span class="pre">sales_tax([1.25,</span> <span class="pre">8,</span> <span class="pre">5])</span></code> returns <code class="docutils literal notranslate"><span class="pre">{1.25:</span> <span class="pre">1.33,</span> <span class="pre">8:</span> <span class="pre">8.48,</span> <span class="pre">5:</span> <span class="pre">5.3}</span></code>.</p>

</div>
<textarea data-lang="python" id="dictionaries_mixedupcode_question10_4_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sales_tax([1.25, 8, 5]),{1.25: 1.33, 8: 8.48, 5: 5.3},"sales_tax([1.25, 8, 5])")
        self.assertEqual(sales_tax([0]),{0: 0.0},"sales_tax([0])")
        self.assertEqual(sales_tax([5]),{5: 5.3},"sales_tax([5])")
        self.assertEqual(sales_tax([18.5, 100, 2000.0]),{18.5: 19.61, 100: 106.0, 2000.0: 2120.0},"sales_tax([18.5, 100, 2000.0])")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="dictionaries_mixedupcode_question10_5" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">discount_and_tax</span></code> that takes in a <code class="docutils literal notranslate"><span class="pre">list_of_costs</span></code> (floating point values) and a <code class="docutils literal notranslate"><span class="pre">discount</span></code> percentage as parameters and
returns a dictionary with the original costs as keys and the final costs after applying the <code class="docutils literal notranslate"><span class="pre">discount</span></code> (first) and the sales tax of 6 percent (second) as values.
The final costs should be rounded to the hundredths place. For example, <code class="docutils literal notranslate"><span class="pre">discount_and_tax([1.25,</span> <span class="pre">8,</span> <span class="pre">5],</span> <span class="pre">20)</span></code> should return <code class="docutils literal notranslate"><span class="pre">{1.25:</span> <span class="pre">1.06,</span> <span class="pre">8:</span> <span class="pre">6.78,</span> <span class="pre">5:</span> <span class="pre">4.24}</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="9"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def discount_and_tax(list_of_costs, discount):
---
    discount_multiplier = 1 - (discount / 100)
---
    discount_multiplier = (discount / 100) #paired
---
    dictionary = {}
---
    for cost in list_of_costs:
---
        dictionary[cost] = round(cost * discount_multiplier * 1.06, 2)
---
        dictionary[cost] = round(cost * discount_multiplier * 0.06, 2) #paired
---
    return dictionary
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=dictionaries_mixedupcode_question10_5_ac data-question_label="10">
<div id=dictionaries_mixedupcode_question10_5_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">discount_and_tax</span></code> that takes in a <code class="docutils literal notranslate"><span class="pre">list_of_costs</span></code> (floating point values) and a <code class="docutils literal notranslate"><span class="pre">discount</span></code> percentage as parameters and
returns a dictionary with the original costs as keys and the final costs after applying the <code class="docutils literal notranslate"><span class="pre">discount</span></code> (first) and the sales tax of 6 percent (second) as values.
The final costs should be rounded to the hundredths place. For example, <code class="docutils literal notranslate"><span class="pre">discount_and_tax([1.25,</span> <span class="pre">8,</span> <span class="pre">5],</span> <span class="pre">20)</span></code> should return <code class="docutils literal notranslate"><span class="pre">{1.25:</span> <span class="pre">1.06,</span> <span class="pre">8:</span> <span class="pre">6.78,</span> <span class="pre">5:</span> <span class="pre">4.24}</span></code>.</p>

</div>
<textarea data-lang="python" id="dictionaries_mixedupcode_question10_5_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(discount_and_tax([1.25, 8, 5], 20),{1.25: 1.06, 8: 6.78, 5: 4.24},"discount_and_tax([1.25, 8, 5], 20)")
        self.assertEqual(discount_and_tax([0], 85),{0: 0.0},"discount_and_tax([0], 85)")
        self.assertEqual(discount_and_tax([5], 15),{5: 4.5},"discount_and_tax([5], 15)")
        self.assertEqual(discount_and_tax([18.5, 100, 2000.0], 1),{18.5: 19.41, 100: 104.94, 2000.0: 2098.8},"discount_and_tax([18.5, 100, 2000.0], 1)")
        self.assertEqual(discount_and_tax([18.5, 100, 2000.0], 99),{18.5: 0.2, 100: 1.06, 2000.0: 21.2},"discount_and_tax([18.5, 100, 2000.0], 99)")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="dictionaries_mixedupcode_question10_6" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">create_dictionary</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">keys_list</span></code> and <code class="docutils literal notranslate"><span class="pre">values_list</span></code> as parameters
and returns a dictionary with the <code class="docutils literal notranslate"><span class="pre">keys_list</span></code> as keys and the <code class="docutils literal notranslate"><span class="pre">values_list</span></code> as values. Assume <code class="docutils literal notranslate"><span class="pre">keys_list</span></code> and <code class="docutils literal notranslate"><span class="pre">values_list</span></code> are the same length, and
they only contain numbers and strings. For example, <code class="docutils literal notranslate"><span class="pre">create_dictionary(['one',</span> <span class="pre">'two',</span> <span class="pre">'three'],</span> <span class="pre">[1,</span> <span class="pre">2,</span> <span class="pre">3])</span></code> should return <code class="docutils literal notranslate"><span class="pre">{'one':</span> <span class="pre">1,</span> <span class="pre">'two':</span> <span class="pre">2,</span> <span class="pre">'three':</span> <span class="pre">3}</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="11"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def create_dictionary(keys_list, values_list):
---
    combination = {}
---
    for i in range(len(keys_list)):
---
    for i in len(keys_list): #paired
---
        combination[keys_list[i]] = values_list[i]
---
        combination[keys_list[i]] = combination[values_list[i]] #paired
---
    return combination
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=dictionaries_mixedupcode_question10_6_ac data-question_label="12">
<div id=dictionaries_mixedupcode_question10_6_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">create_dictionary</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">keys_list</span></code> and <code class="docutils literal notranslate"><span class="pre">values_list</span></code> as parameters
and returns a dictionary with the <code class="docutils literal notranslate"><span class="pre">keys_list</span></code> as keys and the <code class="docutils literal notranslate"><span class="pre">values_list</span></code> as values. Assume <code class="docutils literal notranslate"><span class="pre">keys_list</span></code> and <code class="docutils literal notranslate"><span class="pre">values_list</span></code> are the same length, and
they only contain numbers and strings. For example, <code class="docutils literal notranslate"><span class="pre">create_dictionary(['one',</span> <span class="pre">'two',</span> <span class="pre">'three'],</span> <span class="pre">[1,</span> <span class="pre">2,</span> <span class="pre">3])</span></code> should return <code class="docutils literal notranslate"><span class="pre">{'one':</span> <span class="pre">1,</span> <span class="pre">'two':</span> <span class="pre">2,</span> <span class="pre">'three':</span> <span class="pre">3}</span></code>.</p>

</div>
<textarea data-lang="python" id="dictionaries_mixedupcode_question10_6_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(create_dictionary(['one', 'two', 'three'], [1, 2, 3]),{'one': 1, 'two': 2, 'three': 3},"create_dictionary(['one', 'two', 'three'], [1, 2, 3])")
        self.assertEqual(create_dictionary(['one', 2.5, 'three', 4], [1, 'two and a half', 3, 'four']),{'one': 1, 2.5: 'two and a half', 'three': 3, 4: 'four'},"create_dictionary(['one', 2.5, 'three', 4], [1, 'two and a half', 3, 'four'])")
        self.assertEqual(create_dictionary(['me'], ['I']),{'me': 'I'},"create_dictionary(['me'], ['I'])")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="dictionaries_mixedupcode_question10_7" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">lowest_value_keys</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">dictionary</span></code> as a parameter and returns a list of the keys containing the lowest value of all the keys in the dictionary.
Assume all values of the <code class="docutils literal notranslate"><span class="pre">dictionary</span></code> are integers or floats. For example, <code class="docutils literal notranslate"><span class="pre">lowest_value_keys({'pizza':</span> <span class="pre">11,</span> <span class="pre">'fruits':</span> <span class="pre">0,</span> <span class="pre">'toys':</span> <span class="pre">0,</span> <span class="pre">'rice</span> <span class="pre">bags':</span> <span class="pre">1.9,</span> <span class="pre">'ice':</span> <span class="pre">1.8})</span></code> should return <code class="docutils literal notranslate"><span class="pre">['fruits',</span> <span class="pre">'toys']</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="13"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def lowest_value_keys(dictionary):
---
    keys_list = []
---
    minimum = min(dictionary.values())
---
    minimum = min(dictionary) #paired
---
    minimum = minimum(dictionary) #paired
---
    for key, value in dictionary.items():
---
    for value in dictionary.values(): #paired
---
    for key in dictionary.keys(): #paired
---
        if value == minimum:
---
            keys_list.append(key)
---
    return keys_list
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=dictionaries_mixedupcode_question10_7_ac data-question_label="14">
<div id=dictionaries_mixedupcode_question10_7_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">lowest_value_keys</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">dictionary</span></code> as a parameter and returns a list of the keys containing the lowest value of all the keys in the dictionary.
Assume all values of the <code class="docutils literal notranslate"><span class="pre">dictionary</span></code> are integers or floats. For example, <code class="docutils literal notranslate"><span class="pre">lowest_value_keys({'pizza':</span> <span class="pre">11,</span> <span class="pre">'fruits':</span> <span class="pre">0,</span> <span class="pre">'toys':</span> <span class="pre">0,</span> <span class="pre">'rice</span> <span class="pre">bags':</span> <span class="pre">1.9,</span> <span class="pre">'ice':</span> <span class="pre">1.8})</span></code> should return <code class="docutils literal notranslate"><span class="pre">['fruits',</span> <span class="pre">'toys']</span></code>.</p>

</div>
<textarea data-lang="python" id="dictionaries_mixedupcode_question10_7_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(lowest_value_keys({'pizza': 5, 'fruits': 20, 'toys': 3, 'rice bags': 2}),['rice bags'],"lowest_value_keys({'pizza': 5, 'fruits': 20, 'toys': 3, 'rice bags': 2})")
        self.assertEqual(lowest_value_keys({'pizza': 5, 'fruits': 20, 'toys': 3, 'rice bags': 1.9, 'ice': 1.8}),['ice'],"lowest_value_keys({'pizza': 5, 'fruits': 20, 'toys': 3, 'rice bags': 1.9, 'ice': 1.8})")
        self.assertEqual(lowest_value_keys({'pizza': 0, 'fruits': 20, 'toys': 3, 'rice bags': 1.9, 'ice': 1.8}),['pizza'],"lowest_value_keys({'pizza': 0, 'fruits': 20, 'toys': 3, 'rice bags': 1.9, 'ice': 1.8})")
        self.assertEqual(lowest_value_keys({'pizza': 11, 'fruits': 0, 'toys': 0, 'rice bags': 1.9, 'ice': 1.8}),['fruits', 'toys'],"lowest_value_keys({'pizza': 11, 'fruits': 0, 'toys': 0, 'rice bags': 1.9, 'ice': 1.8})")
        self.assertEqual(lowest_value_keys({'pizza': 100}),['pizza'],"lowest_value_keys({'pizza': 100})")
        self.assertEqual(lowest_value_keys({'pizza': 100, 'pizza2': 100, 'pizza3': 100}),['pizza', 'pizza2', 'pizza3'],"lowest_value_keys({'pizza': 100, 'pizza2': 100, 'pizza3': 100})")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="dictionaries_mixedupcode_question10_8" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">duplicated_data</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">dictionary1</span></code> and <code class="docutils literal notranslate"><span class="pre">dictionary2</span></code> as parameters and returns a dictionary with key-value pairs that are
in both dictionaries. For example, <code class="docutils literal notranslate"><span class="pre">duplicated_data({'Ten':</span> <span class="pre">10,</span> <span class="pre">'Twenty'</span> <span class="pre">:</span> <span class="pre">20,</span> <span class="pre">'Thirty'</span> <span class="pre">:</span> <span class="pre">30},</span> <span class="pre">{'Ten':</span> <span class="pre">10,</span> <span class="pre">'Twenty'</span> <span class="pre">:</span> <span class="pre">20})</span></code> should return <code class="docutils literal notranslate"><span class="pre">{'Ten':</span> <span class="pre">10,</span> <span class="pre">'Twenty':</span> <span class="pre">20}</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="15"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def duplicated_data(dictionary1, dictionary2):
---
    dict = {}
---
    for key, value in dictionary1.items():
---
    for key, value in dictionary2: #paired
---
    for key, value in dictionary2.items(): #paired
---
        if dictionary2.get(key) == value:
---
        if dictionary2[key] == value: #paired
---
            dict[key] = value
---
    return dict
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=dictionaries_mixedupcode_question10_8_ac data-question_label="16">
<div id=dictionaries_mixedupcode_question10_8_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">duplicated_data</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">dictionary1</span></code> and <code class="docutils literal notranslate"><span class="pre">dictionary2</span></code> as parameters and returns a dictionary with key-value pairs that are
in both dictionaries. For example, <code class="docutils literal notranslate"><span class="pre">duplicated_data({'Ten':</span> <span class="pre">10,</span> <span class="pre">'Twenty'</span> <span class="pre">:</span> <span class="pre">20,</span> <span class="pre">'Thirty'</span> <span class="pre">:</span> <span class="pre">30},</span> <span class="pre">{'Ten':</span> <span class="pre">10,</span> <span class="pre">'Twenty'</span> <span class="pre">:</span> <span class="pre">20})</span></code> should return <code class="docutils literal notranslate"><span class="pre">{'Ten':</span> <span class="pre">10,</span> <span class="pre">'Twenty':</span> <span class="pre">20}</span></code>.</p>

</div>
<textarea data-lang="python" id="dictionaries_mixedupcode_question10_8_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(duplicated_data({'Ten': 10, 'Twenty' : 20, 'Thirty' : 30}, {'Ten': 10, 'Twenty' : 20, 'Thirty' : 30}),{'Ten': 10, 'Twenty': 20, 'Thirty': 30},"duplicated_data({'Ten': 10, 'Twenty' : 20, 'Thirty' : 30}, {'Ten': 10, 'Twenty' : 20, 'Thirty' : 30})")
        self.assertEqual(duplicated_data({'Ten': 10, 'Twenty' : 20, 'Thirty' : 30}, {'Ten': 10, 'Twenty' : 20}),{'Ten': 10, 'Twenty': 20},"duplicated_data({'Ten': 10, 'Twenty' : 20, 'Thirty' : 30}, {'Ten': 10, 'Twenty' : 20})")
        self.assertEqual(duplicated_data({'Ten': 10, 'Twenty' : 20, 'Thirty' : 30}, {'Five': 5, 'Six' : 6}),{},"duplicated_data({'Ten': 10, 'Twenty' : 20, 'Thirty' : 30}, {'Five': 5, 'Six' : 6})")
        self.assertEqual(duplicated_data({'Ten': 10, 'Twenty' : 20, 'Five' : 'Five'}, {'Five': 'Five', 'Six' : 'Six'}),{'Five': 'Five'},"duplicated_data({'Ten': 10, 'Twenty' : 20, 'Five' : 'Five'}, {'Five': 'Five', 'Six' : 'Six'})")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="dictionaries_mixedupcode_question10_9" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">total_price</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">quantity_dict</span></code> and <code class="docutils literal notranslate"><span class="pre">price_dict</span></code> and returns the total cost of all items.
Assume keys in both dictionaries are the same, and the values in <code class="docutils literal notranslate"><span class="pre">quantity_dict</span></code> are integers.
For example, <code class="docutils literal notranslate"><span class="pre">total_price({'fries':</span> <span class="pre">7,</span> <span class="pre">'hot</span> <span class="pre">dogs':</span> <span class="pre">9,</span> <span class="pre">'soda':</span> <span class="pre">9},</span> <span class="pre">{'fries':</span> <span class="pre">1.5,</span> <span class="pre">'hot</span> <span class="pre">dogs':</span> <span class="pre">1,</span> <span class="pre">'soda':</span> <span class="pre">1.1})</span></code> should return <code class="docutils literal notranslate"><span class="pre">29.4</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="17"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def total_price(quantity_dict, price_dict):
---
    total_price = 0
---
    for key in quantity_dict.keys():
---
    for value in quantity_dict.values(): #paired
---
        total_price += quantity_dict.get(key) * price_dict.get(key)
---
        total_price += quantity_dict[key] + price_dict[key] #paired
---
    return total_price
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=dictionaries_mixedupcode_question10_9_ac data-question_label="18">
<div id=dictionaries_mixedupcode_question10_9_ac_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">total_price</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">quantity_dict</span></code> and <code class="docutils literal notranslate"><span class="pre">price_dict</span></code> and returns the total cost of all items.
Assume keys in both dictionaries are the same, and the values in <code class="docutils literal notranslate"><span class="pre">quantity_dict</span></code> are integers.
For example, <code class="docutils literal notranslate"><span class="pre">total_price({'fries':</span> <span class="pre">7,</span> <span class="pre">'hot</span> <span class="pre">dogs':</span> <span class="pre">9,</span> <span class="pre">'soda':</span> <span class="pre">9},</span> <span class="pre">{'fries':</span> <span class="pre">1.5,</span> <span class="pre">'hot</span> <span class="pre">dogs':</span> <span class="pre">1,</span> <span class="pre">'soda':</span> <span class="pre">1.1})</span></code> should return <code class="docutils literal notranslate"><span class="pre">29.4</span></code>.</p>

</div>
<textarea data-lang="python" id="dictionaries_mixedupcode_question10_9_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertAlmostEqual(total_price({'fries': 7, 'hot dogs': 9, 'soda': 9}, {'fries': 1.5, 'hot dogs': 1, 'soda': 1.1}),29.4,1,"total_price({'fries': 7, 'hot dogs': 9, 'soda': 9}, {'fries': 1.5, 'hot dogs': 1, 'soda': 1.1})")
        self.assertAlmostEqual(total_price({'fries': 0, 'hot dogs': 48, 'soda': 3}, {'fries': 1.13, 'hot dogs': 3, 'soda': 3.97}),155.91,2,"total_price({'fries': 0, 'hot dogs': 48, 'soda': 3}, {'fries': 1.13, 'hot dogs': 3, 'soda': 3.97})")
        self.assertAlmostEqual(total_price({'pizza': 5}, {'pizza': 2.50}),12.5,1,"total_price({'pizza': 5}, {'pizza': 2.50})")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="dictionaries_mixedupcode_question10_10" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Create a function called <code class="docutils literal notranslate"><span class="pre">scale_recipe</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">quantity_dict</span></code> and <code class="docutils literal notranslate"><span class="pre">scale_factor</span></code> as parameters and
returns a dictionary with <code class="docutils literal notranslate"><span class="pre">quantity_dict</span></code> scaled by the <code class="docutils literal notranslate"><span class="pre">scale_factor</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">scale_recipe({'fries':</span> <span class="pre">7,</span> <span class="pre">'hot</span> <span class="pre">dogs':</span> <span class="pre">9,</span> <span class="pre">'soda':</span> <span class="pre">8},</span> <span class="pre">4)</span></code> should return <code class="docutils literal notranslate"><span class="pre">{'fries':</span> <span class="pre">28,</span> <span class="pre">'hot</span> <span class="pre">dogs':</span> <span class="pre">36,</span> <span class="pre">'soda':</span> <span class="pre">32}</span></code>.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="19"   data-adaptive="true"       data-numbered="left"  style="visibility: hidden;">
        def scale_recipe(quantity_dict, scale_factor):
---
    for key in quantity_dict.keys():
---
        quantity_dict[key] = quantity_dict.get(key) * scale_factor
---
        quantity_dict[key] = quantity_dict.get(key) / scale_factor #paired
---
    return quantity_dict
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=dictionaries_mixedupcode_question10_10_ac data-question_label="20">
<div id=dictionaries_mixedupcode_question10_10_ac_question class="ac_question col-md-12">
<p>Create a function called <code class="docutils literal notranslate"><span class="pre">scale_recipe</span></code> that takes in <code class="docutils literal notranslate"><span class="pre">quantity_dict</span></code> and <code class="docutils literal notranslate"><span class="pre">scale_factor</span></code> as parameters and
returns a dictionary with <code class="docutils literal notranslate"><span class="pre">quantity_dict</span></code> scaled by the <code class="docutils literal notranslate"><span class="pre">scale_factor</span></code>.
For example, <code class="docutils literal notranslate"><span class="pre">scale_recipe({'fries':</span> <span class="pre">7,</span> <span class="pre">'hot</span> <span class="pre">dogs':</span> <span class="pre">9,</span> <span class="pre">'soda':</span> <span class="pre">8},</span> <span class="pre">4)</span></code> should return <code class="docutils literal notranslate"><span class="pre">{'fries':</span> <span class="pre">28,</span> <span class="pre">'hot</span> <span class="pre">dogs':</span> <span class="pre">36,</span> <span class="pre">'soda':</span> <span class="pre">32}</span></code>.</p>

</div>
<textarea data-lang="python" id="dictionaries_mixedupcode_question10_10_ac_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(scale_recipe({'fries': 7, 'hot dogs': 9, 'soda': 8}, 4),{'fries': 28, 'hot dogs': 36, 'soda': 32},"scale_recipe({'fries': 7, 'hot dogs': 9, 'soda': 8}, 4)")
        self.assertEqual(scale_recipe({'fries': 7, 'hot dogs': 9, 'soda': 8}, 0.5),{'fries': 3.5, 'hot dogs': 4.5, 'soda': 4.0},"scale_recipe({'fries': 7, 'hot dogs': 9, 'soda': 8}, 0.5)")
        self.assertEqual(scale_recipe({'brownie mix': 1, 'oil': 1/3, 'egg': 1, 'water': 1/3}, 4),{'brownie mix': 4, 'oil': 1.3333333333333333, 'egg': 4, 'water': 1.3333333333333333},"scale_recipe({'brownie mix': 1, 'oil': 1/3, 'egg': 1, 'water': 1/3}, 4)")

myTests().main()
</textarea>
</div>
</div>
</section>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>



  
    {% if settings.google_ga %}
      <script>
        var _gaq = _gaq || [];
        _gaq.push(['_setAccount', '{{ settings.google_ga }}']);
        _gaq.push(['_trackPageview']);

        (function() {
          var ga = document.createElement('script'); ga.async = true;
          ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
          var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
        })();
      </script>
    {% endif %}
  

  
    <script type="application/json" class="js-hypothesis-config">
    {
      "showHighlights": false
    }
    </script>

    <script async src="https://hypothes.is/embed.js"></script>
    <!--provide space to prevent hypothes.is sidebar from obscuring content-->
    <style>@media screen and (max-width:767px) { body { padding-right: 21px; }}</style>

  


<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>