<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>10.10. Write Code Questions &#8212; Python for Everybody - Interactive</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/637.css?v=bddd2011b20db76a20ea" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.css?v=7c5388ad4389eae1a318" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/runtime.bundle.js?v=aa7c717c38c5405af9e6"></script>
    <script src="../_static/637.bundle.js?v=1123ad75f9b821cdd636"></script>
    <script src="../_static/runestone.bundle.js?v=9c8e466e8cbed4a3af5a"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="11. Tuples" href="../tuples/toctree.html" />
    <link rel="prev" title="10.9. Mixed-Up Code Questions" href="mixedupcode.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'py4e-int';
    eBookConfig.basecourse = 'py4e-int';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = '' ? '' : 'python'
  eBookConfig.runestone_version = '';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          py4e-int
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=py4e-int&page=writecode'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">10.9. </span>Mixed-Up Code Questions" data-toggle="tooltip">
    <a href="mixedupcode.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">11. </span>Tuples' data-toggle="tooltip" >              
    <a href="../tuples/toctree.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
  
<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="write-code-questions">
<h1><span class="section-number">10.10. </span>Write Code Questions<a class="headerlink" href="#write-code-questions" title="Permalink to this headline">¶</a></h1>
<ol class="arabic">
<li><div id='dict_writecode1' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=dict_writecode1q data-question_label="10.10.1">
<div id=dict_writecode1q_question class="ac_question col-md-12">
<p>Write a program that categorizes each mail message by which day of the week
the commit was done. To do this, look for lines that start with “From”. Then,
look for the third word, and keep a running count of each of the days of the week.
At the end of the program, print out the contents of the dictionary <code class="docutils literal notranslate"><span class="pre">mail_count</span></code>
(order does not matter). For example, <code class="docutils literal notranslate"><span class="pre">mail_count['Mon']</span></code> should be <code class="docutils literal notranslate"><span class="pre">2</span></code>.</p>

</div>
<textarea data-lang="python" id="dict_writecode1q_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
mail = ['From stephen.marquard@uct.ac.za Sat Jan  7', 'From gopal.ramasammycook@gmail.com Thurs Jan  5', 'From louis@media.berkeley.edu Tues Jan  3', 'From antranig@caret.cam.ac.uk Sat Jan  7', 'From david.horwitz@uct.ac.za Wed Jan  4', 'From ray@media.berkeley.edu Mon Jan  2', 'From stephen.marquard@uct.ac.za Mon Jan 2', 'From wagnermr@iupui.edu Fri Jan  6']

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(mail_count['Mon'], 2, "mail_count['Mon']")
        self.assertEqual(mail_count['Sat'], 2, "mail_count['Sat']")
        self.assertEqual(mail_count.get('Sun', 0), 0, "mail_count.get('Sun', 0)")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=dict_writecode1a data-question_label="10.10.2">
<div id=dict_writecode1a_question class="ac_question col-md-12">
<p>Write a program that categorizes each mail message by which day of the week
the commit was done. To do this, look for lines that start with “From”. Then,
look for the third word, and keep a running count of each of the days of the week.
At the end of the program, print out the contents of the dictionary <code class="docutils literal notranslate"><span class="pre">mail_count</span></code>
(order does not matter). For example, <code class="docutils literal notranslate"><span class="pre">mail_count['Mon']</span></code> should be <code class="docutils literal notranslate"><span class="pre">2</span></code>.</p>

</div>
<textarea data-lang="python" id="dict_writecode1a_editor" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
# mail list was given
mail = ['From stephen.marquard@uct.ac.za Sat Jan  7', 'From gopal.ramasammycook@gmail.com Thurs Jan  5', 'From louis@media.berkeley.edu Tues Jan  3', 'From antranig@caret.cam.ac.uk Sat Jan  7', 'From david.horwitz@uct.ac.za Wed Jan  4', 'From ray@media.berkeley.edu Mon Jan  2', 'From stephen.marquard@uct.ac.za Mon Jan 2', 'From wagnermr@iupui.edu Fri Jan  6']

# Create dictionary for the emails
mail_count = {}

# Iterate through each email in the list
for email in mail:
    # Separate pieces of each email setting key to the day (third element)
    key = email.split()[2]
    # Add key to dictionary if not included
    if key not in mail_count.keys():
        mail_count[key] = 0
    # Update key
    mail_count[key] += 1
# Print final count
print(mail_count)
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=dict_writecode2 data-question_label="10.10.3">
<div id=dict_writecode2_question class="ac_question col-md-12">
<p>Write a program to read through a mail log, build the dictionary <code class="docutils literal notranslate"><span class="pre">user_count</span></code> to count how many messages have come from each email address, and print the dictionary. For example, <code class="docutils literal notranslate"><span class="pre">user_count['stephen.marquard&#64;uct.ac.za']</span></code> should be <code class="docutils literal notranslate"><span class="pre">4</span></code>.</p>

</div>
<textarea data-lang="python" id="dict_writecode2_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
mail_log = ['From stephen.marquard@uct.ac.za Sat Jan  7', 'From gopal.ramasammycook@gmail.com Thurs Jan  5', 'From stephen.marquard@uct.ac.za Sat Feb  7', 'From louis@media.berkeley.edu Tues Jan  3', 'From stephen.marquard@uct.ac.za Sat Nov  6', 'From antranig@caret.cam.ac.uk Sat Jan  7', 'From david.horwitz@uct.ac.za Wed Jan  4', 'From ray@media.berkeley.edu Mon Jan  2', 'From stephen.marquard@uct.ac.za Mon Jan 2', 'From wagnermr@iupui.edu Fri Jan  6', 'From gopal.ramasammycook@gmail.com Thurs Dec  5', 'From louis@media.berkeley.edu Tues April  1']

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(len(user_count), 7, "len(user_count)")
        self.assertEqual(user_count['stephen.marquard@uct.ac.za'], 4, "user_count['stephen.marquard@uct.ac.za']")

myTests().main()
</textarea>
</div>
</div>
</li>
<li><div id='dict_writecode3' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=dict_writecode3q data-question_label="10.10.4">
<div id=dict_writecode3q_question class="ac_question col-md-12">
<p>Write a program that creates a dictionary <code class="docutils literal notranslate"><span class="pre">letter_count</span></code> that keeps track
of the amount of times each letter appears in the given phrase. Assign the
number of times “e” appears in the phrase to the variable <code class="docutils literal notranslate"><span class="pre">e_counter</span></code>. Make
sure to account for each letter in its lowercase form. For example, <code class="docutils literal notranslate"><span class="pre">e_counter</span></code> should be <code class="docutils literal notranslate"><span class="pre">10</span></code>, and <code class="docutils literal notranslate"><span class="pre">letter_count['e']</span></code> should be <code class="docutils literal notranslate"><span class="pre">10</span></code>.</p>

</div>
<textarea data-lang="python" id="dict_writecode3q_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
phrase = "Exeggcute evolves into Exeggutor which are two extraordinary Pokemon"

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(len(letter_count), 21, "len(letter_count)")
        self.assertEqual(e_counter, 10, "e_counter")
        self.assertEqual(letter_count['e'], 10, "letter_count['e']")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=dict_writecode3a data-question_label="10.10.5">
<div id=dict_writecode3a_question class="ac_question col-md-12">
<p>Write a program that creates a dictionary <code class="docutils literal notranslate"><span class="pre">letter_count</span></code> that keeps track
of the amount of times each letter appears in the given phrase. Assign the
number of times “e” appears in the phrase to the variable <code class="docutils literal notranslate"><span class="pre">e_counter</span></code>. Make
sure to account for each letter in its lowercase form. For example, <code class="docutils literal notranslate"><span class="pre">e_counter</span></code> should be <code class="docutils literal notranslate"><span class="pre">10</span></code>, and <code class="docutils literal notranslate"><span class="pre">letter_count['e']</span></code> should be <code class="docutils literal notranslate"><span class="pre">10</span></code>.</p>

</div>
<textarea data-lang="python" id="dict_writecode3a_editor" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
# phrase string is given
phrase = "Exeggcute evolves into Exeggutor which are two extraordinary Pokemon"

# Create dictionary for letters and their counts
letter_count = {}

# Iterate through words in string
for word in phrase.split():
    # Iterate through each letter
    for letter in word:
        # Put each letter into lowercase
        letter = letter.lower()
        # Add letter as key to dictionary, if not included
        if letter not in letter_count.keys():
            letter_count[letter] = 0
        # Update letter
        letter_count[letter] += 1
# Create e_counter to see how many are in the phrase
e_counter = letter_count['e']
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=dict_writecode4 data-question_label="10.10.6">
<div id=dict_writecode4_question class="ac_question col-md-12">
<p>Write a program that reads the words in the string <code class="docutils literal notranslate"><span class="pre">phrase</span></code> and counts how many times each word appears.
Store the words as keys in the dictionary <code class="docutils literal notranslate"><span class="pre">word_dictionary</span></code>, and then use the in operator as a fast way to
check whether the string is in the dictionary. For example, <code class="docutils literal notranslate"><span class="pre">word_dictionary['Writing']</span></code> should be <code class="docutils literal notranslate"><span class="pre">1</span></code>. (Note: ‘Writing’ and ‘writing’ would be counted as two separate words for this question.)</p>

</div>
<textarea data-lang="python" id="dict_writecode4_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
phrase = "Writing programs or programming is a very creative and rewarding activity  You can write programs for many reasons ranging from making your living to solving a difficult data analysis problem to having fun to helping someone else solve a problem  This book assumes that {\em everyone} needs to know how to program and that once you know how to program, you will figure out what you want to do with your newfound skills"

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(len(word_dictionary), 56, "len(word_dictionary)")
        self.assertEqual(word_dictionary['Writing'], 1, "word_dictionary['Writing']")
        self.assertEqual(word_dictionary['a'], 3, "word_dictionary['a']")

myTests().main()
</textarea>
</div>
</div>
</li>
</ol>
<p>Here is the words.txt file for Question 5.</p>

    <div data-component="reveal" id="words-txt-file"   data-showtitle="Show" data-hidetitle="Hide"  style="visibility: hidden;">
    <div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">Writing</span> <span class="n">programs</span> <span class="ow">or</span> <span class="n">programming</span> <span class="ow">is</span> <span class="n">a</span> <span class="n">very</span> <span class="n">creative</span>
<span class="ow">and</span> <span class="n">rewarding</span> <span class="n">activity</span>  <span class="n">You</span> <span class="n">can</span> <span class="n">write</span> <span class="n">programs</span> <span class="k">for</span>
<span class="n">many</span> <span class="n">reasons</span> <span class="n">ranging</span> <span class="kn">from</span> <span class="nn">making</span> <span class="n">your</span> <span class="n">living</span> <span class="n">to</span> <span class="n">solving</span>
<span class="n">a</span> <span class="n">difficult</span> <span class="n">data</span> <span class="n">analysis</span> <span class="n">problem</span> <span class="n">to</span> <span class="n">having</span> <span class="n">fun</span> <span class="n">to</span> <span class="n">helping</span>
<span class="n">someone</span> <span class="k">else</span> <span class="n">solve</span> <span class="n">a</span> <span class="n">problem</span>  <span class="n">This</span> <span class="n">book</span> <span class="n">assumes</span> <span class="n">that</span>
<span class="p">{</span>\<span class="n">em</span> <span class="n">everyone</span><span class="p">}</span> <span class="n">needs</span> <span class="n">to</span> <span class="n">know</span> <span class="n">how</span> <span class="n">to</span> <span class="n">program</span> <span class="ow">and</span> <span class="n">that</span> <span class="n">once</span>
<span class="n">you</span> <span class="n">know</span> <span class="n">how</span> <span class="n">to</span> <span class="n">program</span><span class="p">,</span> <span class="n">you</span> <span class="n">will</span> <span class="n">figure</span> <span class="n">out</span> <span class="n">what</span> <span class="n">you</span> <span class="n">want</span>
<span class="n">to</span> <span class="n">do</span> <span class="k">with</span> <span class="n">your</span> <span class="n">newfound</span> <span class="n">skills</span>

<span class="n">We</span> <span class="n">are</span> <span class="n">surrounded</span> <span class="ow">in</span> <span class="n">our</span> <span class="n">daily</span> <span class="n">lives</span> <span class="k">with</span> <span class="n">computers</span> <span class="n">ranging</span>
<span class="kn">from</span> <span class="nn">laptops</span> <span class="n">to</span> <span class="n">cell</span> <span class="n">phones</span>  <span class="n">We</span> <span class="n">can</span> <span class="n">think</span> <span class="n">of</span> <span class="n">these</span> <span class="n">computers</span>
<span class="k">as</span> <span class="n">our</span> <span class="n">personal</span> <span class="n">assistants</span> <span class="n">who</span> <span class="n">can</span> <span class="n">take</span> <span class="n">care</span> <span class="n">of</span> <span class="n">many</span> <span class="n">things</span>
<span class="n">on</span> <span class="n">our</span> <span class="n">behalf</span>  <span class="n">The</span> <span class="n">hardware</span> <span class="ow">in</span> <span class="n">our</span> <span class="n">current</span><span class="o">-</span><span class="n">day</span> <span class="n">computers</span>
<span class="ow">is</span> <span class="n">essentially</span> <span class="n">built</span> <span class="n">to</span> <span class="n">continuously</span> <span class="n">ask</span> <span class="n">us</span> <span class="n">the</span> <span class="n">question</span>
<span class="n">What</span> <span class="n">would</span> <span class="n">you</span> <span class="n">like</span> <span class="n">me</span> <span class="n">to</span> <span class="n">do</span> <span class="nb">next</span>

<span class="n">Our</span> <span class="n">computers</span> <span class="n">are</span> <span class="n">fast</span> <span class="ow">and</span> <span class="n">have</span> <span class="n">vasts</span> <span class="n">amounts</span> <span class="n">of</span> <span class="n">memory</span> <span class="ow">and</span>
<span class="n">could</span> <span class="n">be</span> <span class="n">very</span> <span class="n">helpful</span> <span class="n">to</span> <span class="n">us</span> <span class="k">if</span> <span class="n">we</span> <span class="n">only</span> <span class="n">knew</span> <span class="n">the</span> <span class="n">language</span> <span class="n">to</span>
<span class="n">speak</span> <span class="n">to</span> <span class="n">explain</span> <span class="n">to</span> <span class="n">the</span> <span class="n">computer</span> <span class="n">what</span> <span class="n">we</span> <span class="n">would</span> <span class="n">like</span> <span class="n">it</span> <span class="n">to</span>
<span class="n">do</span> <span class="nb">next</span> <span class="n">If</span> <span class="n">we</span> <span class="n">knew</span> <span class="n">this</span> <span class="n">language</span> <span class="n">we</span> <span class="n">could</span> <span class="n">tell</span> <span class="n">the</span>
<span class="n">computer</span> <span class="n">to</span> <span class="n">do</span> <span class="n">tasks</span> <span class="n">on</span> <span class="n">our</span> <span class="n">behalf</span> <span class="n">that</span> <span class="n">were</span> <span class="n">reptitive</span>
<span class="n">Interestingly</span><span class="p">,</span> <span class="n">the</span> <span class="n">kinds</span> <span class="n">of</span> <span class="n">things</span> <span class="n">computers</span> <span class="n">can</span> <span class="n">do</span> <span class="n">best</span>
<span class="n">are</span> <span class="n">often</span> <span class="n">the</span> <span class="n">kinds</span> <span class="n">of</span> <span class="n">things</span> <span class="n">that</span> <span class="n">we</span> <span class="n">humans</span> <span class="n">find</span> <span class="n">boring</span>
<span class="ow">and</span> <span class="n">mind</span><span class="o">-</span><span class="n">numbing</span>
</pre></div>
</div>

    </div>
    <ol class="arabic">
<li><div id='dict_writecode5' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone datafile">
<div class="datafile_caption">Data file: <code>words.txt</code></div>
<pre data-component="datafile" id=words.txt data-hidden data-edit="false" data-rows="20" data-cols="60">
Writing programs or programming is a very creative
and rewarding activity  You can write programs for
many reasons ranging from making your living to solving
a difficult data analysis problem to having fun to helping
someone else solve a problem  This book assumes that
{\em everyone} needs to know how to program and that once
you know how to program, you will figure out what you want
to do with your newfound skills

We are surrounded in our daily lives with computers ranging
from laptops to cell phones  We can think of these computers
as our personal assistants who can take care of many things
on our behalf  The hardware in our current-day computers
is essentially built to continuously ask us the question
What would you like me to do next

Our computers are fast and have vasts amounts of memory and
could be very helpful to us if we only knew the language to
speak to explain to the computer what we would like it to
do next If we knew this language we could tell the
computer to do tasks on our behalf that were reptitive
Interestingly, the kinds of things computers can do best
are often the kinds of things that we humans find boring
and mind-numbing
</pre></div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=dict_writecode5q data-question_label="10.10.8">
<div id=dict_writecode5q_question class="ac_question col-md-12">
<p>Write code that reads in the text from the file words.txt and uses
the dictionary <code class="docutils literal notranslate"><span class="pre">word_count</span></code> to count the amount of times a word appears
in the file. Watch out for repetition using the .lower() function. For example, <code class="docutils literal notranslate"><span class="pre">word_count['and']</span></code> should be <code class="docutils literal notranslate"><span class="pre">5</span></code>.</p>

</div>
<textarea data-lang="python" id="dict_writecode5q_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(len(word_count), 119, "len(word_count)")
        self.assertEqual(word_count['and'], 5, "word_count['and']")
        self.assertEqual(word_count['what'], 3, "word_count['what']")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=dict_writecode5a data-question_label="10.10.9">
<div id=dict_writecode5a_question class="ac_question col-md-12">
<p>Write code that reads in the text from the file words.txt and uses
the dictionary <code class="docutils literal notranslate"><span class="pre">word_count</span></code> to count the amount of times a word appears
in the file. Watch out for repetition using the .lower() function. For example, <code class="docutils literal notranslate"><span class="pre">word_count['and']</span></code> should be <code class="docutils literal notranslate"><span class="pre">5</span></code>.</p>

</div>
<textarea data-lang="python" id="dict_writecode5a_editor" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
# Open the file in "read" mode
with open("words.txt", "r") as filename:
    # Create dictionary to count words
    word_count = {}
    # Separate lines and iterate through them
    lines = filename.readlines()
    for line in lines:
        # Iterate through each word
        for word in line.split():
            # Set each word to lowercase
            word = word.lower()
            # Add word to dictionary if not included
            if word not in word_count.keys():
                word_count[word] = 0
            # Update word
            word_count[word] += 1
# Print final count
print(word_count)
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=dict_writecode6 data-question_label="10.10.10">
<div id=dict_writecode6_question class="ac_question col-md-12">
<p>Write a program that reads the words in the string <code class="docutils literal notranslate"><span class="pre">phrase</span></code> and counts how many times each word appears.
Store the words as keys in the dictionary <code class="docutils literal notranslate"><span class="pre">word_dictionary</span></code>, and then use the in operator as a fast way to
check whether the string is in the dictionary. Make sure to turn all letters in words into lowercase letters in order to avoid any repetition.
For example, <code class="docutils literal notranslate"><span class="pre">word_dictionary.get('Writing',</span> <span class="pre">0)</span></code> should be <code class="docutils literal notranslate"><span class="pre">0</span></code>, and <code class="docutils literal notranslate"><span class="pre">word_dictionary['writing']</span></code> should be <code class="docutils literal notranslate"><span class="pre">1</span></code>.</p>

</div>
<textarea data-lang="python" id="dict_writecode6_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
phrase = "Writing programs or programming is a very creative and rewarding activity  You can write programs for many reasons ranging from making your living to solving a difficult data analysis problem to having fun to helping someone else solve a problem  This book assumes that {\em everyone} needs to know how to program and that once you know how to program, you will figure out what you want to do with your newfound skills"

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(len(word_dictionary), 55, "len(word_dictionary)")
        self.assertEqual(word_dictionary['you'], 4, "word_dictionary['you']")
        self.assertEqual(word_dictionary['writing'], 1, "word_dictionary['writing']")
        self.assertEqual(word_dictionary.get('Writing', 0), 0, "word_dictionary.get('Writing', 0)")

myTests().main()
</textarea>
</div>
</div>
</li>
</ol>
<p>Below is the romeo3.txt file used in Question 7.</p>

    <div data-component="reveal" id="romeo3-txt-file"   data-showtitle="Show" data-hidetitle="Hide"  style="visibility: hidden;">
    <div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">But</span> <span class="n">soft</span> <span class="n">what</span> <span class="n">light</span> <span class="n">through</span> <span class="n">yonder</span> <span class="n">window</span> <span class="n">breaks</span>
<span class="n">It</span> <span class="ow">is</span> <span class="n">the</span> <span class="n">east</span> <span class="ow">and</span> <span class="n">Juliet</span> <span class="ow">is</span> <span class="n">the</span> <span class="n">sun</span>
<span class="n">Arise</span> <span class="n">fair</span> <span class="n">sun</span> <span class="ow">and</span> <span class="n">kill</span> <span class="n">the</span> <span class="n">envious</span> <span class="n">moon</span>
<span class="n">Who</span> <span class="ow">is</span> <span class="n">already</span> <span class="n">sick</span> <span class="ow">and</span> <span class="n">pale</span> <span class="k">with</span> <span class="n">grief</span>
</pre></div>
</div>

    </div>
    <ol class="arabic">
<li><div id='dict_writecode7' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone datafile">
<div class="datafile_caption">Data file: <code>romeo3.txt</code></div>
<pre data-component="datafile" id=romeo3.txt data-hidden data-edit="false" data-rows="20" data-cols="48">
But soft what light through yonder window breaks
It is the east and Juliet is the sun
Arise fair sun and kill the envious moon
Who is already sick and pale with grief
</pre></div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=dict_writecode7q data-question_label="10.10.12">
<div id=dict_writecode7q_question class="ac_question col-md-12">
<p>Write code to read through the lines of the file, break each line into a list of
words, and then loop through each of the words in the line and count each word using
the dictionary <code class="docutils literal notranslate"><span class="pre">counts</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">counts['is']</span></code> should be <code class="docutils literal notranslate"><span class="pre">3</span></code>.</p>

</div>
<textarea data-lang="python" id="dict_writecode7q_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(len(counts), 26, "len(counts)")
        self.assertEqual(counts['is'], 3, "counts['is']")
        self.assertEqual(counts['arise'], 1, "counts['arise']")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=dict_writecode7a data-question_label="10.10.13">
<div id=dict_writecode7a_question class="ac_question col-md-12">
<p>Write code to read through the lines of the file, break each line into a list of
words, and then loop through each of the words in the line and count each word using
the dictionary <code class="docutils literal notranslate"><span class="pre">counts</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">counts['is']</span></code> should be <code class="docutils literal notranslate"><span class="pre">3</span></code>.</p>

</div>
<textarea data-lang="python" id="dict_writecode7a_editor" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
# Open file in "read" mode
with open("romeo3.txt", "r") as filename:
    # Separate and read lines
    lines = filename.readlines()
    # Create count dictionary
    counts = {}
    # Iterate through lines
    for line in lines:
        # Iterate through words in each line
        for word in line.split():
            # Set each word to lowercase
            word = word.lower()
            # Add word to dictionary if not included
            if word not in counts.keys():
                counts[word] = 0
            # Update word
            counts[word] += 1
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=dict_writecode8 data-question_label="10.10.14">
<div id=dict_writecode8_question class="ac_question col-md-12">
<p>Write code that adds the key ‘two’ with a value of ‘dos’ to the dictionary <code class="docutils literal notranslate"><span class="pre">eng2sp</span></code>. For example, <code class="docutils literal notranslate"><span class="pre">eng2sp['two']</span></code> should be <code class="docutils literal notranslate"><span class="pre">'dos'</span></code>.</p>

</div>
<textarea data-lang="python" id="dict_writecode8_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
eng2sp = {'one':'uno'}

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(eng2sp['two'], 'dos', "eng2sp['two']")

myTests().main()
</textarea>
</div>
</div>
</li>
</ol>
<p>The next two questions are associated with the following text file which has an email address followed by the number of messages from that address.</p>

    <div data-component="reveal" id="mbox-short-txt-file"   data-showtitle="Show" data-hidetitle="Hide"  style="visibility: hidden;">
    <div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">gopal</span><span class="o">.</span><span class="n">ramasammycook</span><span class="nd">@gmail</span><span class="o">.</span><span class="n">com</span> <span class="mi">1</span>
<span class="n">louis</span><span class="nd">@media</span><span class="o">.</span><span class="n">berkeley</span><span class="o">.</span><span class="n">edu</span> <span class="mi">3</span>
<span class="n">cwen</span><span class="nd">@iupui</span><span class="o">.</span><span class="n">edu</span> <span class="mi">5</span>
<span class="n">antranig</span><span class="nd">@caret</span><span class="o">.</span><span class="n">cam</span><span class="o">.</span><span class="n">ac</span><span class="o">.</span><span class="n">uk</span> <span class="mi">1</span>
<span class="n">rjlowe</span><span class="nd">@iupui</span><span class="o">.</span><span class="n">edu</span> <span class="mi">2</span>
<span class="n">gsilver</span><span class="nd">@umich</span><span class="o">.</span><span class="n">edu</span> <span class="mi">3</span>
<span class="n">david</span><span class="o">.</span><span class="n">horwitz</span><span class="nd">@uct</span><span class="o">.</span><span class="n">ac</span><span class="o">.</span><span class="n">za</span> <span class="mi">4</span>
<span class="n">wagnermr</span><span class="nd">@iupui</span><span class="o">.</span><span class="n">edu</span> <span class="mi">1</span>
<span class="n">zqian</span><span class="nd">@umich</span><span class="o">.</span><span class="n">edu</span> <span class="mi">4</span>
<span class="n">stephen</span><span class="o">.</span><span class="n">marquard</span><span class="nd">@uct</span><span class="o">.</span><span class="n">ac</span><span class="o">.</span><span class="n">za</span> <span class="mi">2</span>
<span class="n">ray</span><span class="nd">@media</span><span class="o">.</span><span class="n">berkeley</span><span class="o">.</span><span class="n">edu</span> <span class="mi">1</span>
</pre></div>
</div>

    </div>
    <ol class="arabic">
<li><div id='dict_writecode9' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone datafile">
<div class="datafile_caption">Data file: <code>mbox-short.txt3</code></div>
<pre data-component="datafile" id=mbox-short.txt3 data-hidden data-edit="false" data-rows="20" data-cols="31">
gopal.ramasammycook@gmail.com 1
louis@media.berkeley.edu 3
cwen@iupui.edu 5
antranig@caret.cam.ac.uk 1
rjlowe@iupui.edu 2
gsilver@umich.edu 3
david.horwitz@uct.ac.za 4
wagnermr@iupui.edu 1
zqian@umich.edu 4
stephen.marquard@uct.ac.za 2
ray@media.berkeley.edu 1
</pre></div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=dict_writecode9q data-question_label="10.10.16">
<div id=dict_writecode9q_question class="ac_question col-md-12">
<p>Add code to the program below to figure out who has the most messages in the file. After all the data has been read and the dictionary <code class="docutils literal notranslate"><span class="pre">message_count</span></code> has been created, look through the dictionary using a maximum loop (see Chapter 5: Maximum and minimum loops) to find who has the most messages,
and print how many messages the person has. For example, <code class="docutils literal notranslate"><span class="pre">message_count['cwen&#64;iupui.edu']</span></code> should be <code class="docutils literal notranslate"><span class="pre">5</span></code>.</p>

</div>
<textarea data-lang="python" id="dict_writecode9q_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
with open("mbox-short.txt3", "r") as filename:
    message_count = {}

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(message_count['cwen@iupui.edu'], '5', "message_count['cwen@iupui.edu']")
        self.assertEqual(len(message_count), 11, "len(message_count)")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=dict_writecode9a data-question_label="10.10.17">
<div id=dict_writecode9a_question class="ac_question col-md-12">
<p>Add code to the program below to figure out who has the most messages in the file. After all the data has been read and the dictionary <code class="docutils literal notranslate"><span class="pre">message_count</span></code> has been created, look through the dictionary using a maximum loop (see Chapter 5: Maximum and minimum loops) to find who has the most messages,
and print how many messages the person has. For example, <code class="docutils literal notranslate"><span class="pre">message_count['cwen&#64;iupui.edu']</span></code> should be <code class="docutils literal notranslate"><span class="pre">5</span></code>.</p>

</div>
<textarea data-lang="python" id="dict_writecode9a_editor" 
      data-timelimit=25000  data-codelens="true"   data-optional=true
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
# Open file in read mode
with open("mbox-short.txt3", "r") as filename:
    # Create message_count dictionary
    message_count = {}
    # Create variable for lines of the file
    messages = filename.readlines()
    # Iterate through each message (each line)
    for message in messages:
        # Assign the key to the first (0th) element of the message
        key = message.split()[0]
        # Assign the value to the second element of the message
        value = message.split()[1]
        # Check if key is already in dictionary
        if key not in message_count.keys():
            # if not, add key/value pair to dictionary
            message_count[key] = value
# Create variable to count emails
max_emails = 0
# Iterate through keys in dictionary
for key in message_count.keys():
    # Check if key is larger than the max emails
    if int(message_count[key]) >= max_emails:
        # If so, reassign max_emails to that key
        max_emails = int(message_count[key])
print(max_emails)
</textarea>
</div>
</div>
</div>
    </div>

</li>
</ol>

<div class="runestone datafile">
<div class="datafile_caption">Data file: <code>mbox-short.txt2</code></div>
<pre data-component="datafile" id=mbox-short.txt2 data-hidden data-edit="false" data-rows="20" data-cols="31">
gopal.ramasammycook@gmail.com 1
louis@media.berkeley.edu 3
cwen@iupui.edu 5
antranig@caret.cam.ac.uk 1
rjlowe@iupui.edu 2
gsilver@umich.edu 3
david.horwitz@uct.ac.za 4
wagnermr@iupui.edu 1
zqian@umich.edu 4
stephen.marquard@uct.ac.za 2
ray@media.berkeley.edu 1
</pre></div>
<ol class="arabic">
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=dict_writecode10 data-question_label="10.10.19">
<div id=dict_writecode10_question class="ac_question col-md-12">
<p>Write a program to record in the dictionary <code class="docutils literal notranslate"><span class="pre">message_count</span></code> the total number of messages from each domain name (not the whole address, just the part after the &#64; and before the space). At the end of the program, print out the contents of your dictionary.
The domains should be the keys of the dictionary, and the counts of the domains should be the values of the dictionary. For example, <code class="docutils literal notranslate"><span class="pre">message_count['iupui.edu']</span></code> should be <code class="docutils literal notranslate"><span class="pre">8</span></code>.</p>

</div>
<textarea data-lang="python" id="dict_writecode10_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
with open("mbox-short.txt2", "r") as filename:
    message_count = {}

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(message_count['iupui.edu'], '8', "message_count['iupui.edu']")
        self.assertEqual(len(message_count), 6, "len(message_count)")

myTests().main()
</textarea>
</div>
</div>
</li>
</ol>
</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 10.9. Mixed-Up Code Questions' data-toggle="tooltip">
  <a href="mixedupcode.html" aria-label="Previous - 10.9. Mixed-Up Code Questions">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 11. Tuples' data-toggle="tooltip" >
    <a href="../tuples/toctree.html" aria-label="Next - 11. Tuples">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2020 - Based on Python for Everybody.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> .
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>